<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[1. Blatt 1. Teil zuf&#228;llig [{"DOMAINUID":"5F4E202"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Geben Sie zu dem folgenden Paar von Mengen an, welche der vier Relationen gelten.
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\) und \(\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\)
</p>
<p>
</p>
<span style="font-size: 0.9375rem;">
    \( \subset \) [[/lang]][[lang code="en"]]
    <p>
        Specify to the following pair of sets that apply to the four relations.
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\) and \(\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\)
    </p>
    <p>
    </p>
    <span style="font-size: 0.9375rem;">
        \( \subset \) [[/lang]][[input:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans1]][[lang code="de"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( \supset \) [[/lang]][[lang code="en"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( \supset \) [[/lang]][[input:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans2]][[lang code="de"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( \in \) [[/lang]][[lang code="en"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( \in \) [[/lang]][[input:ans3]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans3]][[lang code="de"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( = \) [[/lang]][[lang code="en"]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \( = \) [[/lang]][[input:ans4]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans4]][[lang code="de"]]
    </span>
    <br/>
    <p>
    </p>
    [[/lang]][[lang code="en"]]
</span>
<br/>
<p>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <strong>
        L&#246;sungsweg:
    </strong>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not\subset \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\), weil z.B. \({@a4@} \not\in&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\).
</p>
<p>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \supset&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\),
    <span style="font-size: 0.9375rem;">
        weil \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\).
    </span>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not\in \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\), weil \(\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\) keine Menge \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\) als Element hat.
</p>
\(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not= \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\), weil z.B. \({@a4@} \not\in&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\).
<br/>
<p>
</p>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
    <strong>
        Solution:
    </strong>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not\subset \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\)because e.g. \({@a4@} \not\in&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\).
</p>
<p>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \supset&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\),
    <span style="font-size: 0.9375rem;">
        because \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\).
    </span>
</p>
<p>
    \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not\in \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\)because \(\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\) No set \(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\}\) as an element.
</p>
\(\left\{{@a1@},{@a2@},{@a3@},{@a3@},{@a4@}\right\} \not= \left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\)because e.g. \({@a4@} \not\in&#160;\left\{{@a1@},{@a2@},{@a2@},{@a3@}\right\}\).
<br/>
<p>
</p>
<p>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand_with_step(1,20,1);
a2:rand_with_prohib(-12,9,[a1]);
a3:rand_with_prohib(30,40,[a1,a2]);
a4:rand_with_prohib(120,4000,[a1,a2,a3]);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@a1@},{@a2@},{@a3@},{@a4@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.25</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.25</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.25</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.25</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>