<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[1. Blatt 2. Teil zuf&#228;llig [{"DOMAINUID":"5495BEC"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Geben Sie zu dem folgenden Paar von Mengen an, welche der vier Relationen gelten
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\}\) und \(\left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)
    </p>
    <p>
    </p>
    <span style="font-size: 0.9375rem;">
        \(\subset\) [[input:ans1]][[validation:ans1]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(\supset\) [[input:ans2]][[validation:ans2]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(\in\) [[input:ans3]][[validation:ans3]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(=\) [[input:ans4]][[validation:ans4]]
    </span>
    <br/>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Specify to the following pair of quantities that apply to the four relations
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\}\) and \(\left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)
    </p>
    <p>
    </p>
    <span style="font-size: 0.9375rem;">
        \(\subset\) [[input:ans1]][[validation:ans1]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(\supset\) [[input:ans2]][[validation:ans2]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(\in\) [[input:ans3]][[validation:ans3]]
    </span>
    <br/>
    <span style="font-size: 0.9375rem;">
        \(=\) [[input:ans4]][[validation:ans4]]
    </span>
    <br/>
    <p>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <strong>
            L&#246;sungsweg:
        </strong>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\subset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\),&#160;weil jedes Element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\).
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\supset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\),&#160;weil jedes Element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@}\right\}\).
        <br/>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\), weil \(\left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\) keine Menge \(\left\{{@a1@},{@a2@},{@a3@}\right\}\) als Element hat.
        <br/>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \subset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\),&#160;weil jedes Element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\) und&#160;\({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@}\right\}\).
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <strong>
            Solution:
        </strong>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\subset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)because each element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\).
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\supset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)because each element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@}\right\}\).
        <br/>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \not\in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)because \(\left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\) No quantity \(\left\{{@a1@},{@a2@},{@a3@}\right\}\) as an element.
        <br/>
    </p>
    <p>
        \(\left\{{@a1@},{@a2@},{@a3@}\right\} \subset \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\)because each element \({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a3@},{@a2@},{@a2@}\right\}\) and\({@a1@},&#160;{@a2@},&#160;{@a3@} \in \left\{{@a1@},{@a2@},{@a3@}\right\}\).
        <br/>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>a1:rand_with_step(-70,1000,1);
a2:rand_with_prohib(120,400,[a1]);
a3:rand_with_prohib(-140,-20,[a1,a2]);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    [[feedback:prt1]]
</span>
<span class="multilang" lang="en">
    [[feedback:prt1]]
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@a1@},{@a2@},{@a3@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>