<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[2.07 Umkehrfunktion einer gebrochenrationalen Funktion [{"DOMAINUID":"FAE87CE"}] [{"DOMAINUID":"FAE87CE"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Gegeben sei die Funktion
    <span style="font-size: 0.9375rem;">
        \(f(x)=\dfrac{{@z@}}{{@n@}}\) mit dem Parameter \(a \in \mathbb{R}\).
    </span>
</p>
<p>
    Bilden Sie die Umkehrfunktion von \(f\) und bestimmen Sie anschlie&#223;end \(a\) so, dass beide Funktionen \(f\) und \(f^{-1}\) durch den Koordinatenursprung verlaufen.
</p>
<p>
    \(f^{-1}(x)\,=\,\) [[/lang]][[lang code="en"]]
    <p>
        the function
        <span style="font-size: 0.9375rem;">
            \(f(x)=\dfrac{{@z@}}{{@n@}}\) with the parameter \(a \in \mathbb{R}\).
        </span>
    </p>
    <p>
        Create the reverse function of \(f\) and then determine \(a\) so that both functions \(f\) and \(f^{-1}\) run through the coordinate origin.
    </p>
    <p>
        \(f^{-1}(x)\,=\,\) [[/lang]][[input:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        \(a\,=\,\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        \(a\,=\,\) [[/lang]][[input:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p dir="ltr" style="text-align: left;">
    <strong>
        L&#246;sungsweg:
    </strong>
</p>
<p dir="ltr">
    Wir setzen \(y=\dfrac{{@z@}}{{@n@}}\) und versuchen diese Gleichung nach \(x\) umzustellen.
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@y*(f*x+c*a*x-a*d+e)=a1*x-a*x+a*b+c@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y-a*d*y+e*y=a1*x-a*x+a*b+c@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y=a1*x-a*x+a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y-a1*x+a*x=a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@x*(f*y+c*a*y-a1+a)=a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad \displaystyle {@x=(a*b+c+a*d*y-e*y)/(f*y+c*a*y-a1+a)@}\)
</p>
<p dir="ltr">
    Somit ist die Umkehrfunktion&#160;\(\displaystyle f^{-1}(x) = {@loes1@}\). (Ersetzung von \(y\) durch \(x\) und umgekehrt.)
</p>
<p dir="ltr">
    Ferner&#160;verlaufen die beiden Funktionen \(f\) und \(f^{-1}\) durch den Koordinatenursprung, wenn es gilt:
    <span style="font-size: 0.9375rem;">
        \(f(0) = f^{-1}(0) = 0\).
    </span>
</p>
<p dir="ltr">
</p>
<p dir="ltr">
    \(\displaystyle f(0) = {@(a*b+c)/(-a*d+e)@} = 0\) und&#160;\(\displaystyle f^{-1}(0) = {@(a*b+c)/(-a1+a)@} = 0\).
</p>
<p dir="ltr">
    Somit gilt: \(\displaystyle {@a*b+c=0@}\) mit \({@-a*d+e@}\not=0\) und \({@-a1+a@} \not=0\).
</p>
<p dir="ltr">
    Also ist&#160;\(\displaystyle a=&#160;{@-c/b@}\) mit \(\displaystyle a \not= {@e/d@}\) und \(a \not= {@a1@}\).
</p>
[[/lang]][[lang code="en"]]
<p dir="ltr" style="text-align: left;">
    <strong>
        Solution:
    </strong>
</p>
<p dir="ltr">
    We set \(y=\dfrac{{@z@}}{{@n@}}\) and try this equation \(x\) to change.
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@y*(f*x+c*a*x-a*d+e)=a1*x-a*x+a*b+c@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y-a*d*y+e*y=a1*x-a*x+a*b+c@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y=a1*x-a*x+a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@f*x*y+c*a*x*y-a1*x+a*x=a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad {@x*(f*y+c*a*y-a1+a)=a*b+c+a*d*y-e*y@}\)
</p>
<p dir="ltr">
    \(\quad \Longleftrightarrow \quad \displaystyle {@x=(a*b+c+a*d*y-e*y)/(f*y+c*a*y-a1+a)@}\)
</p>
<p dir="ltr">
    Thus, the inverse function\(\displaystyle f^{-1}(x) = {@loes1@}\). (Replacement of \(y\) by \(x\) and vice versa.)
</p>
<p dir="ltr">
    Furthermore, the two functions run \(f\) and \(f^{-1}\) by the origin of coordinates, if:
    <span style="font-size: 0.9375rem;">
        \(f(0) = f^{-1}(0) = 0\).
    </span>
</p>
<p dir="ltr">
</p>
<p dir="ltr">
    \(\displaystyle f(0) = {@(a*b+c)/(-a*d+e)@} = 0\) and\(\displaystyle f^{-1}(0) = {@(a*b+c)/(-a1+a)@} = 0\).
</p>
<p dir="ltr">
    Therefore this applies: \(\displaystyle {@a*b+c=0@}\) with \({@-a*d+e@}\not=0\) and \({@-a1+a@} \not=0\).
</p>
<p dir="ltr">
    So...\(\displaystyle a=&#160;{@-c/b@}\) with \(\displaystyle a \not= {@e/d@}\) and \(a \not= {@a1@}\).
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>b:rand_with_prohib(-3,3,[0]);
c0:rand_with_prohib(-5,5,[0]);
c:b*c0;
d:rand_with_prohib(-5,5,[0]);
e:rand_with_prohib(-5,5,[0,-c*d/b]);
f:rand_with_prohib(-5,5,[0]);
a1:rand_with_prohib(-5,5,[0,-c/b]);

z:a1*x-a*x+a*b+c;
n:f*x+c*a*x-a*d+e;
 
loes1: (d*a*x+a*b-e*x+c)/(x*f+x*c*a-a1+a);
loes2:-c0;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text><![CDATA[\(f(x)=\dfrac{{@z@}}{{@n@}}\) -> \(f^{-1}(x)={@loes1@}\), \(a={@loes2@}\)]]></text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtig. Gut gemacht.
</p>
[[/lang]][[lang code="en"]]
<p>
    Right. Well done.
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist teilweise richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist nicht richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is not correct.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>loes1</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>2</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>loes2</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>loes1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>loes2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>