<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[2.4: F&#252;r welche p sind die 3 Vektoren linear abh&#228;ngig oder unabh&#228;ngig (R3) [{"DOMAINUID":"3D1E524"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Geben sind die Vektoren
    <br/>
    \(\vec{a_1}=\begin{pmatrix}@a1@\\@a2@ \\@a3@ \end{pmatrix},\quad \vec{a_2}=\begin{pmatrix}@x1@\\@x2@ \\@x3@ \end{pmatrix} \quad\) und \( \quad \vec{a_3}=\begin{pmatrix}@b1@\\@b2@ \\p\end{pmatrix}\)
</p>
<p>
    <br/>
    Geben Sie einen Wert des Parameters \(p \in \mathbb{R}\) ein, so dass \(\vec{a_1},&#160;\vec{a_2}\) und \(\vec{a_3}\) linear abh&#228;ngig sind.
</p>
<p>
    \(p_1=\)[[/lang]][[lang code="en"]]
    <p>
        Enter the vectors
        <br/>
        \(\vec{a_1}=\begin{pmatrix}@a1@\\@a2@ \\@a3@ \end{pmatrix},\quad \vec{a_2}=\begin{pmatrix}@x1@\\@x2@ \\@x3@ \end{pmatrix} \quad\) and \( \quad \vec{a_3}=\begin{pmatrix}@b1@\\@b2@ \\p\end{pmatrix}\)
    </p>
    <p>
        <br/>
        Enter a value of the parameter \(p \in \mathbb{R}\) so that \(\vec{a_1},&#160;\vec{a_2}\) and \(\vec{a_3}\) are linearly dependent.
    </p>
    <p>
        \(p_1=\)[[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
    </p>
    <p>
        Geben Sie noch einen Wert des Parameters \(p \in \mathbb{R}\) ein, so dass \(\vec{a_1},&#160;\vec{a_2}\) und \(\vec{a_3}\) linear unabh&#228;ngig sind.
    </p>
    <p>
        \(p_2=\)[[/lang]][[lang code="en"]]
    </p>
    <p>
    </p>
    <p>
        Enter a value of the parameter \(p \in \mathbb{R}\) so that \(\vec{a_1},&#160;\vec{a_2}\) and \(\vec{a_3}\) are linearly independent.
    </p>
    <p>
        \(p_2=\)[[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text/>
        </stackversion>
        <questionvariables>
            <text>a1: rand_with_prohib(-3,3,[0]);
a2: rand_with_prohib(-3,3,[0]);
a3: rand_with_step(-3,3,1);

b1: rand_with_prohib(-3,3,[0]);
b2: 0;
b3: rand_with_prohib(-3,3,[0]);

va: [a1,a2,a3];
vb: [b1,b2,b3];

M1: matrix([a1,a2,a3],[b1,b2,b3]);
R1: rank(M1);

n: rand_with_prohib(-3,3,[0]);
m: rand_with_prohib(-3,3,[0]);

vx: n*va+m*vb;

x1: vx[1];
x2: vx[2];
x3: vx[3];</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(\vec{a_1}=\begin{pmatrix}@a1@\\@a2@ \\@a3@ \end{pmatrix}, \vec{a_2}=\begin{pmatrix}@x1@\\@x2@ \\@x3@ \end{pmatrix} \) und \( \vec{a_3}=\begin{pmatrix}@b1@\\@b2@ \\p\end{pmatrix}, p1=@b3@\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtige Antwort, gut gemacht!
</p>
[[/lang]][[lang code="en"]]
<p>
    Right answer, well done!
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort ist teilweise korrekt.
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Falsche Antwort.
</p>
[[/lang]][[lang code="en"]]
<p>
    Wrong answer.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>b3</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>2*b3</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>b3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ansb: [b1,b2,ans2];

gl1: 0=c1*va[1]+c2*vx[1]+c3*ansb[1];
gl2: 0=c1*va[2]+c2*vx[2]+c3*ansb[2];
gl3: 0=c1*va[3]+c2*vx[3]+c3*ansb[3];
loesung: solve([gl1,gl2,gl3],[c1,c2,c3]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>b3</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>1.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>