<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[4.1a LGS: 3x3 [{"DOMAINUID":"02CB681"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Bestimmen Sie die L&#246;sung des linearen Gleichungssystems
    <br/>
    <span style="font-size: 0.9375rem;">
        \( \left\{\begin{matrix} {@pm11@} \, {@abs11*x1@} &amp; {@pm12@} &amp; {@abs12*x2@} &amp; {@pm13@} &amp; {@abs13*x3@} &amp; = &amp; {@vektorB[1][1]@} \\
        <br/>
    </span>
    <span style="font-size: 0.9375rem;">
        {@pm21@} \, {@abs21*x1@} &amp; {@pm22@} &amp; {@abs22*x2@} &amp; {@pm23@} &amp; {@abs23*x3@} &amp; = &amp; {@vektorB[2][1]@} \\
        <br/>
        {@pm31@} \, {@abs31*x1@} &amp; {@pm32@} &amp; {@abs32*x2@} &amp; {@pm33@} &amp; {@abs33*x3@} &amp; = &amp; {@vektorB[3][1]@} \\
        <br/>
        \end{matrix} \right.\)
    </span>
</p>
<p>
    <span style="font-size: 0.9375rem;">
        <strong>
            Hinweise:
        </strong>
        <br/>
    </span>
</p>
<ul style="">
    <li style="">
        <span style="font-size: 0.9375rem;">
            Geben Sie die L&#246;sung in der Form eines transponierten Vektors \(\vec{x}^T =\) [\(x_1, x_2, x_3\)] ein.
        </span>
    </li>
    <li style="">
        <span style="font-size: 0.9375rem;">
            Falls keine L&#246;sung existiert, geben Sie als L&#246;sung
        </span>
        <span style="font-size: 0.9375rem;">
            [ ]
        </span>
        <span style="font-size: 0.9375rem;">
            ein.
        </span>
    </li>
</ul>
<p>
    \(\vec{x}^T=\,\) [[/lang]][[lang code="en"]]
    <p>
        Define the solution of the linear equation system
        <br/>
        <span style="font-size: 0.9375rem;">
            \( \left\{\begin{matrix} {@pm11@} \, {@abs11*x1@} &amp; {@pm12@} &amp; {@abs12*x2@} &amp; {@pm13@} &amp; {@abs13*x3@} &amp; = &amp; {@vektorB[1][1]@} \\
            <br/>
        </span>
        <span style="font-size: 0.9375rem;">
            {@pm21@} \, {@abs21*x1@} &amp; {@pm22@} &amp; {@abs22*x2@} &amp; {@pm23@} &amp; {@abs23*x3@} &amp; = &amp; {@vektorB[2][1]@} \\
            <br/>
            {@pm31@} \, {@abs31*x1@} &amp; {@pm32@} &amp; {@abs32*x2@} &amp; {@pm33@} &amp; {@abs33*x3@} &amp; = &amp; {@vektorB[3][1]@} \\
            <br/>
            \end{matrix} \right.\)
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            <strong>
                Notes:
            </strong>
            <br/>
        </span>
    </p>
    <ul style="">
        <li style="">
            <span style="font-size: 0.9375rem;">
                Specify the solution in the form of a transposed vector \(\vec{x}^T =\) [\(x_1, x_2, x_3\)].
            </span>
        </li>
        <li style="">
            <span style="font-size: 0.9375rem;">
                If no solution exists, give as a solution
            </span>
            <span style="font-size: 0.9375rem;">
                [
            </span>
            <span style="font-size: 0.9375rem;">
                one.
            </span>
        </li>
    </ul>
    <p>
        \(\vec{x}^T=\,\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p dir="ltr" style="text-align: left;">
    <strong>
        L&#246;sungsweg:
    </strong>
</p>
<p dir="ltr" style="text-align: left;">
    Das vorgegebene Gleichungssystem l&#228;sst sich als erweiterte Matrix \((\mathbf{A}|\mathbf{b}) = {@erwA@}\) aufschreiben.
</p>
<p dir="ltr" style="text-align: left;">
    Mit Hilfe des Eleminationsverfahren von Gau&#223; wird diese Matrix auf folgende Dreiecksform &#252;berf&#252;hrt:
</p>
<p dir="ltr" style="text-align: left;">
    \((\mathbf{\tilde{A}}|\mathbf{\tilde{b}})={@trA@}\).
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    Somit lassen sich die gesuchten Variablen wie folgt berechnen:
</p>
<p dir="ltr" style="text-align: left;">
    Aus der dritten Gleichung ist \(\qquad {@trA[3,3]@} \cdot x_3 = {@trA[3,4]@} \) \(\qquad \Longleftrightarrow \qquad x_3 = \dfrac{{@trA[3,4]@}}{{@trA[3,3]@}} = {@loes3@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    Aus der zweiten Gleichung mit der Einsetzung von \(x_3\) ist \(\qquad&#160;{@trA[2,2]@} \cdot x_2 + \left({@trA[2,3]@}\right) \cdot \left({@loes3@}\right) = {@trA[2,4]@} \) \(\qquad \Longleftrightarrow \qquad x_2 = \dfrac{{@trA[2,4]@} - \left({@trA[2,3]@}\right) \cdot \left({@loes3@}\right)}{{@trA[2,2]@}} = \dfrac{{@trA[2,4]-trA[2,3]*loes3@}}{{@trA[2,2]@}} = {@loes2@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    Aus der ersten Gleichung mit der Einsetzung von \(x_2\) und \(x_3\) ist \(\qquad&#160;{@trA[1,1]@} \cdot x_1 + \left({@trA[1,2]@}\right) \cdot \left({@loes2@}\right) + \left({@trA[1,3]@}\right) \cdot \left({@loes3@}\right) = {@trA[1,4]@} \) \(\qquad \Longleftrightarrow \qquad x_1 = \dfrac{{@trA[1,4]@} - \left({@trA[1,2]@}\right) \cdot \left({@loes2@}\right) - \left({@trA[1,3]@}\right) \cdot \left({@loes3@}\right)}{{@trA[1,1]@}} = \dfrac{{@trA[1,4]-trA[1,2]*loes2-trA[1,3]*loes3@}}{{@trA[1,1]@}} = {@loes1@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    Somit ergibt sich den L&#246;sungsvektor&#160;\(\vec{x} = {@matrix([loes1],[loes2],[loes3])@}\) und die Antwort \(\vec{x}^T = {@loes@}\).
</p>
[[/lang]][[lang code="en"]]
<p dir="ltr" style="text-align: left;">
    <strong>
        Solution:
    </strong>
</p>
<p dir="ltr" style="text-align: left;">
    The predefined equation system can be used as an expanded matrix \((\mathbf{A}|\mathbf{b}) = {@erwA@}\) write down.
</p>
<p dir="ltr" style="text-align: left;">
    With the aid of the elimination process of Gauss, this matrix is converted to the following triangular form:
</p>
<p dir="ltr" style="text-align: left;">
    \((\mathbf{\tilde{A}}|\mathbf{\tilde{b}})={@trA@}\).
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    Thus, the desired variables can be calculated as follows:
</p>
<p dir="ltr" style="text-align: left;">
    From the third equation \(\qquad {@trA[3,3]@} \cdot x_3 = {@trA[3,4]@} \) \(\qquad \Longleftrightarrow \qquad x_3 = \dfrac{{@trA[3,4]@}}{{@trA[3,3]@}} = {@loes3@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    From the second equation with the establishment of \(x_3\) is \(\qquad&#160;{@trA[2,2]@} \cdot x_2 + \left({@trA[2,3]@}\right) \cdot \left({@loes3@}\right) = {@trA[2,4]@} \) \(\qquad \Longleftrightarrow \qquad x_2 = \dfrac{{@trA[2,4]@} - \left({@trA[2,3]@}\right) \cdot \left({@loes3@}\right)}{{@trA[2,2]@}} = \dfrac{{@trA[2,4]-trA[2,3]*loes3@}}{{@trA[2,2]@}} = {@loes2@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    From the first equation with the establishment of \(x_2\) and \(x_3\) is \(\qquad&#160;{@trA[1,1]@} \cdot x_1 + \left({@trA[1,2]@}\right) \cdot \left({@loes2@}\right) + \left({@trA[1,3]@}\right) \cdot \left({@loes3@}\right) = {@trA[1,4]@} \) \(\qquad \Longleftrightarrow \qquad x_1 = \dfrac{{@trA[1,4]@} - \left({@trA[1,2]@}\right) \cdot \left({@loes2@}\right) - \left({@trA[1,3]@}\right) \cdot \left({@loes3@}\right)}{{@trA[1,1]@}} = \dfrac{{@trA[1,4]-trA[1,2]*loes2-trA[1,3]*loes3@}}{{@trA[1,1]@}} = {@loes1@}\)
    <br/>
</p>
<p dir="ltr" style="text-align: left;">
    This results in the solution vector\(\vec{x} = {@matrix([loes1],[loes2],[loes3])@}\) and the answer \(\vec{x}^T = {@loes@}\).
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a: rand_with_prohib(-4,4,[0]);
b: rand_with_prohib(-4,4,[0]);
c: rand_with_prohib(-4,4,[0]);

d: rand_with_prohib(-4,4,[0]);
e: rand_with_prohib(-4,4,[0, b*d/a]);
f: rand_with_prohib(-4,4,[0]);

g: rand_with_prohib(-4,4,[0]);
h: rand_with_prohib(-4,4,[0,e*g/d]);
i:  rand_with_prohib(-5,5,[0,f*h/e,f*g/d,((a*f-c*d)*h+(c*e-b*f)*g)/(a*e-b*d)]);

A: matrix( [a,b,c],[d,e,f],[g,h,i]);
detA: determinant(A);
rangA: rank(A);

loes1: rand_with_prohib(-4,4,[0]);
loes2: rand_with_step(-4,4,1);
loes3: rand_with_step(-4,4,1);
loesmat: matrix([loes1],[loes2],[loes3]);

vektorB: A.loesmat;

loes: [loes1,loes2,loes3];

pm11: if signum(A[1][1])<0 then "- " else " "; 
abs11: abs(A[1][1]);
pm21: if signum(A[2][1])<0 then "- " else " "; 
abs21: abs(A[2][1]);
pm31: if signum(A[3][1])<0 then "- " else " "; 
abs31: abs(A[3][1]);
pm12: if signum(A[1][2])>0 then "+" else "-"; 
abs12: abs(A[1][2]);
pm13: if signum(A[1][3])>0 then "+" else "-"; 
abs13: abs(A[1][3]);
pm22: if signum(A[2][2])>0 then "+" else "-"; 
abs22: abs(A[2][2]);
pm23: if signum(A[2][3])>0 then "+" else "-"; 
abs23: abs(A[2][3]);
pm32: if signum(A[3][2])>0 then "+" else "-"; 
abs32: abs(A[3][2]);
pm33: if signum(A[3][3])>0 then "+" else "-"; 
abs33: abs(A[3][3]);

erwA: matrix( [a,b,c,vektorB[1][1]],[d,e,f,vektorB[2][1]],[g,h,i,vektorB[3][1]]);
trA: triangularize(erwA);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\({@A@} \cdot \vec{x}={@vektorB@}\)
\(det(A)={@detA@}, Rang(A)={@rangA@}\)
\[L={@loes@}\]</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtige Antwort, gut gemacht!
</p>
[[/lang]][[lang code="en"]]
<p>
    Right answer, well done!
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort ist teilweise korrekt.
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Falsche Antwort.
</p>
[[/lang]][[lang code="en"]]
<p>
    Wrong answer.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>loes</tans>
            <boxsize>50</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ansm: if length(ans1)=3 then matrix([ans1[1]],[ans1[2]],[ans1[3]]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>A.ansm</sans>
                <tans> vektorB</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1892464231</deployedseed>
        <deployedseed>1040885938</deployedseed>
        <deployedseed>1414909774</deployedseed>
        <deployedseed>868861813</deployedseed>
        <deployedseed>1115691451</deployedseed>
        <deployedseed>1471205768</deployedseed>
        <deployedseed>1752924508</deployedseed>
        <deployedseed>920083174</deployedseed>
        <deployedseed>377403241</deployedseed>
        <deployedseed>445774165</deployedseed>
        <deployedseed>950481106</deployedseed>
        <deployedseed>614252968</deployedseed>
        <deployedseed>581490534</deployedseed>
        <deployedseed>1745929992</deployedseed>
        <deployedseed>1916519018</deployedseed>
        <deployedseed>1438784629</deployedseed>
        <deployedseed>1561197374</deployedseed>
        <deployedseed>415254026</deployedseed>
        <deployedseed>662819701</deployedseed>
        <deployedseed>2144253276</deployedseed>
        <deployedseed>1840138966</deployedseed>
        <deployedseed>1555108795</deployedseed>
        <deployedseed>1795377181</deployedseed>
        <deployedseed>792888682</deployedseed>
        <deployedseed>2060769467</deployedseed>
        <deployedseed>357655106</deployedseed>
        <deployedseed>23908224</deployedseed>
        <deployedseed>1197481085</deployedseed>
        <deployedseed>254590376</deployedseed>
        <deployedseed>1075026400</deployedseed>
        <deployedseed>1130244286</deployedseed>
        <deployedseed>801497541</deployedseed>
        <deployedseed>381935837</deployedseed>
        <deployedseed>2134841116</deployedseed>
        <deployedseed>1266961649</deployedseed>
        <deployedseed>413126785</deployedseed>
        <deployedseed>807640007</deployedseed>
        <deployedseed>1523687873</deployedseed>
        <deployedseed>1624098175</deployedseed>
        <deployedseed>1658492070</deployedseed>
        <deployedseed>1265561985</deployedseed>
        <deployedseed>1104780757</deployedseed>
        <deployedseed>402953583</deployedseed>
        <deployedseed>688113409</deployedseed>
        <deployedseed>1423253198</deployedseed>
        <deployedseed>1374845965</deployedseed>
        <deployedseed>1891443890</deployedseed>
        <deployedseed>2756716732</deployedseed>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>