<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[4.4 LGS: Schwimmbecken [{"DOMAINUID":"706A04A"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Ein Schwimmbecken kann durch \(3\) Leitungen gef&#252;llt werden. Um das Schwimmbecken zu f&#252;llen, ben&#246;tigen die beiden ersten Leitungen zusammen @v1@ Minuten. Die erste und die dritte Leitung brauchen zusammen @v2@ Minuten. Die zweite und die dritte Leitung schaffen es zusammen in @v3@ Minuten.
</p>
<p>
    Wie lange braucht jede Leitung um den Schwimmbecken alleine zum F&#252;llen?
</p>
<p>
    1. Leitung braucht: [[/lang]][[lang code="en"]]
    <p>
        A swimming pool can be filled through \(3\) lines. To fill the swimming pool, the first two lines need @v1@ minutes. The first and third lines need @v2@ minutes together. The second and third lines fill it together in @v3@ minutes.
    </p>
    <p>
        How long does every line around the swimming pool need to fill?
    </p>
    <p>
        Line 1 needs: [[/lang]][[input:ans1]][[lang code="de"]]&#160;\(min\) [[/lang]][[lang code="en"]]\(min\) [[/lang]][[validation:ans1]][[lang code="de"]]
    </p>
    <p>
        2. Leitung braucht: [[/lang]][[lang code="en"]]
    </p>
    <p>
        Line 2 needs: [[/lang]][[input:ans2]][[lang code="de"]]&#160;\(min\) [[/lang]][[lang code="en"]]\(min\) [[/lang]][[validation:ans2]][[lang code="de"]]
    </p>
    <p>
        3. Leitung braucht: [[/lang]][[lang code="en"]]
    </p>
    <p>
        Line 3 needs: [[/lang]][[input:ans3]][[lang code="de"]]&#160;\(min\) [[/lang]][[lang code="en"]]\(min\) [[/lang]][[validation:ans3]][[lang code="de"]]
    </p>
    <p>
        [[/lang]][[lang code="en"]]
    </p>
    <p>
        [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        Nach wieviel Minuten ist das Becken gef&#252;llt, wenn alle drei Leitungen zusammen benutzt werden?
    </p>
    <p>
        Nach [[/lang]][[lang code="en"]]
    </p>
    <p>
        After how much minutes is the basin filled when all three lines are used together?
    </p>
    <p>
        After [[/lang]][[input:ans4]][[lang code="de"]]&#160;\(min\) [[/lang]][[lang code="en"]]\(min\) [[/lang]][[validation:ans4]][[lang code="de"]]
    </p>
    <p>
        [[/lang]][[lang code="en"]]
    </p>
    <p>
        [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
        <strong>
            Hinweis:
        </strong>
        Runden Sie die Antworten bei Bedarf auf die ganze Minute.
    </p>
    [[/lang]][[lang code="en"]]
</p>
<p>
    <strong>
        Note:
    </strong>
    Complete the answers for the whole minute as needed.
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text/>
        </stackversion>
        <questionvariables>
            <text>v1: rand([15,30,45,60,75,90]);
v2: 4/3*v1;
v3: 2*v1;

gl1: v1*x+v1*y=M;
gl2: v2*x+v2*z=M;
gl3: v3*y+v3*z=M;

loesung: solve([gl1,gl2,gl3],[x,y,z]);

loes1: M/rhs(loesung[1][1]);
loes2: M/rhs(loesung[1][2]);
loes3: M/rhs(loesung[1][3]);

loes4: M/(rhs(loesung[1][1])+rhs(loesung[1][2])+rhs(loesung[1][3]));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\[{@gl1@}\] \[{@gl2@}\] \[{@gl3@}\] \[L&#246;sung: {@loesung@}\]</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtig!
</p>
[[/lang]][[lang code="en"]]
<p>
    Right!
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist teilweise richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist leider falsch.
</p>
[[/lang]][[lang code="en"]]
<p>
    Unfortunately, the answer is wrong.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>loes1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>loes2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>loes3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>loes4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>{ans1,ans2,ans3}</sans>
                <tans>{loes1,loes2,loes3}</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans4</sans>
                <tans>loes4</tans>
                <testoptions>1</testoptions>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>8937606592</deployedseed>
        <deployedseed>3091188516</deployedseed>
        <deployedseed>1586949795</deployedseed>
        <deployedseed>3094650980</deployedseed>
        <deployedseed>8857929480</deployedseed>
        <deployedseed>8976542548</deployedseed>
        <deployedseed>5098088034</deployedseed>
        <deployedseed>8554085978</deployedseed>
        <deployedseed>4197089601</deployedseed>
        <deployedseed>8782782313</deployedseed>
        <deployedseed>3739597477</deployedseed>
        <deployedseed>7453842458</deployedseed>
        <deployedseed>6202245478</deployedseed>
        <deployedseed>6263335601</deployedseed>
        <deployedseed>5223227952</deployedseed>
        <deployedseed>5577250265</deployedseed>
        <deployedseed>6858865541</deployedseed>
        <deployedseed>6743675385</deployedseed>
    </question>
    

</quiz>