<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[5.1b  LGS 3x3, Eine Nullzeile [{"DOMAINUID":"21EADBF"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Bestimmen Sie alle L&#246;sungen des Gleichungssystems:
</p>
<p>
    \[&#160;\begin{matrix}&#160; @A[1][1]*x1@ &amp; @pm12@@A[1][2]*x2@ &amp;&#160;@pm13@@A[1][3]*x3@ &amp; = &amp; @A[1][4]@ \\&#160; @A[2][1]*x1@ &amp;&#160;@pm22@@A[2][2]*x2@ &amp;&#160;@pm23@@A[2][3]*x3@ &amp; = &amp; @A[2][4]@ \\&#160; @A[3][1]*x1@ &amp;&#160;@pm32@@A[3][2]*x2@ &amp;&#160;@pm33@@A[3][3]*x3@ &amp; = &amp; @A[3][4]@ \\&#160;\end{matrix} \]
</p>
<p>
</p>
<p>
    <strong>
        Hinweise:
    </strong>
</p>
<ul>
    <li>
        Geben Sie die L&#246;sungen in der Form \([x_1, x_2, x_3]\) ein. Bitte die eckigen Klammern nicht vergessen.
    </li>
    <li>
        Wenn es keine L&#246;sung gibt, dann geben Sie bitte [ ] ein.
    </li>
</ul>
<p>
</p>
<p>
    \(L=\)&#160;[[/lang]][[lang code="en"]]
    <p>
        Determine all solutions of the equation system:
    </p>
    <p>
        \[&#160;\begin{matrix}&#160; @A[1][1]*x1@ &amp; @pm12@@A[1][2]*x2@ &amp;&#160;@pm13@@A[1][3]*x3@ &amp; = &amp; @A[1][4]@ \\&#160; @A[2][1]*x1@ &amp;&#160;@pm22@@A[2][2]*x2@ &amp;&#160;@pm23@@A[2][3]*x3@ &amp; = &amp; @A[2][4]@ \\&#160; @A[3][1]*x1@ &amp;&#160;@pm32@@A[3][2]*x2@ &amp;&#160;@pm33@@A[3][3]*x3@ &amp; = &amp; @A[3][4]@ \\&#160;\end{matrix} \]
    </p>
    <p>
    </p>
    <p>
        <strong>
            Notes:
        </strong>
    </p>
    <ul>
        <li>
            Give the solutions in the form \([x_1, x_2, x_3]\) one. Please don't forget the square brackets.
        </li>
        <li>
            If there is no solution, please enter [ ].
        </li>
    </ul>
    <p>
    </p>
    <p>
        \(L=\)[[/lang]][[input:ans1]][[lang code="de"]]&#160;[[/lang]][[lang code="en"]][[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text/>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a11: rand_with_prohib(-5,5,[0]);
a12: rand_with_prohib(-5,5,[0]);
a13: rand_with_prohib(-5,5,[0]);

a21: 0;
a22: rand_with_prohib(-5,5,[0]);
a23: rand_with_prohib(-5,5,[0]);

a31: 0;
a32: 0;
a33: 0;


b1: rand_with_prohib(-5,5,[0]);
b2: rand_with_prohib(-5,5,[0]);
b3: 0;

A1: matrix([a11,a12,a13,b1],[a21,a22,a23,b2],[a31,a32,a33,b3]);

s1: rand_with_prohib(-2,2,[0]);
A2: rowop(A1,3,1,s1);

s2: rand_with_prohib(-2,2,[0]);
A3: rowop(A2,3,2,s2);

s3: rand_with_prohib(-2,2,[0]);
A: rowop(A3,2,1,s3);

matA: submatrix(A,4);
vektorB: col(A,4);

pm12: if A[1][2]>0 then "+" else " ";
pm13: if A[1][3]>0 then "+" else " ";
pm22: if A[2][2]>0 then "+" else " ";
pm23: if A[2][3]>0 then "+" else " ";
pm32: if A[3][2]>0 then "+" else " ";
pm33: if A[3][3]>0 then "+" else " ";

gl1: A[1][1]*x1+A[1][2]*x2+A[1][3]*x3=A[1][4];
gl2: A[2][1]*x1+A[2][2]*x2+A[2][3]*x3=A[2][4];
gl3: A[3][1]*x1+A[3][2]*x2+A[3][3]*x3=A[3][4];
ls: solve([gl1,gl2,gl3],[x1,x2,x3]);

loes: expand( [ rhs(ls[1][1]) , rhs(ls[1][2]), rhs(ls[1][3]) ] );]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\[@A@\]
L&#246;sung: \(@loes@\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtige Antwort, gut gemacht!
</p>
[[/lang]][[lang code="en"]]
<p>
    Right answer, well done!
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort ist teilweise korrekt.
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Falsche Antwort.
</p>
[[/lang]][[lang code="en"]]
<p>
    Wrong answer.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>loes</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ansm: if length(ans1)=3 then matrix([ans1[1]],[ans1[2]],[ans1[3]]);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>matA.ansm</sans>
                <tans>vektorB</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>length(listofvars(ans1))</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort \(@ans1@\) ist eine spezifische L&#246;sung des Gleichungssystems.
</p>
<p>
    Gefragt wird nach der allgemeinen L&#246;sung.
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer \(@ans1@\) is a specific solution of the equation system.
</p>
<p>
    The question is about the general solution.
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>1.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1788287284</deployedseed>
        <deployedseed>576492072</deployedseed>
        <deployedseed>1885153380</deployedseed>
        <deployedseed>408123529</deployedseed>
        <deployedseed>1408428442</deployedseed>
        <deployedseed>1476357967</deployedseed>
        <deployedseed>1268122836</deployedseed>
        <deployedseed>1371991493</deployedseed>
        <deployedseed>655297154</deployedseed>
        <deployedseed>1519211889</deployedseed>
        <deployedseed>186710872</deployedseed>
        <deployedseed>1281589498</deployedseed>
        <deployedseed>525493515</deployedseed>
        <deployedseed>1563792141</deployedseed>
        <deployedseed>465380207</deployedseed>
        <deployedseed>1356128424</deployedseed>
        <deployedseed>801802172</deployedseed>
        <deployedseed>892089591</deployedseed>
        <deployedseed>2078093997</deployedseed>
        <deployedseed>1705384004</deployedseed>
        <deployedseed>81436082</deployedseed>
        <deployedseed>788667439</deployedseed>
        <deployedseed>1236107602</deployedseed>
        <deployedseed>676839239</deployedseed>
        <deployedseed>1510060679</deployedseed>
        <deployedseed>140957626</deployedseed>
        <deployedseed>874775886</deployedseed>
        <deployedseed>621743218</deployedseed>
        <deployedseed>1384631502</deployedseed>
        <deployedseed>1829741886</deployedseed>
        <deployedseed>1394027782</deployedseed>
        <deployedseed>385779974</deployedseed>
        <deployedseed>249073495</deployedseed>
        <deployedseed>446650460</deployedseed>
        <deployedseed>1245254111</deployedseed>
        <deployedseed>977125031</deployedseed>
        <deployedseed>1721048361</deployedseed>
        <deployedseed>462931861</deployedseed>
        <deployedseed>310873125</deployedseed>
        <deployedseed>708647801</deployedseed>
        <deployedseed>370613688</deployedseed>
        <deployedseed>2045424244</deployedseed>
        <deployedseed>678855505</deployedseed>
        <deployedseed>256337017</deployedseed>
        <deployedseed>844753091</deployedseed>
        <deployedseed>860645529</deployedseed>
        <deployedseed>1786446724</deployedseed>
        <deployedseed>371901984</deployedseed>
        <deployedseed>1192076959</deployedseed>
        <deployedseed>761547006</deployedseed>
        <deployedseed>1824188025</deployedseed>
        <deployedseed>83119314</deployedseed>
        <deployedseed>681107344</deployedseed>
        <deployedseed>148097582</deployedseed>
        <deployedseed>1336833604</deployedseed>
        <deployedseed>1822804994</deployedseed>
        <deployedseed>369822542</deployedseed>
        <deployedseed>112177314</deployedseed>
        <deployedseed>1927058873</deployedseed>
        <deployedseed>1899219269</deployedseed>
        <deployedseed>968769161</deployedseed>
        <deployedseed>815074938</deployedseed>
    </question>
    

</quiz>