<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Alle: wurzel, funktionen</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Geben Sie den folgenden Ausdruck in das Eingabefeld ein.
    <br/>
</p>
<p>
    \(\dfrac{\pi^2+\sqrt{e^x-{@a@}}}{\tan({@a1@}x)+\ln(x-{@a2@})}\) &#160; [[/lang]][[lang code="en"]]
    <p>
    </p>
    <p>
        Enter the following expression in the input field.
        <br/>
    </p>
    <p>
        \(\dfrac{\pi^2+\sqrt{e^x-{@a@}}}{\tan({@a1@}x)+\ln(x-{@a2@})}\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]]
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            Hinweise:
            <br/>
        </strong>
    </p>
    <ul>
        <li>
            Die Zahl \(\pi\) wird mit
            <strong>
                pi
            </strong>
            eingegeben
        </li>
        <li>
            Potenz wird mit
            <strong>
                ^
            </strong>
            <span>
                und
            </span>
            <strong>
            </strong>
            Wurzel wird mit
            <strong>
                sqrt
            </strong>
            eingegeben, z.B.
            <strong>
                sqrt(x^2+1)
            </strong>
            f&#252;r \(\sqrt{x^2+1}\)
        </li>
    </ul>
    [[/lang]][[lang code="en"]]
</p>
<p>
    <br/>
</p>
<p>
    <strong>
        Notes:
        <br/>
    </strong>
</p>
<ul>
    <li>
        The number \(\pi\) with
        <strong>
            pi
        </strong>
        Entered
    </li>
    <li>
        potency is
        <strong>
            ^
        </strong>
        <span>
            and
        </span>
        <strong>
        </strong>
        Root becomes with
        <strong>
            sqrt
        </strong>
        input, for example
        <strong>
            sqrt(x^2+1)
        </strong>
        for \(\sqrt{x^2+1}\)
    </li>
</ul>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>a:rand_with_step(1,5,1);
a1: rand([5,7]);
a2: rand([2,4]);
los:(pi^2+sqrt(exp(x)-a))/(tan(a1*x)+ln(x-a2));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>@a@</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Falsche Antwort.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>los</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>los</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>