<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Aufgabe3</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Auf&#160; die Menge \( M= \{a,b,c\} \) ist die Relation \( R= \{&#160; (a,b), (a,c), (a,a), (b,b), (c,c) \} \) gegeben. F&#252;gen Sie m&#246;glichst wenige weitere Paare hinzu, sodass eine totale Ordnung entsteht.
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            Gesuchte Element: [[input:ans1]][[validation:ans1]]
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            <br/>
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            <em>
                <strong>
                    Hinweis zur Eingabe:
                </strong>
                Z.B.&#160; [1,2]&#160; wird als
                <code>
                    cc(1,2)
                </code>
                eingegeben.
            </em>
            <br/>
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        To the amount \( M= \{a,b,c\} \) is the relationship \( R= \{&#160; (a,b), (a,c), (a,a), (b,b), (c,c) \} \) given. Add as few more pairs as possible, so that a total order is created.
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            Selected element: [[input:ans1]][[validation:ans1]]
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            <br/>
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            <em>
                <strong>
                    Note to input:
                </strong>
                For example [1,2],
                <code>
                    cc(1,2)
                </code>
                entered.
            </em>
            <br/>
        </span>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr" style="text-align: left;">
        Antwortm&#246;glichkeiten sind:
    </p>
    <p dir="ltr" style="text-align: left;">
        Der gesuchte Element ist: \( [b,c]\) oder eben &#160;\(&#160;[c,b]\). &#160;Man kann sich oben noch aussuchen, wie \(b\) und \(c\) in Relation stehen.
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr" style="text-align: left;">
        Answer options are:
    </p>
    <p dir="ltr" style="text-align: left;">
        The searched element is: \( [b,c]\) or even\(&#160;[c,b]\). You can still look at the top as \(b\) and \(c\) in relation.
        <br/>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>range9:makelist(i,i,1,9);
liste:random_permutation(range9);
n1:3+rand(3);
X:rest(liste,-length(liste)+n1);
R:cc(b,c);
R1:cc(c,b);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    [[feedback:prt1]]
</span>
<span class="multilang" lang="en">
    [[feedback:prt1]]
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@X@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>R</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>R</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>R1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>