<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Aufz&#228;hlende Form_Vorkurs</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Seien \(A = \{x | x \hbox{ ist eine ungerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als 0 und kleiner als } {@a@}\}\), und \( B = \{x | x \hbox{ ist eine gerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als } {@b@} \hbox{ und kleiner als } {@c@} \}\). Bestimmen Sie die folgenden Mengen \(A\) und \(B\).
    </p>
    <p>
        \(A = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        \(B = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Be \(A = \{x | x \hbox{ ist eine ungerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als 0 und kleiner als } {@a@}\}\)and \( B = \{x | x \hbox{ ist eine gerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als } {@b@} \hbox{ und kleiner als } {@c@} \}\). Determine the following quantities \(A\) and \(B\).
    </p>
    <p>
        \(A = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        \(B = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <strong>
            L&#246;s&#252;ng:
        </strong>
        <br/>
    </p>
    <p>
        \(\begin{align*}A &amp; = \{x | x \hbox{ ist eine ungerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als 0 und kleiner als } {@a@}\}\\ &amp; = {@ta1@}. \end{align*}\)
    </p>
    <p>
        \(\begin{align*}B &amp; = \{x | x \hbox{ ist eine gerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als } {@b@} \hbox{ und kleiner als } {@c@} \}\\ &amp; = {@ta2@}. \end{align*}\)
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <strong>
            Lace:
        </strong>
        <br/>
    </p>
    <p>
        \(\begin{align*}A &amp; = \{x | x \hbox{ ist eine ungerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als 0 und kleiner als } {@a@}\}\\ &amp; = {@ta1@}. \end{align*}\)
    </p>
    <p>
        \(\begin{align*}B &amp; = \{x | x \hbox{ ist eine gerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als } {@b@} \hbox{ und kleiner als } {@c@} \}\\ &amp; = {@ta2@}. \end{align*}\)
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>/*Erstellt von Gozel Judakova, OVGU Magdeburg*/

a:rand(5)+10;
b:rand(5)+15;
c:rand(5)+30;
L1:makelist(i,i,1,a-1);
L2:makelist(i,i,b+1,c-1);
A:setify(L1);
B:setify(L2);
ta1:subset(A,oddp);
ta2:subset(B,evenp);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\[A = \{x | x \hbox{ ist eine ungerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als 0 und kleiner als } {@a@}\}\]
\[ B = \{x | x \hbox{ ist eine gerade ganze Zahl und } x \hbox{ ist gr&#246;&#223;er als } {@b@} \hbox{ und kleiner als } {@c@} \}\]
\( A = {@ta1@}, \quad B={@ta2@} \)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-check">
        </i>
    </span>
    Deine Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-check">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{?, ?, ?, ...}</syntaxhint>
            <syntaxattribute>1</syntaxattribute>
            <forbidwords>subset</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{?, ?, ?, ...}</syntaxhint>
            <syntaxattribute>1</syntaxattribute>
            <forbidwords>subset</forbidwords>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>439398470</deployedseed>
        <deployedseed>1518990853</deployedseed>
        <deployedseed>1663661717</deployedseed>
        <deployedseed>241964910</deployedseed>
        <deployedseed>128602194</deployedseed>
        <deployedseed>290659520</deployedseed>
        <deployedseed>1178953760</deployedseed>
        <deployedseed>541389342</deployedseed>
        <deployedseed>636801768</deployedseed>
        <deployedseed>2084640002</deployedseed>
        <deployedseed>1344897941</deployedseed>
        <deployedseed>1221120020</deployedseed>
        <deployedseed>1294565485</deployedseed>
        <deployedseed>1701133430</deployedseed>
        <deployedseed>439653355</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>M1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>fa1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
        </qtest>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>