<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Aufz&#228;hlende Form angeben [{"DOMAINUID":"BD1CD04"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]Stellen Sie die folgenden Mengen in aufz&#228;hlender Form dar.
<p>
</p>
<p>
    a) \(M_1 = \left\{n^{@n5@}+{@n1@} \ \Big|\&#160; n\in\mathbb N, n &lt; {@n2@} \right\} = \) [[/lang]][[lang code="en"]]Set the following sets in a listing form.
    <p>
    </p>
    <p>
        (a) \(M_1 = \left\{n^{@n5@}+{@n1@} \ \Big|\&#160; n\in\mathbb N, n &lt; {@n2@} \right\} = \) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        b) \(M_2 = \left\{{@n3@} \cdot z \ \Big|\ z \in \mathbb Z, |z|\le {@n4@} \right\} = \) [[/lang]][[lang code="en"]]
    </p>
    <p>
        (b) \(M_2 = \left\{{@n3@} \cdot z \ \Big|\ z \in \mathbb Z, |z|\le {@n4@} \right\} = \) [[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
    </p>
    <p>
        <strong>
            Hinweis:
            <br/>
        </strong>
    </p>
    <ul>
        <li>
            Schreiben Sie die Menge mit geschweiften Klammern. Trennen Sie die L&#246;sungen mit Komma, z.B. {1,3}.
        </li>
        <li>
            Schreiben Sie die leere Menge als { }.
        </li>
    </ul>
    [[/lang]][[lang code="en"]]
</p>
<p>
</p>
<p>
    <strong>
        Note:
        <br/>
    </strong>
</p>
<ul>
    <li>
        Write the set with curly brackets. Separate the solutions with comma, e.g. {1,3}.
    </li>
    <li>
        Write the empty set as { }.
    </li>
</ul>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<strong>
    L&#246;sungsweg:
</strong>
<p>
</p>
<p>
    a) F&#252;r jede nat&#252;rliche Zahl, die \(n &lt; {@n2@}\) ist, wird die Potenz \(n^{@n5@}\) gebildet und die Zahl \({@n1@}\) dazu addiert. So bekommt man die Menge \(M_1 = {@M1@}\).
</p>
<p>
    b) F&#252;r jede ganze Zahl, die \({@-n4@} \le z \le {@n4@}\) ist, wird das \({@n3@}\)-fache dieser Zahl gebildet. So bekommt man die Menge \(M_2 = {@M2@}\).
</p>
[[/lang]][[lang code="en"]]
<strong>
    Solution:
</strong>
<p>
</p>
<p>
    (a) For any natural number \(n &lt; {@n2@}\) is the potency \(n^{@n5@}\) and the number \({@n1@}\) added. How to get the set \(M_1 = {@M1@}\).
</p>
<p>
    (b) For each integer \({@-n4@} \le z \le {@n4@}\) it will \({@n3@}\)- more of this number. How to get the set \(M_2 = {@M2@}\).
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>/*Erstellt von Stephan Bach, OTH Amberg-Weiden*/

n1:rand(4)+1;
n2:rand(2)+3;
n3:rand_with_prohib(2,5,[n1]);
n4:rand_with_prohib(1,3,[n2])
n5:rand([2,3])

M1:makelist(i^n5+n1,i,1,n2-1);
M1:setify(M1);
M2:makelist(n3*i,i,-n4,n4);
M2:setify(M2);

/*m&#246;gliche Falschantwort f&#252;r Fragetests*/
fa1:setify(makelist(i^n5+n1,i,0,n2-1));
fa2:setify(makelist(n3*i,i,0,n4));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text><![CDATA[a) \(M_1=\{n^{@n5@}+{@n1@} \ |\  n\in\mathbb N, n<{@n2@}\}={@M1@}, \hspace{0.5cm}\)
b) \(M_2=\{{@n3@} z\ |\ z \in \mathbb Z, |z| \le {@n4@}\}={@M2@}\)]]></text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<strong>
    Richtige Antwort, gut gemacht!
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<strong>
    Right answer, well done!
</strong>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<strong>
    Ihre Antwort ist teilweise korrekt.
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<strong>
    Your answer is partly correct.
</strong>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<strong>
    Falsche Antwort.
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<strong>
    Wrong answer.
</strong>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>M1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>M2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>M1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>fa1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Beachten Sie, dass die Zahl \(0\) kein Element der nat&#252;rlichen Zahlen ist und daher \(0^{@n5@}+{@n1@}={@0^n5+n1@}\notin M_1\).
</p>
<p>
    Die korrekte Darstellung der Menge ist \(M_1={@M1@}\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    Note that the number \(0\) is not an element of natural numbers and therefore \(0^{@n5@}+{@n1@}={@0^n5+n1@}\notin M_1\).
</p>
<p>
    The correct representation of the set is \(M_1={@M1@}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Die korrekte Darstellung der Menge ist \(M_1={@M1@}\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    The correct representation of the set is \(M_1={@M1@}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>M2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Die korrekte Darstellung der Menge ist \(M_2={@M2@}\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    The correct representation of the set is \(M_2={@M2@}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>439398470</deployedseed>
        <deployedseed>1518990853</deployedseed>
        <deployedseed>1663661717</deployedseed>
        <deployedseed>241964910</deployedseed>
        <deployedseed>128602194</deployedseed>
        <deployedseed>290659520</deployedseed>
        <deployedseed>1178953760</deployedseed>
        <deployedseed>541389342</deployedseed>
        <deployedseed>636801768</deployedseed>
        <deployedseed>2084640002</deployedseed>
        <deployedseed>1344897941</deployedseed>
        <deployedseed>1221120020</deployedseed>
        <deployedseed>1294565485</deployedseed>
        <deployedseed>1701133430</deployedseed>
        <deployedseed>439653355</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>M1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>M2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>fa1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>fa2</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
        </qtest>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>