<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Aussage negieren [{"DOMAINUID":"56D6538"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Negieren Sie die folgenden Aussagen.
    <br/>
</p>
<p>
    a) Aussage \(A\) = "Alle Hunde haben vier Beine"
</p>
<p>
    \(\lnot A=\) [[/lang]][[lang code="en"]]
    <p>
        Negate the following statements.
        <br/>
    </p>
    <p>
        a) Statement \(A\) = "All dogs have four legs"
    </p>
    <p>
        \(\lnot A=\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[input:ans2]][[lang code="de"]]
        <span style="font-size: 0.9375rem;">
            [[/lang]][[lang code="en"]]
            <span style="font-size: 0.9375rem;">
                [[/lang]][[feedback:prt1]][[lang code="de"]]
            </span>
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            b) Aussage \(B\) = "Es gibt mindestens ein Geb&#228;ude auf der Welt, das h&#246;her ist als 500 Meter"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot B=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            b) Statement \(B\) = "There is at least one building in the world that is higher than 500 meters"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot B=\) [[/lang]][[input:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[input:ans4]][[lang code="de"]]
        <span style="font-size: 0.9375rem;">
            [[/lang]][[lang code="en"]]
            <span style="font-size: 0.9375rem;">
                [[/lang]][[feedback:prt2]][[lang code="de"]]
            </span>
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            c) Aussage \(C\) = "Es gibt keinen Mathematiker, der sich noch nie verrechnet hat"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot C=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            c) Statement \(C\) = "There is no mathematician who has never been charged"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot C=\) [[/lang]][[input:ans5]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[input:ans6]][[lang code="de"]]
        <span style="font-size: 0.9375rem;">
            [[/lang]][[lang code="en"]]
            <span style="font-size: 0.9375rem;">
                [[/lang]][[feedback:prt3]][[lang code="de"]]
            </span>
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            d) Aussage \(D\) = "F&#252;r alle Pferdebesitzer gilt, dass alle ihre Pferde wiehern"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot D=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            d) statement \(D\) = "For all horse owners it is true that all their horses are like"
        </span>
        <br/>
    </p>
    <p>
        \(\lnot D=\) [[/lang]][[input:ans7]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[input:ans8]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[input:ans9]][[lang code="de"]]
        <span style="font-size: 0.9375rem;">
            [[/lang]][[lang code="en"]]
            <span style="font-size: 0.9375rem;">
                [[/lang]][[feedback:prt4]][[lang code="de"]]
            </span>
        </span>
    </p>
    <p>
        [[/lang]][[lang code="en"]]
    </p>
    <p>
        [[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans4]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans5]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans6]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans7]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans8]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans9]][[lang code="de"]]
        <br/>
    </p>
    [[/lang]][[lang code="en"]]
    <br/>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        Hinweis zur Negation von Aussagen:
    </b>
</p>
<p>
    Aussagen mit Quantoren (\(\forall\) und \(\exists\)) werden negiert, indem man jeden der \(\forall\) in ein \(\exists\) umgewandelt und umgekehrt. Die inneren Aussagen werden ganz normal negiert, die Aussagen des Quantors bleiben unber&#252;hrt.
    <br/>
</p>
<p>
    <u>
        Beispiel:
    </u>
    \(\lnot (\forall x\in M : x&gt;5 )\Leftrightarrow \exists x\in M: x\le 5\)
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        Note on the negation of statements:
    </b>
</p>
<p>
    Statements with quantums (\(\forall\) and \(\exists\)) are negated by making every \(\forall\) in a \(\exists\) and vice versa. The inner statements are negated quite normally, the statements of the quantifier remain unaffected.
    <br/>
</p>
<p>
    <u>
        Example:
    </u>
    \(\lnot (\forall x\in M : x&gt;5 )\Leftrightarrow \exists x\in M: x\le 5\)
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*Erstellt von Jonas Lache, Ruhr-Universit&#228;t Bochum*/
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
L1_1:
[[1,true,tr("Es gibt mindestens einen Hund,","There is at least one dog,")],
[2,false,tr("Es gibt genau einen Hund,","There is exactly one dog")],
[3,false,tr("F&#252;r alle Hunde gilt, dass","For all dogs applies, that")],
[4,false,tr("Es gibt keinen Hund,","There are no dogs,")]];
L1: random_permutation(L1_1);

L2_1: 
[[1,true,tr("der nicht vier Beine hat","that does not have four legs")],
[2,false,tr("der mehr als vier Beine hat","that has more than four legs")],
[3,false,tr("der weniger als vier Beine hat","that has less than four legs")],
[4,false,tr("der genau vier Beine hat","that has exactly four legs")]];
L2: random_permutation(L2_1);

L3_1: 
[[1,true,tr("F&#252;r jedes Geb&#228;ude auf der Welt gilt,","For all buildings applies,")],
[2,false,tr("Es gibt genau ein Geb&#228;ude auf der Welt, f&#252;r das gilt,","There exists exactly one building in the world for that applies,")],
[3,false,tr("Es gibt mindestens ein Geb&#228;ude auf der Welt, f&#252;r das gilt,","There exists at least one building in the world for that applies,")],
[4,false,tr("Es gibt kein Geb&#228;ude auf der Welt, f&#252;r das gilt,","There is no building in the world for that applies,")]];
L3: random_permutation(L3_1);

L4_1: 
[[1,true,tr("dass es niedriger als 500 Meter oder genau 500 Meter hoch ist","that is lower than or exactly 500m tall")],
[2,false,tr("dass es genau 500 Meter hoch ist","that is exactly 500m tall")],
[3,false,tr("dass es niedriger als 500 Meter ist","that is lower than 500m tall")],
[4,false,tr("dass es h&#246;her als 500 Meter oder genau 500 Meter hoch ist","that is taller than or exactly 500m tall")]];
L4: random_permutation(L4_1);

L5_1: 
[[1,true,tr("Es gibt mindestens einen Mathematiker, f&#252;r den gilt,","There is at least one mathematician for that applies,")],
[2,false,tr("Es gibt genau einen Mathematiker, f&#252;r den gilt, ","There is exactly one mathematician for that applies,")],
[3,false,tr("F&#252;r jeden Mathematiker gilt,","For all mathematicians applies,")],
[4,false,tr("Es gibt keinen Mathematiker, f&#252;r den gilt,","There is no mathematician for that applies,")]];
L5: random_permutation(L5_1);

L6_1: 
[[1,true,tr("dass er sich noch nie verrechnet hat","that they have never miscalculated")],
[2,false,tr("dass er sich jedes Mal verrechnet","that they always miscalculate")],
[3,false,tr("dass er sich schon genau einmal verrechnet hat","that they have miscalculated exactly once before")],
[4,false,tr("dass er sich schon mindestens einmal verrechnet hat","that they have miscalculated at least once before")]];
L6: random_permutation(L6_1);


L7_1:
[[1,true,tr("Es gibt mindestens einen Pferdebesitzer, f&#252;r den gilt,","There exists at least one horse owner for that applies,")],
[2,false,tr("Es gibt genau einen Pferdebesitzer, f&#252;r den gilt, ","There exists exactly one horse owner for that applies,")],
[3,false,tr("F&#252;r jeden Pferdebesitzer gilt,","For all horse owner applies,")],
[4,false,tr("Es gibt keinen Pferdebesitzer, f&#252;r den gilt,","There exists no horse owner for that applies,")]];
L7: random_permutation(L7_1);

L8_1:
[[1,true,tr("dass er mindstens ein Pferd besitzt, f&#252;r das gilt,","that they own at least one horse for that applies,")],
[2,false,tr("dass nur Pferde besitzt, von denen f&#252;r jedes gilt,","that they own only horses for that applies,")],
[3,false,tr("dass er genau ein Pferd besitzt, f&#252;r das gilt,","that they own exactly one horse for that applies,")],
[4,false,tr("dass er kein Pferd besitzt, f&#252;r das gilt","that they own no horse for that applies,")]];

L8: random_permutation(L8_1);
L9_1: [[1,true,tr("dass es nicht wiehert","that it does not neighs")],
[2,false,tr("dass es schon genau einmal gewiehert hat","that it has neighed exactly once before")],
[3,false,tr("dass es schon mindestens einmal gewiehert hat","that it has neighed at least once before")]];
L9: random_permutation(L9_1);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@L1@}, {@L2@} {@L3@}, {@L4@} {@L5@}, {@L6@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Falsche Antwort.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>radio</type>
            <tans>L1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>radio</type>
            <tans>L2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>radio</type>
            <tans>L3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>radio</type>
            <tans>L4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>radio</type>
            <tans>L5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>radio</type>
            <tans>L6</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans7</name>
            <type>radio</type>
            <tans>L7</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans8</name>
            <type>radio</type>
            <tans>L8</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans9</name>
            <type>radio</type>
            <tans>L9</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[SAns1: sconcat((L1_1[ans1])[3]," ",(L2_1[ans2])[3]);]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns1@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es gibt mindestens einen Hund, der nicht vier Beine hat".
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns1@}" is unfortunately not correct.
    <br/>
    "There's at least one dog that doesn't have four legs."
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]Ihre Antwort "{@SAns1@}" ist leider nicht korrekt.
<br/>
Richtig w&#228;re "Es gibt mindestens einen Hund, der nicht vier Beine hat".
<br/>
[[/lang]][[lang code="en"]]Your answer "{@SAns1@}" is unfortunately not correct.
<br/>
"There's at least one dog that doesn't have four legs."
<br/>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[SAns2: sconcat((L3_1[ans3])[3]," ",(L4_1[ans4])[3]);]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]Ihre Antwort "{@SAns2@}" ist leider nicht korrekt.
<br/>
Richtig w&#228;re "F&#252;r jedes Geb&#228;ude auf der Welt gilt, dass es niedriger als 500 Meter oder genau 500 Meter hoch ist"
<br/>
[[/lang]][[lang code="en"]]Your answer "{@SAns2@}" is unfortunately not correct.
<br/>
"For every building in the world, it is lower than 500 meters or exactly 500 meters high"
<br/>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns2@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "F&#252;r jedes Geb&#228;ude auf der Welt gilt, dass es niedriger als 500 Meter oder genau 500 Meter hoch ist"
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns2@}" is unfortunately not correct.
    <br/>
    "For every building in the world, it is lower than 500 meters or exactly 500 meters high"
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[SAns3: sconcat((L5_1[ans5])[3]," ",(L6_1[ans6])[3]);]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns3@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es gibt mindestens einen Mathematiker, f&#252;r den gilt, dass er sich noch nie verrechnet hat"
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns3@}" is unfortunately not correct.
    <br/>
    "There's at least one mathematician to whom he has never been charged"
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns3@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es gibt mindestens einen Mathematiker, f&#252;r den gilt, dass er sich noch nie verrechnet hat"
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns3@}" is unfortunately not correct.
    <br/>
    "There's at least one mathematician to whom he has never been charged"
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[SAns4: sconcat((L7_1[ans7])[3]," ",(L8_1[ans8])[3]," ",(L9_1[ans9])[3]);]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans7</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Ihre Antwort "{@SAns4@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es gibt mindestens einen Pferdebesitzer, f&#252;r den gilt, dass er mindstens ein Pferd besitzt, f&#252;r das gilt, dass es nicht wiehert"
    <br/>
</p>
<br/>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    Your answer "{@SAns4@}" is unfortunately not correct.
    <br/>
    "There is at least one horse owner, for whom it has at least one horse, for which it does not weigh"
    <br/>
</p>
<br/>
<p>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans8</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns4@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es 
gibt mindestens einen Pferdebesitzer, f&#252;r den gilt, dass er mindstens 
ein Pferd besitzt, f&#252;r das gilt, dass es nicht wiehert"
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns4@}" is unfortunately not correct.
    <br/>
    Right would be:
gives at least one horse owner to whom it applies at least
a horse, for which it does not weigh"
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans9</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort "{@SAns4@}" ist leider nicht korrekt.
    <br/>
    Richtig w&#228;re "Es 
gibt mindestens einen Pferdebesitzer, f&#252;r den gilt, dass er mindstens 
ein Pferd besitzt, f&#252;r das gilt, dass es nicht wiehert"
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer "{@SAns4@}" is unfortunately not correct.
    <br/>
    Right would be:
There is at least one horse owner to whom it applies that they have at least
one horse,  which it does not neigh"
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>