<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Beschreibende Form angeben  [{"DOMAINUID":"4368EDD"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]Stellen Sie die folgenden Mengen in beschreibender Form dar.
<br/>
Erg&#228;nzen Sie daf&#252;r jeweils noch die allgemeine Form der Elemente, d.h. f&#252;llen Sie die L&#252;cken aus.
<br/>
<p>
</p>
<p>
    a) \(M_1=\{ {@l1[1]@},{@l1[2]@},{@l1[3]@},{@l1[4]@},{@l1[5]@},\ldots\} = \Bigg\{ \)[[/lang]][[lang code="en"]]Set the following sets in a descriptive form.
    <br/>
    Add the general form of the elements, i.e. fill the gaps.
    <br/>
    <p>
    </p>
    <p>
        (a) \(M_1=\{ {@l1[1]@},{@l1[2]@},{@l1[3]@},{@l1[4]@},{@l1[5]@},\ldots\} = \Bigg\{ \)[[/lang]][[input:ans1]][[lang code="de"]] \(\Big|\ n \in \mathbb N_0 = \mathbb N \cup \{0\} \Bigg\} \)[[/lang]][[lang code="en"]] \(\Big|\ n \in \mathbb N_0 = \mathbb N \cup \{0\} \Bigg\} \)[[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
    </p>
    <p>
        b) \(M_2=\{ {@l2[1]@},{@l2[2]@},{@l2[3]@},{@l2[4]@},{@l2[5]@},\ldots\} = \Bigg \{ \)[[/lang]][[lang code="en"]]
    </p>
    <p>
    </p>
    <p>
        (b) \(M_2=\{ {@l2[1]@},{@l2[2]@},{@l2[3]@},{@l2[4]@},{@l2[5]@},\ldots\} = \Bigg \{ \)[[/lang]][[input:ans2]][[lang code="de"]] \(\Big|\ n \in \mathbb N \Bigg\} \)[[/lang]][[lang code="en"]] \(\Big|\ n \in \mathbb N \Bigg\} \)[[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
        c) \(M_3=\left\{ \displaystyle {@l3[1]@},{@l3[2]@},{@l3[3]@},{@l3[4]@},{@l3[5]@},\ldots\right\} = \Bigg \{
                \)[[/lang]][[lang code="en"]]
    </p>
    <p>
        (c) \(M_3=\left\{ \displaystyle {@l3[1]@},{@l3[2]@},{@l3[3]@},{@l3[4]@},{@l3[5]@},\ldots\right\} = \Bigg \{
                \)[[/lang]][[input:ans3]][[lang code="de"]] \(\Big|\ n \in \mathbb N\Bigg\}\) [[/lang]][[lang code="en"]] \(\Big|\ n \in \mathbb N\Bigg\}\) [[/lang]][[validation:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt3]][[lang code="de"]]
    </p>
    <p style="margin-top:2em;">
        <strong>
            Hinweis:
        </strong>
        <em>
        </em>
        <span style="font-size: 0.9375rem;">
            Beachten Sie, dass \(n\) bei Aufgabe a) Element der Menge \(\mathbb N_0=\mathbb N \cup \{0\}\) und bei den Aufgaben b) und c) Element der Menge \(\mathbb N\) ist.
        </span>
    </p>
    <p>
    </p>
    [[/lang]][[lang code="en"]]
</p>
<p style="margin-top:2em;">
    <strong>
        Note:
    </strong>
    <em>
    </em>
    <span style="font-size: 0.9375rem;">
        Note that \(n\) for task a) element of the set \(\mathbb N_0=\mathbb N \cup \{0\}\) (b) and (c) element of set \(\mathbb N\) is.
    </span>
</p>
<p>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        Diese Aufgabe kann wie folgt gel&#246;st werden:
    </b>
</p>
<p>
    Hier kann man die L&#246;sung durch geschicktes Ausprobieren herausfinden.
    <br/>
</p>
<p>
    <u>
        Beispiel (Aufgabe a)):
    </u>
    <br/>
    Es ist ein Term \(t(x)\) gesucht mit \(t(0)={@l1[1]@}\), \(t(1)={@l1[2]@}\),&#160;\(t(2)={@l1[3]@}\) usw.
    <br/>
    Geht man zun&#228;chst davon aus, dass der Term die Form \(t(x) = a \cdot x+b\) hat, folgt aus \(t(0)={@l1[1]@}\) sofort, dass \(b={@l1[1]@}\) gilt (also \(t(x)=ax+{@l1[1]@}\)).
    <br/>
    Setzt man nun \(t({\color{Green}1})={\color{Red} {@l1[2]@} }\) ein, erh&#228;lt man \({\color{Red}{@l1[2]@}}=a\cdot {\color{Green}1}+{@l1[1]@} \Leftrightarrow a={@n1@}\). Daraus folgt, dass der richtige Term \(t(x)={@ta1@}\) sein
    <i>
        k&#246;nnte
    </i>
    . Jetzt m&#252;ssen wir noch &#252;berpr&#252;fen, ob auch \(t(2)={@l1[3]@}\) usw. gilt. Dies ist hier der Fall, also sind wir fertig.
</p>
<p>
    Die anderen Teilaufgaben b) und c) funktionieren analog durch geschicktes Ausprobieren und &#220;berpr&#252;fen von Vermutungen.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        This object can be achieved as follows:
    </b>
</p>
<p>
    Here you can find out the solution through skillful testing.
    <br/>
</p>
<p>
    <u>
        Example (order a):
    </u>
    <br/>
    It's a term \(t(x)\) searched with \(t(0)={@l1[1]@}\), \(t(1)={@l1[2]@}\),\(t(2)={@l1[3]@}\) etc.
    <br/>
    If one assumes that the term is the form \(t(x) = a \cdot x+b\) the following \(t(0)={@l1[1]@}\) immediately that \(b={@l1[1]@}\) the following applies: \(t(x)=ax+{@l1[1]@}\)).
    <br/>
    You have to sit now \(t({\color{Green}1})={\color{Red} {@l1[2]@} }\) one, one gets \({\color{Red}{@l1[2]@}}=a\cdot {\color{Green}1}+{@l1[1]@} \Leftrightarrow a={@n1@}\). It follows that the correct term \(t(x)={@ta1@}\) they were
    <i>
        could be
    </i>
    . Now we have to check whether \(t(2)={@l1[3]@}\) etc. This is the case here, so we're done.
</p>
<p>
    The other subtasks b) and c) function analogously by skillful testing and checking guesses.
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>/*Erstellt von Stephan Bach, OTH Amberg-Weiden*/
/*Ver&#228;ndert von Jonas Lache, Ruhr-Universit&#228;t Bochum (Allgemeines Feedback)*/

n1:rand_with_step(3,7,1);
r:rand_with_step(1,n1-1,1);
n2:rand_with_prohib(3,5,[n1]);
n0:rand_with_step(2,4,1);

ta1:n1*n+r;
l1:makelist(ev(ta1,n=i),i,0,4);
oa1:n1*n-(n1-r);

ta2:n2*(n+n0);
l2:makelist(ev(ta2,n=i),i,1,5);

ta3:rand([n/(n+1),(n+2)/(n+1)]);
l3:makelist(ev(ta3,n=i),i,1,5);</text>
        </questionvariables>
        <specificfeedback format="plain_text">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>a) \(M_1=\{{@l1[1]@},{@l1[2]@},{@l1[3]@},{@l1[4]@},...\} = \{{@ta1@} | n \in \mathbb N_0 \}, \hspace{0.5cm}\)
b) \(M_2=\{ {@l2[1]@},{@l2[2]@},{@l2[3]@},{@l2[4]@},...\} = \{ {@ta2@} |\ n \in \mathbb N\}, \hspace{0.5cm} \)
c)  \(M_3=\{ {@l3[1]@},{@l3[2]@},{@l3[3]@},{@l3[4]@},{@l3[5]@},...\} = \{ {@ta3@}|\ n \in \mathbb N\} \)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Falsche Antwort.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>ta1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>ta2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>ta3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>2</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>oa1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    F&#252;r \(n \in \mathbb N\) w&#228;re Ihre Antwort richtig. Es gilt aber \(n \in \mathbb N_0\), und f&#252;r \(n=0\) ist \({@ans1@}={@ev(ans1,n=0)@} \notin M_1\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    For \(n \in \mathbb N\) your answer would be correct. However, \(n \in \mathbb N_0\)and for \(n=0\) is \({@ans1@}={@ev(ans1,n=0)@} \notin M_1\).
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>ta2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>ta3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1369694768</deployedseed>
        <deployedseed>1730328752</deployedseed>
        <deployedseed>4424115</deployedseed>
        <deployedseed>1144504690</deployedseed>
        <deployedseed>2065974215</deployedseed>
        <deployedseed>1054860652</deployedseed>
        <deployedseed>1624134075</deployedseed>
        <deployedseed>36945099</deployedseed>
        <deployedseed>588478007</deployedseed>
        <deployedseed>412881549</deployedseed>
        <deployedseed>1453916964</deployedseed>
        <deployedseed>48370691</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>ta1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>ta2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>ta3</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>oa1</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value/>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>prt1-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
            <tag>
                <text>L&#246;sungsidee</text>
            </tag>
        </tags>
    </question>
    

</quiz>