<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Betragsgleichung</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bestimmen Sie die L&#246;sungsmenge \(\mathbb{L}\) der Gleichung
        <span style="font-size: 0.9375rem;">
            \(\left\vert x + {@v1@}\right\vert = {@v2*x@}-{@v3@}\).
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            \(\mathbb{L}\,=\,\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
            <br/>
        </span>
    </p>
    <p>
        <strong>
        </strong>
    </p>
    <p>
        <strong>
            Hinweise:
        </strong>
    </p>
    <ul>
        <li>
            Schreiben Sie die Menge mit geschweiften Klammern und trennen Sie die L&#246;sungen mit Komma, z.B. {1, 3}
        </li>
        <li>
            Schreiben Sie leere Menge als {}.
        </li>
    </ul>
</span>
<span class="multilang" lang="en">
    <p>
        Determine the solution quantity \(\mathbb{L}\) the equation
        <span style="font-size: 0.9375rem;">
            \(\left\vert x + {@v1@}\right\vert = {@v2*x@}-{@v3@}\).
        </span>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            \(\mathbb{L}\,=\,\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
            <br/>
        </span>
    </p>
    <p>
        <strong>
        </strong>
    </p>
    <p>
        <strong>
            Notes:
        </strong>
    </p>
    <ul>
        <li>
            Write the amount with curly clips and separate the solutions with comma, e.g. {1, 3}
        </li>
        <li>
            Letter You empty quantity as {}.
        </li>
    </ul>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr">
        <strong>
            L&#246;sungsweg:
        </strong>
    </p>
    <p dir="ltr">
        Die Gleichung \(\left\vert x + {@v1@}\right\vert = {@v2*x@}-{@v3@}\) kann z.B. mit Hilfe der Fallunterscheidung gel&#246;st werden.
    </p>
    <p dir="ltr">
        <strong>
            <u>
                Fall 1:
            </u>
        </strong>
        \(\displaystyle \quad x \geq {@-v1@}\).
    </p>
    <p dir="ltr">
        \(x+{@v1@}={@v2*x@}-{@v3@}\) \(\quad \Longleftrightarrow&#160;\quad (1-{@v2@})\cdot x = -{@v3@}-{@v1@}\)&#160;\(\quad \Longleftrightarrow&#160;\quad ({@1-v2@})\cdot x = {@-v3-v1@}&#160;\)&#160;\(\quad \Longleftrightarrow&#160;\quad x = \displaystyle{@l1@}&#160;\) ist eine L&#246;sung, da \( \displaystyle{@l1@} \geq {@-v1@}\).
        <br/>
    </p>
    <p dir="ltr">
    </p>
    <p dir="ltr">
        <u>
            <strong>
            </strong>
        </u>
    </p>
    <p dir="ltr">
        <strong>
            <u>
                Fall 2:
            </u>
        </strong>
        \(\displaystyle \quad x &lt; {@-v1@}\).
    </p>
    <p dir="ltr">
    </p>
    <p dir="ltr">
        \(-(x+{@v1@})={@v2*x@}-{@v3@}\) \(\quad \Longleftrightarrow&#160;\quad -x-{@v1@}={@v2*x@}-{@v3@}\)&#160;\(\quad \Longleftrightarrow&#160;\quad (-1-{@v2@}) \cdot x=-{@v3@}+{@v1@}\)&#160;&#160;\(\quad \Longleftrightarrow&#160;\quad ({@-1-v2@}) \cdot x={@-v3+v1@}\)
        <span style="font-size: 0.9375rem;">
            \(\quad \Longleftrightarrow&#160;\quad x= \displaystyle{@l2@}\) ist keine L&#246;sung, da \( \displaystyle{@l2@} &gt; {@-v1@}\).
        </span>
    </p>
    <p dir="ltr">
        <span style="font-size: 0.9375rem;">
            Aus den beiden F&#228;llen ergibt sich dann die L&#246;sungsmenge \(\mathbb{L} = \displaystyle {@loesung@}\).
            <br/>
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr">
        <strong>
            Solution:
        </strong>
    </p>
    <p dir="ltr">
        The equation \(\left\vert x + {@v1@}\right\vert = {@v2*x@}-{@v3@}\) can be solved, for example, with the aid of the case differentiation.
    </p>
    <p dir="ltr">
        <strong>
            <u>
                Case 1:
            </u>
        </strong>
        \(\displaystyle \quad x \geq {@-v1@}\).
    </p>
    <p dir="ltr">
        \(x+{@v1@}={@v2*x@}-{@v3@}\) \(\quad \Longleftrightarrow&#160;\quad (1-{@v2@})\cdot x = -{@v3@}-{@v1@}\)\(\quad \Longleftrightarrow&#160;\quad ({@1-v2@})\cdot x = {@-v3-v1@}&#160;\)\(\quad \Longleftrightarrow&#160;\quad x = \displaystyle{@l1@}&#160;\) is a solution, since \( \displaystyle{@l1@} \geq {@-v1@}\).
        <br/>
    </p>
    <p dir="ltr">
    </p>
    <p dir="ltr">
        <u>
            <strong>
            </strong>
        </u>
    </p>
    <p dir="ltr">
        <strong>
            <u>
                Case 2:
            </u>
        </strong>
        \(\displaystyle \quad x &lt; {@-v1@}\).
    </p>
    <p dir="ltr">
    </p>
    <p dir="ltr">
        \(-(x+{@v1@})={@v2*x@}-{@v3@}\) \(\quad \Longleftrightarrow&#160;\quad -x-{@v1@}={@v2*x@}-{@v3@}\)\(\quad \Longleftrightarrow&#160;\quad (-1-{@v2@}) \cdot x=-{@v3@}+{@v1@}\)\(\quad \Longleftrightarrow&#160;\quad ({@-1-v2@}) \cdot x={@-v3+v1@}\)
        <span style="font-size: 0.9375rem;">
            \(\quad \Longleftrightarrow&#160;\quad x= \displaystyle{@l2@}\) is not a solution, since \( \displaystyle{@l2@} &gt; {@-v1@}\).
        </span>
    </p>
    <p dir="ltr">
        <span style="font-size: 0.9375rem;">
            The solution quantity is then obtained from the two cases. \(\mathbb{L} = \displaystyle {@loesung@}\).
            <br/>
        </span>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[v1: rand_with_step(1,9,1);
v2: rand_with_step(2,9,1); 
v3: rand_with_step(1,9,1); 
gl: |x+v1|=v2*x-v3
l1: (v1+v3)/(v2-1);
l2: (v3-v1)/(v2+1);
loesung: if(l1>=-v1) and (l2<-v1) then {l1,l2} else if(l1>=-v1) and (l2>-v1) then {l1} else if(l1<-v1) and (l2<-v1) then {l2} else {};]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@gl@}
L&#246;sung: {@loesung@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise richtig.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>loesung</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>loesung</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>