<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Determinante, Basis komplex</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Betrachte folgende Matrix
</p>
<p>
    \(A={@A@}\)
</p>
<p>
    <br/>
</p>
<p>
    <b>
        (a)
    </b>
    Berechen sie die Determinante von \(A\).
</p>
<p>
    \(\det(A)=\) [[/lang]][[lang code="en"]]
    <p>
    </p>
    <p>
        View the following matrix
    </p>
    <p>
        \(A={@A@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <b>
            (a)
        </b>
        Calculate the determinant of \(A\).
    </p>
    <p>
        \(\det(A)=\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        <br/>
    </p>
    <p>
        <b>
            (b)
        </b>
        Die Matrix hat den Rang \(\text{rg}(A)=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        <br/>
    </p>
    <p>
        <b>
            (b)
        </b>
        The matrix has the rank \(\text{rg}(A)=\) [[/lang]][[input:ans2]][[lang code="de"]] und die Spaltenvektoren bilden somit&#160; [[/lang]][[lang code="en"]] and the column vectors thus form [[/lang]][[input:ans3]][[lang code="de"]] Basis des \(\mathbb{C}^3\).[[/lang]][[lang code="en"]] Based on \(\mathbb{C}^3\).[[/lang]][[validation:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans3]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt3]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <strong>
        L&#246;sung
    </strong>
</h4>
<p>
    <br/>
</p>
<p>
    <strong>
        (a)
    </strong>
    \(A\) ist eine Dreiecksmatrix (Alle Eintr&#228;ge unter der Diagonale sind Null). F&#252;r solche Matrizen ist die Determinante das Produkt der Diagonaleintr&#228;ge:
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \(\det(A)=(1+{@b@}i)(-i)({@a@}+i)={@det@}\)
    </p>
    <p>
        <br/>
    </p>
</div>
<span style="">
    <span style="">
        <strong>
            (b)
        </strong>
        Der Rang der Matrix ist voll, also \(\text{rg}(A)=3\). Das l&#228;sst sich entweder direkt aus der Matrix ablesen, da sie keine Gau&#223;-Umformung mehr ben&#246;tigt oder man argumentiert mit \(\det(A)\neq 0\). Die Spaltenvektoren bilden bei vollem Rang immer auch eine Basis, da sie damit linear unabh&#228;ngig sind.
    </span>
</span>
<p>
    <span style="">
        <span style="">
            <br/>
        </span>
    </span>
</p>
<p>
    <span style="">
        <span style="">
            <br/>
        </span>
    </span>
</p>
[[/lang]][[lang code="en"]]
<h4>
    <strong>
        Solution
    </strong>
</h4>
<p>
    <br/>
</p>
<p>
    <strong>
        (a)
    </strong>
    \(A\) is a triangular matrix (all entries under the diagonal are zero). For such matrices, the determinant is the product of the diagonal entries:
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \(\det(A)=(1+{@b@}i)(-i)({@a@}+i)={@det@}\)
    </p>
    <p>
        <br/>
    </p>
</div>
<span style="">
    <span style="">
        <strong>
            (b)
        </strong>
        The rank of the matrix is full, so \(\text{rg}(A)=3\). This can be seen either directly from the matrix, as it no longer needs Gaussian transformation or you argue with \(\det(A)\neq 0\). The column vectors always form a base with full rank, since they are therefore linearly independent.
    </span>
</span>
<p>
    <span style="">
        <span style="">
            <br/>
        </span>
    </span>
</p>
<p>
    <span style="">
        <span style="">
            <br/>
        </span>
    </span>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a: -2-rand(3)
b: 2+rand(3)
A: matrix([1+b*%i,%i,-1],[0,-%i, 2+3*%i],[0,0,a+%i])
det: expand((1+b*%i)*(-%i)*(a+%i))

a1: 3
a2: 1
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]")
al: [[1, true, tr("eine","one")], [2, false, tr("keine","none")]]]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>Rang</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>det</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>a1</tans>
            <boxsize>2</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>radio</type>
            <tans>al</tans>
            <boxsize>2</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>det</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>a1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>a2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Rang</text>
            </tag>
            <tag>
                <text>TE:3:Basis</text>
            </tag>
            <tag>
                <text>TE:3:Determinante</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:2</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>