<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Entscheiden, ob Funktionen injektiv/surjektiv (2) [{"DOMAINUID":"D4A8030"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Entscheiden Sie bei folgenden Funktionen, ob sie injektiv oder surjektiv sind, indem Sie den Wahrheitsgehalt folgender Aussagen beurteilen.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    In the following functions, decide whether they are subjective or surjective by assessing the truth of the following statements.
    <br/>
</p>
[[/lang]][[validation:ans1]][[lang code="de"]]




[[/lang]][[lang code="en"]][[/lang]][[validation:ans2]][[lang code="de"]]
<table>
    <tbody>
        <tr>
            <td>
                Die Funktion \(f:\mathbb{R}\to\mathbb{R}\) mit \(f(x)={@f(x)@}\) ist injektiv.
            </td>
            <td>
                [[/lang]][[lang code="en"]]
                <table>
                    <tbody>
                        <tr>
                            <td>
                                The function \(f:\mathbb{R}\to\mathbb{R}\) with \(f(x)={@f(x)@}\) is injective.
                            </td>
                            <td>
                                [[/lang]][[input:ans1]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                Die Funktion \(g:\mathbb{R}\to\mathbb{R}\) mit \(g(x)={@g(x)@}\) ist surjektiv.
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                The function \(g:\mathbb{R}\to\mathbb{R}\) with \(g(x)={@g(x)@}\) is surjective.
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans2]][[lang code="de"]]
                            </td>
                        </tr>
                    </tbody>
                </table>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
    </tbody>
</table>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[nf:5+rand(4);
ng:5+rand(4);
c:-5+rand(10);
f(x):=x^nf-c;
g(x):=x^ng+x^2;
f_injective:is(mod(nf,2)=1);
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
f_text:[tr("nicht injektiv","not injective"),tr("injectiv")][1+mod(nf,2)];
g_surjective:is(mod(ng,2)=1);
g_text:[tr("nicht surjektiv","not surjective"),tr("surjektiv","surjective")][1+mod(ng,2)];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(f(x)={@f(x)@}\)
\(g(x)={@g(x)@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>boolean</type>
            <tans>f_injective</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>boolean</type>
            <tans>g_surjective</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>2.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>f_injective</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Funktion \(f(x)={@f(x)@}\) ist {@f_text@}. Das haben Sie richtig gemacht.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The function \(f(x)={@f(x)@}\) is {@f_text@}. You did that right.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Funktion \(f(x)={@f(x)@}\) ist {@f_text@}. Das haben Sie leider falsch gemacht.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The function \(f(x)={@f(x)@}\) is {@f_text@}. You've done that wrong.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>g_surjective</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Funktion \(g(x)={@g(x)@}\) ist {@g_text@}. Das haben Sie richtig gemacht.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The function \(g(x)={@g(x)@}\) is {@g_text@}. You did that right.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Funktion \(g(x)={@g(x)@}\) ist {@g_text@}. Das haben Sie leider falsch gemacht.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The function \(g(x)={@g(x)@}\) is {@g_text@}. You've done that wrong.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>