<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Gleichungen: Wurzelgleichungen (1)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bestimmen Sie die L&#246;sungsmenge der Gleichung {@gl@}
    </p>
    <p>
        \(\mathbb{L} = \) [[input:ans]] [[validation:ans]][[feedback:prt1]]
    </p>
    <p>
    </p>
    <p>
        <strong>
            Hinweise:
        </strong>
    </p>
    <ul>
        <li>
            Schreiben Sie die Menge mit geschweiften Klammern und trennen Sie die L&#246;sungen mit Komma, z.B. {1, 3}
        </li>
        <li>
            Schreiben Sie leere Menge als {}.
        </li>
        <li>
            Geben Sie die exakte L&#246;sung (keine Dezimalzahlen) ein.
        </li>
    </ul>
</span>
<span class="multilang" lang="en">
    <p>
        Determine the solution of the equation {@gl@}
    </p>
    <p>
        \(\mathbb{L} = \) [[input:ans]] [[validation:ans]][[feedback:prt1]]
    </p>
    <p>
    </p>
    <p>
        <strong>
            Notes:
        </strong>
    </p>
    <ul>
        <li>
            Write the amount with curly clips and separate the solutions with comma, e.g. {1, 3}
        </li>
        <li>
            Letter You empty quantity as {}.
        </li>
        <li>
            Enter the exact solution (no decimal numbers).
        </li>
    </ul>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr">
        <strong>
            L&#246;sungsweg:
        </strong>
    </p>
    <p dir="ltr">
        Die vorgegebene Gleichung&#160;\({@sqrt(x^2+v1)=x+v2@}\) kann beidseitig quadriert werden, unter der beiden Voraussetzungen:
    </p>
    <p dir="ltr">
    </p>
    <ul>
        <li>
            \({@x^2+v1@} \geq 0\), was f&#252;r alle \(x \in \mathbb{R}\) gilt.
        </li>
        <li>
            \({@x+v2@} \geq 0 \quad \Longleftrightarrow&#160;\quad x \geq {@-v2@}\)
        </li>
    </ul>
    \(x^2+{@v1@} = x^2 + {@2*x*v2+v2*v2@}\) \(\quad \Longleftrightarrow&#160;\quad {@v1 = 2*x*v2+v2*v2@}\)&#160;\(\quad \Longleftrightarrow&#160;\quad {@2*x*v2+v2*v2-v1=0@}\)&#160;\(\quad \Longleftrightarrow&#160;\quad {@2*x*v2=-v2^2+v1@}\)&#160;\(\quad \Longleftrightarrow&#160;\quad \displaystyle x={@(-v2*v2+v1)/(2*v2)@}\) ist die L&#246;sung, weil \(\displaystyle {@l1@} \geq {@-v2@}\).
    <br/>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr">
        <strong>
            Solution:
        </strong>
    </p>
    <p dir="ltr">
        The given equation\({@sqrt(x^2+v1)=x+v2@}\) can be squared on both sides, under the two conditions:
    </p>
    <p dir="ltr">
    </p>
    <ul>
        <li>
            \({@x^2+v1@} \geq 0\), what for all \(x \in \mathbb{R}\) Article 2
        </li>
        <li>
            \({@x+v2@} \geq 0 \quad \Longleftrightarrow&#160;\quad x \geq {@-v2@}\)
        </li>
    </ul>
    \(x^2+{@v1@} = x^2 + {@2*x*v2+v2*v2@}\) \(\quad \Longleftrightarrow&#160;\quad {@v1 = 2*x*v2+v2*v2@}\)\(\quad \Longleftrightarrow&#160;\quad {@2*x*v2+v2*v2-v1=0@}\)\(\quad \Longleftrightarrow&#160;\quad {@2*x*v2=-v2^2+v1@}\)\(\quad \Longleftrightarrow&#160;\quad \displaystyle x={@(-v2*v2+v1)/(2*v2)@}\) is the solution because \(\displaystyle {@l1@} \geq {@-v2@}\).
    <br/>
    <p>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>v1: rand_with_step(1,8,1);
v2: rand_with_step(1,4,1);
gl: sqrt(x^2+v1)=x+v2;
l1: (v1-v2^2)/(2*v2);
loesung: {l1};</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr">
    </p>
    <br/>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr">
    </p>
    <br/>
    <p>
    </p>
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@gl@}L&#246;sung: {@loesung@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise richtig.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans</name>
            <type>algebraic</type>
            <tans>loesung</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint>{?,?,...}</syntaxhint>
            <syntaxattribute>1</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans</sans>
                <tans>loesung</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>1-0-T </trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>1-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Quadrieren Sie die Gleichung zuerst auf beiden Seiten, um die Quadradwurzel aufzul&#246;sen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Quadrants First, use the equation on both sides to dissolve the quadricycle root.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>1-0-T</expectedanswernote>
            </expected>
        </qtest>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
            <tag>
                <text>R&#252;ckmeldebaum</text>
            </tag>
            <tag>
                <text>TE:1:Grundlagen</text>
            </tag>
            <tag>
                <text>TE:2:Gleichungen</text>
            </tag>
            <tag>
                <text>TE:3:Wurzelgleichungen</text>
            </tag>
            <tag>
                <text>te:3:L&#246;sungsmenge</text>
            </tag>
            <tag>
                <text>src:F09-IPK</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:1</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>