<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Gleichungen Log,e 1 [{"DOMAINUID":"36C704E"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span>
            Geben Sie alle reellen L&#246;sungen der Gleichung \(\displaystyle {@gl2@} = {@-d@}\) an.
        </span>
    </p>
    <p>
        Diese Gleichung l&#228;sst sich durch Potenzieren auf eine lineare Gleichung zur&#252;ckf&#252;hren.
    </p>
    <p>
        Diese lineare Gleichung ist: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        <br/>
    </p>
    <p>
        Die L&#246;sungsmenge der Gleichung ist also \(\mathbb{L}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        <br/>
    </p>
    <p>
    </p>
    <p>
        <strong>
            Hinweise:
        </strong>
    </p>
    <ul>
        <li>
            Schreiben Sie die Menge mit geschweiften Klammern und trennen Sie die L&#246;sungen mit Komma, z.B. {1, 3}
        </li>
        <li>
            Schreiben Sie leere Menge als {}.
        </li>
    </ul>
</span>
<span class="multilang" lang="en">
    <p>
        <span>
            Specify all real solutions of the equation \(\displaystyle {@gl2@} = {@-d@}\) on.
        </span>
    </p>
    <p>
        This equation can be traced back to a linear equation by potentiating.
    </p>
    <p>
        This linear equation is: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        <br/>
    </p>
    <p>
        The solution quantity of the equation is thus \(\mathbb{L}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        <br/>
    </p>
    <p>
    </p>
    <p>
        <strong>
            Notes:
        </strong>
    </p>
    <ul>
        <li>
            Write the amount with curly clips and separate the solutions with comma, e.g. {1, 3}
        </li>
        <li>
            Letter You empty quantity as {}.
        </li>
    </ul>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <strong>
            L&#246;sungsweg:
        </strong>
    </p>
    <p>
        \(\displaystyle {@gl2@} = {@-d@}\) \(\quad \Longleftrightarrow&#160;\quad -\ln({@a@} \cdot x-{@b@}) = -{@d@}&#160;\cdot \ln({@f@})\)
        <span style="font-size: 0.9375rem;">
            \(\quad \Longleftrightarrow&#160;\quad \ln({@a@}&#160;\cdot&#160;x-{@b@}) = \ln({@f@}^{@d@})\)
        </span>
        <span style="font-size: 0.9375rem;">
            <span>
                \(\quad \Longleftrightarrow&#160;\quad {@a@}&#160;\cdot&#160;x-{@b@} = {@f@}^{@d@}\)
            </span>
            <span>
                <span>
                    \(\quad \Longleftrightarrow&#160;\quad x = {@LM1@}\)
                </span>
                <br/>
            </span>
            <br/>
            <br/>
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <strong>
            Solution:
        </strong>
    </p>
    <p>
        \(\displaystyle {@gl2@} = {@-d@}\) \(\quad \Longleftrightarrow&#160;\quad -\ln({@a@} \cdot x-{@b@}) = -{@d@}&#160;\cdot \ln({@f@})\)
        <span style="font-size: 0.9375rem;">
            \(\quad \Longleftrightarrow&#160;\quad \ln({@a@}&#160;\cdot&#160;x-{@b@}) = \ln({@f@}^{@d@})\)
        </span>
        <span style="font-size: 0.9375rem;">
            <span>
                \(\quad \Longleftrightarrow&#160;\quad {@a@}&#160;\cdot&#160;x-{@b@} = {@f@}^{@d@}\)
            </span>
            <span>
                <span>
                    \(\quad \Longleftrightarrow&#160;\quad x = {@LM1@}\)
                </span>
                <br/>
            </span>
            <br/>
            <br/>
        </span>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>a: 1+rand(5);
b: 1+rand(5);
c: 1+rand(5);
d: rand([1,2,3]);
f: 1+rand(5);
lb(x):=log(x)/log(1/f);
lg(x):=log(x)/log(10);
tans1:a*x-b=f^d
tans: a*x-b
gl1:lb(tans)= -d;
gl2:lb(tans);
LM:setify(map(rhs,radcan(solve(tans1,x))));
sol:solve(tans1,x);
LM1:rhs(sol[1]);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>Berechne \({@x@}={@solve(gl1,x)@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise richtig.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>arccos</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>tans1</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>solve</forbidwords>
            <allowwords>lb,lg</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>LM</tans>
            <boxsize>20</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>solve</forbidwords>
            <allowwords>lb,lg</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>SysEquiv</answertest>
                <sans>[ans1]</sans>
                <tans>[tans1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>+</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>LM</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2</deployedseed>
        <deployedseed>5</deployedseed>
        <deployedseed>10</deployedseed>
        <deployedseed>18</deployedseed>
        <deployedseed>33</deployedseed>
        <deployedseed>44</deployedseed>
        <deployedseed>77</deployedseed>
        <deployedseed>99</deployedseed>
        <deployedseed>111</deployedseed>
        <deployedseed>666</deployedseed>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>