<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Inverse lineare Funktion [{"DOMAINUID":"4BFAC29"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<div>
    <span style="font-size: 0.9375rem;">
        Bestimmen Sie die Umkehrfunktion
    </span>
    <span style="font-size: 0.9375rem;">
        \(f^{-1}\) der Funktion \(\displaystyle f(x) = {@v1@} \cdot x + {@v2@}\).
    </span>
    <br/>
</div>
<p>
    \(f^{-1}(x) = \) [[/lang]][[lang code="en"]]
    <div>
        <span style="font-size: 0.9375rem;">
            Determine the inverse function
        </span>
        <span style="font-size: 0.9375rem;">
            \(f^{-1}\) the function \(\displaystyle f(x) = {@v1@} \cdot x + {@v2@}\).
        </span>
        <br/>
    </div>
    <p>
        \(f^{-1}(x) = \) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <strong>
        L&#246;sungsweg:
    </strong>
</p>
<p>
    Gesucht wird die Umkehrfunktion der Funktion \(\displaystyle f(x)={@v1@}\cdot x + {@v2@}\).
</p>
<p>
    Dazu l&#246;sen wir die Gleichung
    <span style="font-size: 0.9375rem;">
        \(\displaystyle y={@v1@}\cdot x + {@v2@}\)&#160;nach \(x\) auf.
    </span>
</p>
<p>
    Wir erhalten
    <span style="font-size: 0.9375rem;">
        \(\displaystyle x={@1/v1@}\cdot ( {@y-v2@})\).
    </span>
</p>
<p>
    Jetzt benennen wir die Variablen um und erhalten
    <span style="font-size: 0.9375rem;">
        \(\qquad \displaystyle f^{-1}(x) = {@tans1@}\).
    </span>
</p>
[[/lang]][[lang code="en"]]
<p>
    <strong>
        Solution:
    </strong>
</p>
<p>
    The aim is to reverse the function \(\displaystyle f(x)={@v1@}\cdot x + {@v2@}\).
</p>
<p>
    We solve the equation
    <span style="font-size: 0.9375rem;">
        \(\displaystyle y={@v1@}\cdot x + {@v2@}\)After \(x\) up.
    </span>
</p>
<p>
    We receive
    <span style="font-size: 0.9375rem;">
        \(\displaystyle x={@1/v1@}\cdot ( {@y-v2@})\).
    </span>
</p>
<p>
    Now we rename and receive the variables
    <span style="font-size: 0.9375rem;">
        \(\qquad \displaystyle f^{-1}(x) = {@tans1@}\).
    </span>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>a:rand([1, 8]);
b: rand([2,12]);
v2: rand ([1,9]);
v1:a/b;
tans1: (x-v2)/v1;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    [[/lang]][[lang code="en"]]
    <p>
        [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\({@v1@} \cdot x + {@v2@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<strong>
    Richtige Antwort, gut gemacht!
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
<strong>
    Right answer, well done!
</strong>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<strong>
    Ihre Antwort ist teilweise korrekt.
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
<strong>
    Your answer is partly correct.
</strong>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<strong>
    Falsche Antwort.
</strong>
[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
<strong>
    Wrong answer.
</strong>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>arccos</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>tans1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>tans1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>