<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Ist x ein Element/Teilmenge der Menge?  [{"DOMAINUID":"635F157"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Beantworten Sie folgende Fragen:
</p>
<table>
    <tbody>
        <tr>
            <td>
                a)
            </td>
            <td>
                Ist \(0\) ein Element von \(\mathbb{Z}\)?
                <br/>
            </td>
            <td>
                [[/lang]][[lang code="en"]]
                <p>
                    Answer the following questions:
                </p>
                <table>
                    <tbody>
                        <tr>
                            <td>
                                (a)
                            </td>
                            <td>
                                I \(0\) an element of \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans1]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt1]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                b)
                            </td>
                            <td>
                                Ist \(\emptyset\) ein Element von \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (b)
                            </td>
                            <td>
                                Is \(\emptyset\) an element of \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans2]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt2]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                c)
                            </td>
                            <td>
                                Ist \(\{\emptyset\}\) ein Element von \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (c)
                            </td>
                            <td>
                                Is \(\{\emptyset\}\) an element of \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans3]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt3]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                d)
                            </td>
                            <td>
                                Ist \(0\) eine Teilmenge von \(\mathbb{Z}\)?
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (d)
                            </td>
                            <td>
                                Is \(0\) a subset of \(\mathbb{Z}\)?
                            </td>
                            <td>
                                [[/lang]][[input:ans4]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt4]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                e)
                            </td>
                            <td>
                                Ist \(\emptyset\) eine Teilmenge von \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (e)
                            </td>
                            <td>
                                I \(\emptyset\) a subset of \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans5]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt5]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                f)
                            </td>
                            <td>
                                Ist \(\{\emptyset\}\) eine Teilmenge von \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (f)
                            </td>
                            <td>
                                Is \(\{\emptyset\}\) a subset of \(\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans6]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt6]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                g)
                            </td>
                            <td>
                                Ist \(\emptyset\) ein Element von \(\emptyset\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                g)
                            </td>
                            <td>
                                Is \(\emptyset\) an element of \(\emptyset\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans7]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt7]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                h)
                            </td>
                            <td>
                                Ist \(\emptyset\) ein Element von \(\{\emptyset\}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                h)
                            </td>
                            <td>
                                Is \(\emptyset\) an element of \(\{\emptyset\}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans8]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt8]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                i)
                            </td>
                            <td>
                                Ist \(\emptyset\) ein Element von \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                (i)
                            </td>
                            <td>
                                Is \(\emptyset\) an element of \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans9]][[lang code="de"]]
                                <br/>
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                                <br/>
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt9]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                j)
                            </td>
                            <td>
                                Ist \(\emptyset\) eine Teilmenge von \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                                j)
                            </td>
                            <td>
                                Is \(\emptyset\) a subset of \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\)?
                                <br/>
                            </td>
                            <td>
                                [[/lang]][[input:ans10]][[lang code="de"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[lang code="en"]]
                            </td>
                        </tr>
                        <tr>
                            <td>
                            </td>
                            <td colspan="2">
                                [[/lang]][[feedback:prt10]][[lang code="de"]]
                            </td>
                        </tr>
                    </tbody>
                </table>
                <p>
                </p>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
    </tbody>
</table>
<p>
</p>
[[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans4]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans5]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans6]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans7]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans8]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans9]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans10]][[lang code="de"]]
<br/>
[[/lang]][[lang code="en"]]
<br/>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>n1:rand(10);
n2:15+rand(10);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@n1@},{@n2@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans10</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans7</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans8</name>
            <type>boolean</type>
            <tans>true</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans9</name>
            <type>boolean</type>
            <tans>false</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]Denn: Die Menge \(\mathbb{Z}\) umfasst alle Zahlen \(\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\).
<br/>
[[/lang]][[lang code="en"]]Because: The set \(\mathbb{Z}\) All numbers \(\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\).
<br/>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: Die Menge&#160;\(\mathbb{Z}\) umfasst alle Zahlen \(\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: The set\(\mathbb{Z}\) All numbers \(\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt10</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans10</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt10-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Denn: Die leere Menge \(\emptyset\) ist eine Teilmenge von jeder Menge.
    <br/>
</p>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    For: The empty set \(\emptyset\) is a part of any set.
    <br/>
</p>
<p>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt10-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Denn: Die leere Menge \(\emptyset\) ist eine Teilmenge von jeder Menge.
    <br/>
</p>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    For: The empty set \(\emptyset\) is a part of any set.
    <br/>
</p>
<p>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\emptyset\) ist eine Menge und \(\mathbb{Z}=\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\) hat nur Zahlen und nicht die
    Menge \(\emptyset\) als Elemente.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\emptyset\) is a set and \(\mathbb{Z}=\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\) has only numbers and not the
set \(\emptyset\) as elements.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\emptyset\) ist eine Menge und \(\mathbb{Z}=\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\) hat nur Zahlen und nicht die
    Menge \(\emptyset\) als Elemente.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\emptyset\) is a set and \(\mathbb{Z}=\{&#160; \ldots,-3,-2,-1,0,1,2,3,\ldots \}\) has only numbers and not the
set \(\emptyset\) as elements.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\{\emptyset\}\) ist eine Menge, die nur die leere Menge \(\emptyset\) als Element enth&#228;lt. Diese Menge ist nicht als Element in \(\mathbb{Z}=\{&#160; 
\dots,-5,-4,-3,-2,-1,0,1,2,3,4,5,\dots \}\) enthalten.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\{\emptyset\}\) is a set that only contains the empty set \(\emptyset\) as an element. This set is not as an element in \(\mathbb{Z}=\{&#160; 
\dots,-5,-4,-3,-2,-1,0,1,2,3,4,5,\dots \}\) included.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\{\emptyset\}\) ist eine 
Menge, die nur die leere Menge \(\emptyset\) als Element enth&#228;lt. Diese 
Menge ist nicht als Element in \(\mathbb{Z}=\{&#160; 
\dots,-5,-4,-3,-2,-1,0,1,2,3,4,5,\dots \}\) enthalten.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\{\emptyset\}\) is a
set that only contains the empty set \(\emptyset\) as an element. This
set is not as an element in \(\mathbb{Z}=\{&#160; 
\dots,-5,-4,-3,-2,-1,0,1,2,3,4,5,\dots \}\) included.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(0\) ist eine Zahl und keine Menge und kann daher keine Teilmenge von \(\mathbb{Z}\) sein.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(0\) is a number and no set and therefore cannot be a subset of \(\mathbb{Z}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(0\) ist eine Zahl und keine Menge und kann daher keine Teilmenge von \(\mathbb{Z}\) sein.
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(0\) is a number and no set and therefore cannot be a subset of \(\mathbb{Z}\).
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: Die leere Menge \(\emptyset\) ist eine Teilmenge von jeder Menge.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    For: The empty set \(\emptyset\) is a subset of any set.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: Die leere Menge \(\emptyset\) ist eine Teilmenge von jeder Menge.
</p>
[[/lang]][[lang code="en"]]
<p>
    For: The empty set \(\emptyset\) is a subset of any set.
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\{\emptyset\}\) ist eine Menge, die nur die leere Menge \(\emptyset\) als Element enth&#228;lt. Diese Menge ist keine Teilmenge von \(\mathbb{Z}\), denn \(\mathbb{Z}\) enth&#228;lt \(\emptyset\) nicht als Element.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\{\emptyset\}\) is a set that contains only the empty set \(\emptyset\) as an element. This set is not a subset of \(\mathbb{Z}\)for \(\mathbb{Z}\) does not contain \(\emptyset\) as an element.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\{\emptyset\}\) ist eine
                Menge, die nur die leere Menge \(\emptyset\) als Element enth&#228;lt. Diese
                Menge ist keine Teilmenge von \(\mathbb{Z}\), denn \(\mathbb{Z}\)
                enth&#228;lt \(\emptyset\) nicht als Element.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\{\emptyset\}\) is a
set that contains only the empty set \(\emptyset\) as an element. This
set is not a subset \(\mathbb{Z}\)for \(\mathbb{Z}\) does not contain \(\emptyset\) as an element.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt7</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans7</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Denn: \(\emptyset\) ist die leere Menge und enth&#228;lt daher keine Elemente (auch nicht sich selbst).
    <br/>
</p>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    Because: \(\emptyset\) is the empty set and therefore contains no elements (also not itself).
    <br/>
</p>
<p>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt7-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\emptyset\) ist die leere Menge und enth&#228;lt daher keine Elemente (auch nicht sich selbst).
    <br/>
</p>
<p>
</p>
<p>
</p>
<br/>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\emptyset\) is the empty set and therefore contains no elements (also not itself).
    <br/>
</p>
<p>
</p>
<p>
</p>
<br/>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt8</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans8</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt8-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: \(\{\emptyset\}\) ist die Menge, die die leere Menge \(\emptyset\) als Element enth&#228;lt. Somit gilt \(\emptyset \in \{\emptyset\}\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: \(\{\emptyset\}\) is the set that the empty set \(\emptyset\) is an element of. This applies \(\emptyset \in \{\emptyset\}\).
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt8-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    Denn: \(\{\emptyset\}\) ist die Menge, die die leere Menge \(\emptyset\) als Element enth&#228;lt. Somit gilt \(\emptyset \in \{\emptyset\}\).
    <br/>
</p>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    Because: \(\{\emptyset\}\) is the set that the empty set \(\emptyset\) is an element of. This applies \(\emptyset \in \{\emptyset\}\).
    <br/>
</p>
<p>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt9</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans9</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt9-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: Die Menge \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\) besteht aus allen Tupeln \((x,y)\), mit \(x\in \{{@n1@},\ldots,{@n2@}\}\) und \(y \in \mathbb{Z}\). Daher ist die leere Menge \(\emptyset\) kein Element dieser Menge.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: The set \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\) consists of all tuples \((x,y)\)with \(x\in \{{@n1@},\ldots,{@n2@}\}\) and \(y \in \mathbb{Z}\). Therefore, the empty set \(\emptyset\) is not an element of this set.
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt9-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Denn: Die Menge \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\) besteht aus 
allen Tupeln \((x,y)\), mit \(x\in \{{@n1@},\ldots,{@n2@}\}\) und \(y \in
    \mathbb{Z}\). Daher ist die leere Menge \(\emptyset\) kein Element 
dieser Menge.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    Because: The set \(\{{@n1@},\ldots,{@n2@}\}\times\mathbb{Z}\) consists of:
all tuples \((x,y)\)with \(x\in \{{@n1@},\ldots,{@n2@}\}\) and \(y \in
    \mathbb{Z}\). Therefore, the empty set \(\emptyset\) is no element of
this set.
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>R&#252;ckmeldebaum</text>
            </tag>
        </tags>
    </question>
    

</quiz>