<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Kombinatorik 1</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Am Ende einer Konferenz (... noch vor Corona ...) sch&#252;ttelten sich \({@n@}\) Teilnehmer pressegerecht die H&#228;nde.
    </p>
    <p>
        Wie viele Sch&#252;ttelvorg&#228;nge gab es?
    </p>
    <p>
        [[input:ans1]] [[validation:ans1]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        At the end of a conference (... still before Corona...) \({@n@}\) Participants press their hands properly.
    </p>
    <p>
        How many shaking processes were there?
    </p>
    <p>
        [[input:ans1]] [[validation:ans1]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr" style="text-align: left;">
        <strong>
            L&#246;sungsweg:
        </strong>
    </p>
    <p dir="ltr" style="text-align: left;">
        \(&#160; &#160;\dfrac{{@n@}!}{2! \cdot ({@n@}-2)!}=&#160;&#160;\dfrac{{@n@}!}{2! \cdot {@t@}!}&#160; =&#160;\dfrac{{@t@}! \cdot {@n1@}\cdot {@n@}}{2! \cdot {@t@}!}&#160; =&#160; {@ta1@}\)
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr" style="text-align: left;">
        <strong>
            Solution:
        </strong>
    </p>
    <p dir="ltr" style="text-align: left;">
        \(&#160; &#160;\dfrac{{@n@}!}{2! \cdot ({@n@}-2)!}=&#160;&#160;\dfrac{{@n@}!}{2! \cdot {@t@}!}&#160; =&#160;\dfrac{{@t@}! \cdot {@n1@}\cdot {@n@}}{2! \cdot {@t@}!}&#160; =&#160; {@ta1@}\)
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>n:rand_with_step(10,20,1);
ta1:binomial(n,2);
t:n-2;
n1:n-1;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        [[feedback:prt1]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        [[feedback:prt1]]
    </p>
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@n@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    <strong>
        Richtige Antwort, gut gemacht!
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    <strong>
        Right answer, well done!
    </strong>
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    <strong>
        Ihre Antwort ist teilweise korrekt.
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    <strong>
        Your answer is partly correct.
    </strong>
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    <strong>
        Falsche Antwort.
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    <strong>
        Wrong answer.
    </strong>
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>ta1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>ta1</tans>
                <testoptions>.01</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>+</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>TE:1:Wahrscheinlichkeitsrechnung</text>
            </tag>
            <tag>
                <text>TE:2:Kombinatorik</text>
            </tag>
            <tag>
                <text>TE:3:Summen Produkte</text>
            </tag>
            <tag>
                <text>src:F04-IVW</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:3</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:2</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>