<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Komplexe Zahlen: Gegeben ist z1, geben sie z2 mit Bedingungen [{"DOMAINUID":"A24C23B"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Gegeben ist die komplexe Zahl \( \, z_1={@z1@} \, \).
    </p>
    <p>
        Geben Sie eine komplexe Zahl \( \, z_2 \neq 0 \, \) so, dass \(\ldots\)
    </p>
    <p>
        a) \(\ldots\) die komplexe Zahl \( z_1+z_2 \, \) reell ist.
    </p>
    <p>
        \(z_2=\,\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        b) \(\ldots\) die komplexe Zahl&#160;\(z_1 \cdot z_2 \, \) reell ist.
    </p>
    <p>
        \(z_2=\,\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
    <p>
        c) \(\ldots\) die beiden komplexen Zahlen \( \, z_1 + z_2 \, \) und \( \, z_1 \cdot z_2 \, \) reell sind.
    </p>
    <p>
        \(z_2=\,\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        given is the complex number \( \, z_1={@z1@} \, \).
    </p>
    <p>
        Give a complex number \( \, z_2 \neq 0 \, \) so that \(\ldots\)
    </p>
    <p>
        (a) \(\ldots\) the complex number \( z_1+z_2 \, \) reell.
    </p>
    <p>
        \(z_2=\,\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        (b) \(\ldots\) the complex number\(z_1 \cdot z_2 \, \) reell.
    </p>
    <p>
        \(z_2=\,\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
    <p>
        (c) \(\ldots\) the two complex numbers \( \, z_1 + z_2 \, \) and \( \, z_1 \cdot z_2 \, \) are real.
    </p>
    <p>
        \(z_2=\,\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <strong>
            L&#246;sungsidee:
        </strong>
    </p>
    <p>
        Eine
        <strong>
            m&#246;gliche L&#246;sung
        </strong>
        der Aufgabe w&#228;re die folgende.
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            a) Wenn \(z_1={@z1@}\) die vorgegebene komplexe Zahl ist, dann kann z.B. die komplexe Zahl \(z_2 = {@loes1@} \neq 0 \) gew&#228;hlt werden, so dass die komplexe Zahl \(z_1+z_2 = ({@z1@}) + ({@loes1@}) = {@z1+loes1@}\) reell ist.
        </span>
        <br/>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            b) Wenn \(z_1={@z1@}\) die vorgegebene komplexe Zahl ist, dann kann z.B. die komplexe Zahl \(z_2 = {@loes2@} \neq 0 \) gew&#228;hlt werden, so dass die komplexe Zahl&#160;\(z_1 \cdot z_2 = {@z1*loes2@} = {@l2@}\) reell ist.
        </span>
        <br/>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            c) Wenn \(z_1={@z1@}\) die vorgegebene komplexe Zahl ist, dann kann z.B. die komplexe Zahl \(z_2 = {@loes3@} \neq 0 \) gew&#228;hlt werden, so dass die beiden komplexen Zahlen \(z_1+z_2 = ({@z1@}) + ({@loes3@}) = {@z1+loes3@}\) und \(z_1 \cdot z_2 = {@z1*loes3@} = {@l3@}\) reell sind.
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <strong>
            Solution idea:
        </strong>
    </p>
    <p>
        One
        <strong>
            possible solution
        </strong>
        the task would be the following.
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            a) If \(z_1={@z1@}\) the predetermined complex number, then, for example, the complex number \(z_2 = {@loes1@} \neq 0 \) are chosen so that the complex number \(z_1+z_2 = ({@z1@}) + ({@loes1@}) = {@z1+loes1@}\) reell.
        </span>
        <br/>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            b) If \(z_1={@z1@}\) the predetermined complex number, then, for example, the complex number \(z_2 = {@loes2@} \neq 0 \) are chosen so that the complex number\(z_1 \cdot z_2 = {@z1*loes2@} = {@l2@}\) reell.
        </span>
        <br/>
    </p>
    <p>
        <span style="font-size: 0.9375rem;">
            c) If \(z_1={@z1@}\) the predetermined complex number, then, for example, the complex number \(z_2 = {@loes3@} \neq 0 \) are chosen so that the two complex numbers \(z_1+z_2 = ({@z1@}) + ({@loes3@}) = {@z1+loes3@}\) and \(z_1 \cdot z_2 = {@z1*loes3@} = {@l3@}\) are real.
        </span>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>a: rand_with_step(1,9,1);
b: rand_with_prohib(-9,9,[0]);
z1: a * %i + b;
loes1: -a * %i;
loes2: a * %i - b;
loes3: -a * %i + b;

l2: rectform(z1*loes2);
l3: rectform(z1*loes3);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>Gegeben ist die komplexe Zahl \( \, z_1={@z1@} \, \).</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Richtige Antwort.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Right answer.
    </p>
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Antwort ist teilweise richtig.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The answer is partly correct.
    </p>
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Antwort ist nicht richtig.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The answer is not correct.
    </p>
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans> loes1</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans> loes2</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans> loes3</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>3.4000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>imagpart(z1+ans1)</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>3.3000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>imagpart(z1*ans2)</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>3.3000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>imagpart(z1+ans3)</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>imagpart(z1*ans3)</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>L&#246;sungsidee</text>
            </tag>
        </tags>
    </question>
    

</quiz>