<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Determinante 3x3</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]Bestimmen Sie die Determinante der Matrix.
<p>
    <br/>
</p>
<p>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p style="">
    </p>
    <p>
        \( A= \) {@A1@}
    </p>
    <p>
    </p>
</div>
<p style="">
    <br/>
</p>
<p>
    \( \det(A)= \) [[/lang]][[lang code="en"]]Determine the determinant of the matrix.
    <p>
        <br/>
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p style="">
        </p>
        <p>
            \( A= \) {@A1@}
        </p>
        <p>
        </p>
    </div>
    <p style="">
        <br/>
    </p>
    <p>
        \( \det(A)= \) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        Musterl&#246;sung:
    </b>
</p>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    <span>
        \(\det(A)=\det\left({@A1@}\right)\)
    </span>
</p>
<p>
    <span>
        \(=({@a@}-9-12)-({@b@}-1+36)\)
    </span>
</p>
<p>
    <span>
        (Regel von Sarrus, gilt nur f&#252;r 3x3 Matrizen)
    </span>
</p>
<p>
    <span>
        \(={@c@}\)
    </span>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        Sample solution:
    </b>
</p>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    <span>
        \(\det(A)=\det\left({@A1@}\right)\)
    </span>
</p>
<p>
    <span>
        \(=({@a@}-9-12)-({@b@}-1+36)\)
    </span>
</p>
<p>
    <span>
        (Regulation of Sarrus, applies only to 3x3 matrices)
    </span>
</p>
<p>
    <span>
        \(={@c@}\)
    </span>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>A:rand_with_step(-10,3,1);
A1:matrix([-2,3,-1],[2,1,-3],[1,6,A]);

a:-2*A
b: 6*A
c: (a-9-12)-(b-1+36)</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>1</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>(-8)*(A)-56</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>(-8)*(A)-56</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Determinante</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>