<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Eigenwerte (reell)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Bestimmen Sie f&#252;r die Matrix
</p>
<p>
    <br/>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
    </p>
    <p>
    </p>
    <p>
        \({@(A)@} =\) {@A1@}
    </p>
    <p>
    </p>
    <p>
    </p>
</div>
<p>
</p>
<p>
</p>
<p>
    <br/>
</p>
<p>
    <b>
        (a)
    </b>
    die Eigenwerte:
</p>
<p>
    Geben Sie die Eigenwerte in aufsteigender Reihenfolge (\(\lambda_1&lt;\lambda_2&lt;\lambda_3\)) an.
</p>
<p>
    <br/>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <div>
        <p>
            \(\lambda_1 =\) [[/lang]][[lang code="en"]]
            <p>
                Determine for the matrix
            </p>
            <p>
                <br/>
            </p>
            <div class="editor-indent" style="margin-left: 30px;">
                <p>
                </p>
                <p>
                </p>
                <p>
                    \({@(A)@} =\) {@A1@}
                </p>
                <p>
                </p>
                <p>
                </p>
            </div>
            <p>
            </p>
            <p>
            </p>
            <p>
                <br/>
            </p>
            <p>
                <b>
                    (a)
                </b>
                the own values:
            </p>
            <p>
                Specify the values in ascending order (\(\lambda_1&lt;\lambda_2&lt;\lambda_3\))
            </p>
            <p>
                <br/>
            </p>
            <div class="editor-indent" style="margin-left: 30px;">
                <div>
                    <p>
                        \(\lambda_1 =\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(\lambda_2 =\) [[/lang]][[lang code="en"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(\lambda_2 =\) [[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt2]][[lang code="de"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(\lambda_3 =\) [[/lang]][[lang code="en"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(\lambda_3 =\) [[/lang]][[input:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt3]][[lang code="de"]]
                    </p>
                </div>
            </div>
            <p>
                <br/>
            </p>
            <b>
                (b)
            </b>
            Geben Sie Komponenten des normierten Eigenvektors \( \frac{\vec{v}}{||\vec{v}||} \) zum Eigenwert 1 an:
            <p>
                <nobr>
                    <span role="math">
                        <br/>
                    </span>
                </nobr>
            </p>
            <div class="editor-indent" style="margin-left: 30px;">
                <div>
                    <p>
                        <span id="MathJax-Element-2-Frame">
                            <nobr>
                                <span id="MathJax-Span-56" role="math">
                                    <span id="MathJax-Span-57">
                                        <span id="MathJax-Span-58">
                                            <span id="MathJax-Span-59">
                                                <span id="MathJax-Span-60">
                                                </span>
                                            </span>
                                        </span>
                                    </span>
                                </span>
                            </nobr>
                        </span>
                        \(v_1 =\) [[/lang]][[lang code="en"]]
                    </p>
                </div>
            </div>
            <p>
                <br/>
            </p>
            <b>
                (b)
            </b>
            Specify components of the normalized eigenvector \( \frac{\vec{v}}{||\vec{v}||} \) for own value 1
            <p>
                <nobr>
                    <span role="math">
                        <br/>
                    </span>
                </nobr>
            </p>
            <div class="editor-indent" style="margin-left: 30px;">
                <div>
                    <p>
                        <span id="MathJax-Element-2-Frame">
                            <nobr>
                                <span id="MathJax-Span-56" role="math">
                                    <span id="MathJax-Span-57">
                                        <span id="MathJax-Span-58">
                                            <span id="MathJax-Span-59">
                                                <span id="MathJax-Span-60">
                                                </span>
                                            </span>
                                        </span>
                                    </span>
                                </span>
                            </nobr>
                        </span>
                        \(v_1 =\) [[/lang]][[input:ans4]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans4]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt4]][[lang code="de"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(v_2 =\) [[/lang]][[lang code="en"]]
                    </p>
                </div>
                <div>
                    <p>
                        \(v_2 =\) [[/lang]][[input:ans5]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans5]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt5]][[lang code="de"]]
                    </p>
                </div>
                <div>
                    <p>
                        <span>
                            \(v_3 =\) [[/lang]][[lang code="en"]]
                        </span>
                    </p>
                </div>
                <div>
                    <p>
                        <span>
                            \(v_3 =\) [[/lang]][[input:ans6]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans6]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt6]][[lang code="de"]]
                        </span>
                    </p>
                </div>
            </div>
            <p>
            </p>
            <p>
            </p>
            [[/lang]][[lang code="en"]]
        </p>
    </div>
</div>
<p>
</p>
<p>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <b>
        L&#246;sung
    </b>
</h4>
<p>
    <br/>
</p>
<p>
    <span>
    </span>
</p>
<p style="">
    <b>
        (a)
    </b>
    Es muss gelten: \(p(\lambda)=\det(A- \lambda E)=0\)
</p>
<p style="">
    <br/>
</p>
<p style="">
    \(p(\lambda)=\det(A- \lambda E)\)
</p>
<p style="">
    \(=\det\left({@a1@}\right)\)
    <br/>
</p>
<p style="">
    \(=(-3-\lambda)\cdot \det\left({@a2@}\right)\)
</p>
<p style="">
    (Laplace Entwicklung nach der dritten Zeile)
</p>
<p style="">
    \(=(-3-\lambda)(2-\lambda)(1-\lambda)\)
</p>
<p style="">
    <br/>
</p>
<p style="">
    Damit der Ausdruck \(0\) wird, muss entweder 1., 2. oder 3. gelten (Linearfaktoren):
</p>
<p style="">
</p>
<ol style="">
    <li>
        \((-3-\lambda)=0\)&#160; &#160; &#160;\(\Leftrightarrow\)&#160; &#160; &#160; \(\lambda_1=-3\)
    </li>
    <li>
        \((1-\lambda)=0\)&#160; &#160; &#160; &#160; &#160;\(\Leftrightarrow\)&#160; &#160; &#160; \(\lambda_2=1\)
    </li>
    <li>
        \((2-\lambda)=0\)&#160; &#160; &#160; &#160; &#160;\(\Leftrightarrow\)&#160; &#160; &#160; \(\lambda_3=2\)
    </li>
</ol>
<br/>
<p>
</p>
<p>
    <span>
        <b>
            (b)
        </b>
        Eigenvektor zu \(\lambda_2=1\):
    </span>
</p>
<p>
    <span>
    </span>
    Berechne \(\ker(A-\lambda_2 E)\)&#160; &#160;(Das hei&#223;t: \(A-\lambda_2 E=O\))
</p>
<p>
    <br/>
</p>
<p>
    \(A-E={@a3@}\xrightarrow{II-I}{@a5@}\)
</p>
<p>
    \(\xrightarrow{ \frac{II}{-8},\frac{III}{-4}}{@a6@}\xrightarrow{I-5II, III-II}{@a7@}\)
</p>
<p>
    <br/>
</p>
<p>
    Da der \(A-\lambda E=O\) sein muss, folgt \(v={@a8@}\) (oder jedes Vielfache).
</p>
<p>
    Der normierte Eigenvektor ist definiert als \(\frac{v}{||v||}\).
</p>
<p>
    Da hier \(||v||=1\) gilt, ist \(v\) auch der normierte Eigenvektor.
</p>
<p>
    <br/>
</p>
<p>
    <br/>
</p>
[[/lang]][[lang code="en"]]
<h4>
    <b>
        Solution
    </b>
</h4>
<p>
    <br/>
</p>
<p>
    <span>
    </span>
</p>
<p style="">
    <b>
        (a)
    </b>
    It must apply: \(p(\lambda)=\det(A- \lambda E)=0\)
</p>
<p style="">
    <br/>
</p>
<p style="">
    \(p(\lambda)=\det(A- \lambda E)\)
</p>
<p style="">
    \(=\det\left({@a1@}\right)\)
    <br/>
</p>
<p style="">
    \(=(-3-\lambda)\cdot \det\left({@a2@}\right)\)
</p>
<p style="">
    (Laplace development after third line)
</p>
<p style="">
    \(=(-3-\lambda)(2-\lambda)(1-\lambda)\)
</p>
<p style="">
    <br/>
</p>
<p style="">
    So the expression \(0\) shall apply either 1., 2. or 3. (linear factors):
</p>
<p style="">
</p>
<ol style="">
    <li>
        \((-3-\lambda)=0\)\(\Leftrightarrow\) \(\lambda_1=-3\)
    </li>
    <li>
        \((1-\lambda)=0\)\(\Leftrightarrow\) \(\lambda_2=1\)
    </li>
    <li>
        \((2-\lambda)=0\)\(\Leftrightarrow\) \(\lambda_3=2\)
    </li>
</ol>
<br/>
<p>
</p>
<p>
    <span>
        <b>
            (b)
        </b>
        Own vector \(\lambda_2=1\):
    </span>
</p>
<p>
    <span>
    </span>
    Calculation \(\ker(A-\lambda_2 E)\)(That means: \(A-\lambda_2 E=O\))
</p>
<p>
    <br/>
</p>
<p>
    \(A-E={@a3@}\xrightarrow{II-I}{@a5@}\)
</p>
<p>
    \(\xrightarrow{ \frac{II}{-8},\frac{III}{-4}}{@a6@}\xrightarrow{I-5II, III-II}{@a7@}\)
</p>
<p>
    <br/>
</p>
<p>
    There \(A-\lambda E=O\) must be followed \(v={@a8@}\) (or multiple).
</p>
<p>
    The normalized eigenvector is defined as \(\frac{v}{||v||}\).
</p>
<p>
    Here. \(||v||=1\) is applicable, \(v\) also the normalized eigenvector.
</p>
<p>
    <br/>
</p>
<p>
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>A1:matrix([2,0,5],[1,1,-3],[0,0,-3]);
A2:[vals, vecs] : eigenvectors (A1);
A3:vals[1];
E1:A3[1];
E2:A3[2];
E3:A3[3];
A4:vecs[2];
A6:A4[1];
nE1:A6[1];
nE2:A6[2];
nE3:A6[3];

a1: matrix([2-lambda,0,5],[1,1-lambda,-3],[0,0,-3-lambda])
a2: matrix([2-lambda,0],[1, 1-lambda])
a3: matrix([2-1,0,5],[1,1-1,-3],[0,0,-3-1])
a4: matrix([1,0,5],[1,0,-3],[0,0,-4])
a5: matrix([1,0,5],[0,0,-8],[0,0,-4])
a6: matrix([1,0,5],[0,0,1],[0,0,1])
a7: matrix([1,0,0],[0,0,1],[0,0,0])
a8: matrix([0],[1],[0])</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>E1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>E2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>E3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>nE1</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>nE2</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>nE3</tans>
            <boxsize>10</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>E1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>E2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>E3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>nE1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>nE2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>nE3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Eigenvektoren</text>
            </tag>
            <tag>
                <text>TE:3:Eigenwerte</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>


</quiz>