<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Matrizenrechnung</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Gegeben seien die Matrizen
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
    </p>
    \( A={@A1@} , B={@B1@} \)&#160;sowie der Vektor&#160; \( c = {@C1@} \).
    <p>
        <br/>
    </p>
    <p>
        Bestimmen Sie
        <br/>
    </p>
    <p>
        <b>
            (a)
        </b>
        \(A\cdot B =\)
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
        </p>
        <p>
            [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
        <p>
            <br/>
        </p>
        <p>
        </p>
    </div>
    <p>
    </p>
    <p>
        <b>
            (b)
        </b>
        die L&#246;sung des linearen Gleichungssystems
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \( A\cdot x=c \)
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                <br/>
            </p>
            <p>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        \( x_1=\)&#160;[[input:ans2]][[validation:ans2]][[feedback:prt2]]
    </p>
    <p>
        \( x_2= \)&#160;[[input:ans3]][[validation:ans3]][[feedback:prt3]]
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        the matrices
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
    </p>
    \( A={@A1@} , B={@B1@} \)and the vector \( c = {@C1@} \).
    <p>
        <br/>
    </p>
    <p>
        Determine
        <br/>
    </p>
    <p>
        <b>
            (a)
        </b>
        \(A\cdot B =\)
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
        </p>
        <p>
            [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
        </p>
        <p>
            <br/>
        </p>
        <p>
        </p>
    </div>
    <p>
    </p>
    <p>
        <b>
            (b)
        </b>
        the solution of the linear equation system
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \( A\cdot x=c \)
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                <br/>
            </p>
            <p>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        \( x_1=\)[[input:ans2]][[validation:ans2]][[feedback:prt2]]
    </p>
    <p>
        \( x_2= \)[[input:ans3]][[validation:ans3]][[feedback:prt3]]
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <h4>
        <b>
            L&#246;sung
        </b>
    </h4>
    <p>
        <b>
            <br/>
        </b>
    </p>
    <p>
        <span>
            <b style="font-weight: bold;">
                (a)
            </b>
            \(A\cdot B = {@A1@} \cdot {@B1@}={@erg1@}\)
            <br/>
        </span>
    </p>
    <p>
        <span>
            <br/>
        </span>
    </p>
    <p>
        <span>
        </span>
    </p>
    <p>
        <b>
            (b)
        </b>
        die L&#246;sung des linearen Gleichungssystems
    </p>
    <div>
        <p>
            \( A\cdot x=c\)
        </p>
        <p>
            <br/>
        </p>
        <p>
            Mit Gauss:
        </p>
        <p>
            \(\left[\begin{array}{cc|c} 2 &amp; -3 &amp; {@a1@} \\ 1 &amp; 4 &amp; 2 \end{array}\right] \xrightarrow{\frac{I}{2}} \left[ \begin{array}{cc|c}1 &amp; -\frac{3}{2} &amp; {@c@} \\ 1 &amp; 4 &amp; 2 \end{array} \right] \xrightarrow{II-I} \left[ \begin{array}{cc|c} 1 &amp; -\frac{3}{2} &amp; {@c@} \\ 0 &amp; \frac{11}{2} &amp; {@a@} \end{array}\right]\)
        </p>
        <p>
            <br/>
        </p>
        <p>
            <b>
                II:
            </b>
            \(\frac{11}{2} x_{2}={@a@}\)&#160; &#160; &#160; &#160; &#160;\(\Leftrightarrow\)&#160; &#160;\(x_{2}={@b@}\)
        </p>
        <p>
            <b>
                I:
            </b>
            \(x_1 - \frac{3}{2}x_2={@c@}\)&#160; &#160; &#160; &#160; ,\(x_2\) einsetzen:
        </p>
        <p>
            \(x_{1}+{@d@}={@c@}\)&#160; &#160; &#160;\(\Leftrightarrow\)&#160; &#160;\(x_{1}={@e@}\)
        </p>
    </div>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <h4>
        <b>
            Solution
        </b>
    </h4>
    <p>
        <b>
            <br/>
        </b>
    </p>
    <p>
        <span>
            <b style="font-weight: bold;">
                (a)
            </b>
            \(A\cdot B = {@A1@} \cdot {@B1@}={@erg1@}\)
            <br/>
        </span>
    </p>
    <p>
        <span>
            <br/>
        </span>
    </p>
    <p>
        <span>
        </span>
    </p>
    <p>
        <b>
            (b)
        </b>
        the solution of the linear equation system
    </p>
    <div>
        <p>
            \( A\cdot x=c\)
        </p>
        <p>
            <br/>
        </p>
        <p>
            With Gauss:
        </p>
        <p>
            \(\left[\begin{array}{cc|c} 2 &amp; -3 &amp; {@a1@} \\ 1 &amp; 4 &amp; 2 \end{array}\right] \xrightarrow{\frac{I}{2}} \left[ \begin{array}{cc|c}1 &amp; -\frac{3}{2} &amp; {@c@} \\ 1 &amp; 4 &amp; 2 \end{array} \right] \xrightarrow{II-I} \left[ \begin{array}{cc|c} 1 &amp; -\frac{3}{2} &amp; {@c@} \\ 0 &amp; \frac{11}{2} &amp; {@a@} \end{array}\right]\)
        </p>
        <p>
            <br/>
        </p>
        <p>
            <b>
                II:
            </b>
            \(\frac{11}{2} x_{2}={@a@}\)\(\Leftrightarrow\)\(x_{2}={@b@}\)
        </p>
        <p>
            <b>
                I:
            </b>
            \(x_1 - \frac{3}{2}x_2={@c@}\),\(x_2\) use:
        </p>
        <p>
            \(x_{1}+{@d@}={@c@}\)\(\Leftrightarrow\)\(x_{1}={@e@}\)
        </p>
    </div>
    <p>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2019052700</text>
        </stackversion>
        <questionvariables>
            <text>a1:2+rand(50)
b1:2+rand(50)
A1:matrix([2,-3],[1,4]);
B1:matrix([b1,1,2],[6,0,1]);
C1:matrix([a1],[2]);

erg1:A1.B1;

A: matrix([2,-3,a1],[1,4,2])
B: matrix([1,(-3/2),(a1/2)],[1,4,2])
C: matrix([1,(-3/2),(a1/2)],[0,(11/2),2-(a1/2)])
a: 2-a1/2
b: 4/11-a1/11
c: a1/2
d: -6/11+3*a1/22
e:c-d</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>matrix</type>
            <tans>erg1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>6/11+4/11*a1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>4/11-1/11*a1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>erg1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>F1:6/11+4/11*a1</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans2</sans>
                <tans>F1</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>F2:4/11-1/11*a1</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans3</sans>
                <tans>F2</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Multiplikation</text>
            </tag>
            <tag>
                <text>te:3:L&#246;sungsmenge</text>
            </tag>
            <tag>
                <text>TE:3:Lineare Gleichungssysteme</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>