<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Determinante a</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p style="">
    Bestimmen Sie den Parameter \(a\), so&#160;dass die Determinante der Matrix \(A\)&#160;Null ergibt.
</p>
<p style="">
    <br/>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \( {@(A)@} =\)&#160; {@A1@}
    </p>
    <p>
        <br/>
    </p>
</div>
<p>
    \( a= \) [[/lang]][[lang code="en"]]
    <p style="">
        Define the parameter \(a\)so that the determinant of the matrix \(A\)Zero.
    </p>
    <p style="">
        <br/>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \( {@(A)@} =\) {@A1@}
        </p>
        <p>
            <br/>
        </p>
    </div>
    <p>
        \( a= \) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <b>
        L&#246;sung
    </b>
</h4>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    <span>
        \(\det(A)=\det\left({@A1@}\right)\)
    </span>
</p>
<p>
    <span>
        \(=(-2+({@a1@}))-(-2+{@a2@})\)
    </span>
</p>
<p>
    <span>
        (Regel von Sarrus)
    </span>
</p>
<p>
    <span>
        \(={@a3@}\)
    </span>
</p>
<p>
    <span>
        \(=0\)
    </span>
</p>
<p>
    <span>
        (Nach Aufgabenstellung)
    </span>
</p>
<p>
    <span>
        <br/>
    </span>
</p>
<p>
    <span>
        \(\Leftrightarrow \ \ a={@a4@}\)
    </span>
</p>
[[/lang]][[lang code="en"]]
<h4>
    <b>
        Solution
    </b>
</h4>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    <span>
        \(\det(A)=\det\left({@A1@}\right)\)
    </span>
</p>
<p>
    <span>
        \(=(-2+({@a1@}))-(-2+{@a2@})\)
    </span>
</p>
<p>
    <span>
        (Rule of Sarrus)
    </span>
</p>
<p>
    <span>
        \(={@a3@}\)
    </span>
</p>
<p>
    <span>
        \(=0\)
    </span>
</p>
<p>
    <span>
        (According to task)
    </span>
</p>
<p>
    <span>
        <br/>
    </span>
</p>
<p>
    <span>
        \(\Leftrightarrow \ \ a={@a4@}\)
    </span>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>A2:rand([-5,-4,-3,-2,-1,1,2,3,4,5])

A1:matrix([1,0,-1],[a,1,A2],[2,A2,-2]);

a1: -A2*a
a2: A2^2
a3: (-2+a1)-(-2+a2)
a4: -A2</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>(-1)*(A2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>(-1)*(A2)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Gleichung</text>
            </tag>
            <tag>
                <text>TE:3:Determinante</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>