<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Inverse</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <strong>
        <span style="font-size:20px;">
            <strong>
            </strong>
        </span>
    </strong>
    <span style="font-size: 20px;">
    </span>
    <span>
        <span style="font-size: 16px;">
        </span>
    </span>
    <p style="">
        Berechnen Sie die Inverse der folgenden Matrizen.
    </p>
    <b>
        a)
    </b>
    \(A = \begin{bmatrix}
1 &amp; {@a@} &amp; 0  \\
{@a@}&amp; 1 &amp; {@a@}  \\
0 &amp; {@a@} &amp; 1  \\
\end{bmatrix} \)
    <p>
        <br data-cke-eol="1"/>
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[input:ans1]]&#160; [[input:ans2]]&#160; [[input:ans3]]
            </p>
        </div>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(A^{-1} =\) [[input:ans4]]&#160; [[input:ans5]]&#160; [[input:ans6]]
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[input:ans7]]&#160; [[input:ans8]]&#160; [[input:ans9]]
            </p>
            <p>
                [[validation:ans1]][[validation:ans2]][[validation:ans3]]
            </p>
            <p>
                [[validation:ans4]][[validation:ans5]][[validation:ans6]]
            </p>
            <p>
                [[validation:ans7]][[validation:ans8]][[validation:ans9]]
            </p>
            <p>
            </p>
            <p>
                [[feedback:prt1]]&#160;[[feedback:prt2]][[feedback:prt3]]
                <br/>
            </p>
            <p>
                [[feedback:prt4]][[feedback:prt5]][[feedback:prt6]]
                <br/>
            </p>
            <p>
                [[feedback:prt7]][[feedback:prt8]][[feedback:prt9]]
            </p>
            <br/>
            <p>
            </p>
            <p>
            </p>
            <p>
                <br/>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        <span>
            <span>
                <b>
                    b)
                </b>
                \(B = \begin{bmatrix} \dfrac{1}{{@b@}}  &amp; - \dfrac{{@bplus1@}}{{@b@}} &amp; 1 \\ 0&amp; \dfrac{{@bplus1@}}{{@b@}}&amp;-1\\ 0 &amp; -1 &amp; 1 \\ \end{bmatrix} \)
            </span>
        </span>
    </p>
    <p>
        <span>
            <span>
                <br/>
            </span>
        </span>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
        </p>
        <div>
            <div>
                <p>
                    [[input:ans10]]&#160; [[input:ans11]]&#160; [[input:ans12]]
                </p>
            </div>
        </div>
        <p>
        </p>
        <p>
        </p>
        <div>
            <span>
                \(B^{-1} =\) [[input:ans13]]&#160; [[input:ans14]]&#160; [[input:ans15]]
            </span>
        </div>
        <p>
        </p>
        <p>
        </p>
        <div>
            <div>
                <p>
                    [[input:ans16]]&#160; [[input:ans17]]&#160; [[input:ans18]]
                </p>
            </div>
        </div>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans10]][[validation:ans11]][[validation:ans12]]
            </p>
        </div>
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans13]][[validation:ans14]][[validation:ans15]]
            </p>
        </div>
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans16]][[validation:ans17]][[validation:ans18]]
            </p>
            <p>
            </p>
            <p>
                [[feedback:prt10]][[feedback:prt11]][[feedback:prt12]]
                <br/>
            </p>
            <p>
                [[feedback:prt13]][[feedback:prt14]][[feedback:prt15]]
                <br/>
            </p>
            <p>
                [[feedback:prt16]][[feedback:prt17]][[feedback:prt18]]
            </p>
            <br/>
            <p>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        <b>
            Hinweis:
        </b>
        Geben Sie die Werte gegebenfalls in Br&#252;chen an.
    </p>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <strong>
        <span style="font-size:20px;">
            <strong>
            </strong>
        </span>
    </strong>
    <span style="font-size: 20px;">
    </span>
    <span>
        <span style="font-size: 16px;">
        </span>
    </span>
    <p style="">
        Calculate the inverse of the following matrices.
    </p>
    <b>
        (a)
    </b>
    \(A = \begin{bmatrix}
1 &amp; {@a@} &amp; 0  \\
{@a@}&amp; 1 &amp; {@a@}  \\
0 &amp; {@a@} &amp; 1  \\
\end{bmatrix} \)
    <p>
        <br data-cke-eol="1"/>
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[input:ans1]] [[input:ans2]] [[input:ans3]]
            </p>
        </div>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(A^{-1} =\) [[input:ans4]] [[input:ans5]] [[input:ans6]]
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[input:ans7]] [[input:ans8]] [[input:ans9]]
            </p>
            <p>
                [[validation:ans1]][[validation:ans2]][[validation:ans3]]
            </p>
            <p>
                [[validation:ans4]][[validation:ans5]][[validation:ans6]]
            </p>
            <p>
                [[validation:ans7]][[validation:ans8]][[validation:ans9]]
            </p>
            <p>
            </p>
            <p>
                [[feedback:prt1]][[feedback:prt2]][[feedback:prt3]]
                <br/>
            </p>
            <p>
                [[feedback:prt4]][[feedback:prt5]][[feedback:prt6]]
                <br/>
            </p>
            <p>
                [[feedback:prt7]][[feedback:prt8]][[feedback:prt9]]
            </p>
            <br/>
            <p>
            </p>
            <p>
            </p>
            <p>
                <br/>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        <span>
            <span>
                <b>
                    (b)
                </b>
                \(B = \begin{bmatrix} \dfrac{1}{{@b@}}  &amp; - \dfrac{{@bplus1@}}{{@b@}} &amp; 1 \\ 0&amp; \dfrac{{@bplus1@}}{{@b@}}&amp;-1\\ 0 &amp; -1 &amp; 1 \\ \end{bmatrix} \)
            </span>
        </span>
    </p>
    <p>
        <span>
            <span>
                <br/>
            </span>
        </span>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
        </p>
        <div>
            <div>
                <p>
                    [[input:ans10]] [[input:ans11]] [[input:ans12]]
                </p>
            </div>
        </div>
        <p>
        </p>
        <p>
        </p>
        <div>
            <span>
                \(B^{-1} =\) [[input:ans13]] [[input:ans14]] [[input:ans15]]
            </span>
        </div>
        <p>
        </p>
        <p>
        </p>
        <div>
            <div>
                <p>
                    [[input:ans16]] [[input:ans17]] [[input:ans18]]
                </p>
            </div>
        </div>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans10]][[validation:ans11]][[validation:ans12]]
            </p>
        </div>
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans13]][[validation:ans14]][[validation:ans15]]
            </p>
        </div>
        <div class="editor-indent" style="margin-left: 30px;">
            <p>
                [[validation:ans16]][[validation:ans17]][[validation:ans18]]
            </p>
            <p>
            </p>
            <p>
                [[feedback:prt10]][[feedback:prt11]][[feedback:prt12]]
                <br/>
            </p>
            <p>
                [[feedback:prt13]][[feedback:prt14]][[feedback:prt15]]
                <br/>
            </p>
            <p>
                [[feedback:prt16]][[feedback:prt17]][[feedback:prt18]]
            </p>
            <br/>
            <p>
            </p>
        </div>
    </div>
    <p>
    </p>
    <p>
        <b>
            Note:
        </b>
        Specify the values in fractures.
    </p>
    <p>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <b>
            Musterl&#246;sung:
        </b>
    </p>
    <p>
        <br/>
    </p>
    <p>
        <b>
            a)
        </b>
        \(A = \begin{bmatrix} 1 &amp; {@a@} &amp; 0 \\ {@a@}&amp; 1 &amp; {@a@} \\ 0 &amp; {@a@} &amp; 1 \\ \end{bmatrix} \)
        <br/>
    </p>
    <p>
        Mit Gauss:
    </p>
    <p>
        <br/>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ {@a@} &amp; 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; {@a@} &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{II-{@a@}I} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; {@a1@} &amp; {@a@} &amp; {@a2@} &amp; 1 &amp; 0 \\ 0 &amp; {@a@} &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{II/({@a1@}), III/({@a@})}\)
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 1 &amp; {@a5@} &amp; 0 &amp; 0 &amp; {@a5@} \end{array}\right] \xrightarrow{III-II} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 0 &amp; {@a6@} &amp; {@a3@} &amp; {@a4@} &amp; {@a5@} \end{array}\right] \xrightarrow{III \cdot ({@a11@})}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\xrightarrow{II - ({@a3@} \cdot III)} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; 0 &amp; {@a12@} &amp; {@a13@} &amp; {@a14@} \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right] \xrightarrow{I- {@a@}II}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@a15@} &amp; {@a16@} &amp; {@a17@} \\ 0 &amp; 1 &amp; 0 &amp; {@a12@} &amp; {@a13@} &amp; {@a14@} \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\)
        </p>
        <p>
            <br/>
        </p>
    </div>
    <span style="">
        Das hei&#223;t:
    </span>
    <p>
        <span style="">
            \( A^{-1}=\left[ \begin{array}{ccc} {@a15@} &amp; {@a16@} &amp; {@a17@} \\ {@a12@} &amp; {@a13@} &amp; {@a14@} \\ {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\)
        </span>
    </p>
    <p>
        <span style="">
            <br/>
        </span>
    </p>
    <p>
        <span style="">
            <br/>
        </span>
        <span style="">
            <b>
                b)
            </b>
            \(B = \begin{bmatrix} {@b1@} &amp; - {@b2@} &amp; 1 \\ 0&amp; {@b2@}&amp;-1\\ 0 &amp; -1 &amp; 1 \\ \end{bmatrix} \)
            <br/>
        </span>
    </p>
    <p>
        Mit Gauss:
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} {@b1@} &amp; - {@b2@} &amp; 1 &amp; 1 &amp; 0 &amp; 0 \\ 0&amp; {@b2@}&amp;-1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; -1 &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{I+II} \left[ \begin{array}{ccc|ccc} {@b1@} &amp; 0 &amp; 0 &amp; 1 &amp; 1 &amp; 0 \\ 0&amp; {@b2@}&amp;-1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; -1 &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{{@b@}I, III+{@b3@}II}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; {@b2@} &amp; -1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; 0 &amp; {@b4@} &amp; 0 &amp; {@b3@} &amp;1 \end{array}\right] \xrightarrow{{@b3@}II, {@b7@}III} \left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; 1 &amp; {@b5@} &amp; 0 &amp; {@b3@} &amp; 0 \\ 0 &amp; 0 &amp; 1 &amp; 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right] \xrightarrow{II+{@b3@}III}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; 1 &amp; 0 &amp; 0 &amp; {@b8@} &amp; {@b9@} \\ 0 &amp; 0 &amp; 1 &amp; 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right]\)
        </p>
    </div>
    <span style="">
        <br/>
    </span>
    <p>
    </p>
    <p>
        <span style="">
            Das hei&#223;t:
        </span>
    </p>
    <p>
        <span style="">
            \(B^{-1}=\left[ \begin{array}{ccc} {@b@} &amp; {@b@} &amp; 0 \\ 0 &amp; {@b8@} &amp; {@b9@} \\ 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right]\)
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <b>
            Sample solution:
        </b>
    </p>
    <p>
        <br/>
    </p>
    <p>
        <b>
            (a)
        </b>
        \(A = \begin{bmatrix} 1 &amp; {@a@} &amp; 0 \\ {@a@}&amp; 1 &amp; {@a@} \\ 0 &amp; {@a@} &amp; 1 \\ \end{bmatrix} \)
        <br/>
    </p>
    <p>
        With Gauss:
    </p>
    <p>
        <br/>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ {@a@} &amp; 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; {@a@} &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{II-{@a@}I} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; {@a1@} &amp; {@a@} &amp; {@a2@} &amp; 1 &amp; 0 \\ 0 &amp; {@a@} &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{II/({@a1@}), III/({@a@})}\)
        </p>
    </div>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 1 &amp; {@a5@} &amp; 0 &amp; 0 &amp; {@a5@} \end{array}\right] \xrightarrow{III-II} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 0 &amp; {@a6@} &amp; {@a3@} &amp; {@a4@} &amp; {@a5@} \end{array}\right] \xrightarrow{III \cdot ({@a11@})}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; {@a3@} &amp; {@a4@} &amp; {@a10@} &amp; 0 \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\xrightarrow{II - ({@a3@} \cdot III)} \left[ \begin{array}{ccc|ccc} 1 &amp; {@a@} &amp; 0 &amp; 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; 0 &amp; {@a12@} &amp; {@a13@} &amp; {@a14@} \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right] \xrightarrow{I- {@a@}II}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@a15@} &amp; {@a16@} &amp; {@a17@} \\ 0 &amp; 1 &amp; 0 &amp; {@a12@} &amp; {@a13@} &amp; {@a14@} \\ 0 &amp; 0 &amp; 1 &amp; {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\)
        </p>
        <p>
            <br/>
        </p>
    </div>
    <span style="">
        That means:
    </span>
    <p>
        <span style="">
            \( A^{-1}=\left[ \begin{array}{ccc} {@a15@} &amp; {@a16@} &amp; {@a17@} \\ {@a12@} &amp; {@a13@} &amp; {@a14@} \\ {@a7@} &amp; {@a8@} &amp; {@a9@} \end{array}\right]\)
        </span>
    </p>
    <p>
        <span style="">
            <br/>
        </span>
    </p>
    <p>
        <span style="">
            <br/>
        </span>
        <span style="">
            <b>
                (b)
            </b>
            \(B = \begin{bmatrix} {@b1@} &amp; - {@b2@} &amp; 1 \\ 0&amp; {@b2@}&amp;-1\\ 0 &amp; -1 &amp; 1 \\ \end{bmatrix} \)
            <br/>
        </span>
    </p>
    <p>
        With Gauss:
    </p>
    <p>
    </p>
    <div class="editor-indent" style="margin-left: 30px;">
        <p>
            \(\left[ \begin{array}{ccc|ccc} {@b1@} &amp; - {@b2@} &amp; 1 &amp; 1 &amp; 0 &amp; 0 \\ 0&amp; {@b2@}&amp;-1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; -1 &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{I+II} \left[ \begin{array}{ccc|ccc} {@b1@} &amp; 0 &amp; 0 &amp; 1 &amp; 1 &amp; 0 \\ 0&amp; {@b2@}&amp;-1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; -1 &amp; 1 &amp; 0 &amp; 0 &amp;1 \end{array}\right] \xrightarrow{{@b@}I, III+{@b3@}II}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; {@b2@} &amp; -1 &amp; 0 &amp; 1 &amp; 0 \\ 0 &amp; 0 &amp; {@b4@} &amp; 0 &amp; {@b3@} &amp;1 \end{array}\right] \xrightarrow{{@b3@}II, {@b7@}III} \left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; 1 &amp; {@b5@} &amp; 0 &amp; {@b3@} &amp; 0 \\ 0 &amp; 0 &amp; 1 &amp; 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right] \xrightarrow{II+{@b3@}III}\)
        </p>
        <p>
            \(\left[ \begin{array}{ccc|ccc} 1 &amp; 0 &amp; 0 &amp; {@b@} &amp; {@b@} &amp; 0 \\ 0&amp; 1 &amp; 0 &amp; 0 &amp; {@b8@} &amp; {@b9@} \\ 0 &amp; 0 &amp; 1 &amp; 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right]\)
        </p>
    </div>
    <span style="">
        <br/>
    </span>
    <p>
    </p>
    <p>
        <span style="">
            That means:
        </span>
    </p>
    <p>
        <span style="">
            \(B^{-1}=\left[ \begin{array}{ccc} {@b@} &amp; {@b@} &amp; 0 \\ 0 &amp; {@b8@} &amp; {@b9@} \\ 0 &amp; {@b6@} &amp; {@b7@} \end{array}\right]\)
        </span>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>a:rand_with_step(2,5,1);
b:rand_with_step(1,5,1);
bplus1:b+1;

a1: 1-a^2
a2:-a
a3: a/(1-a^2)
a4: -a3
a5: 1/a
a6: a5-a3
a7:a3/a6
a10: 1/a1
a8: -a10/a6
a9: a5/a6
a11: 1/a6
a12: a4-a7*a3
a13: a10-a3*a8
a14: -a3*a9
a15: 1-a*a12
a16:-a*a13
a17: -a*a14

b1: 1/b
b2: (1+b)/b
b3: b/(1+b)
b4:1-b/(1+b)
b5:-b3
b6: b3/b4
b7: 1/b4
b8:  b3+b3*b6
b9: b7*b3</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>(a^2-1)/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans10</name>
            <type>algebraic</type>
            <tans>b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans11</name>
            <type>algebraic</type>
            <tans>b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans12</name>
            <type>algebraic</type>
            <tans>0</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans13</name>
            <type>algebraic</type>
            <tans>0</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans14</name>
            <type>algebraic</type>
            <tans>b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans15</name>
            <type>algebraic</type>
            <tans>b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans16</name>
            <type>algebraic</type>
            <tans>0</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans17</name>
            <type>algebraic</type>
            <tans>b</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans18</name>
            <type>algebraic</type>
            <tans>b+1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>a/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>-a^2/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>a/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>-1/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>algebraic</type>
            <tans>a/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans7</name>
            <type>algebraic</type>
            <tans>-a^2/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans8</name>
            <type>algebraic</type>
            <tans>a/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans9</name>
            <type>algebraic</type>
            <tans>(a^2-1)/(2*a^2-1)</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>(a^2-1)/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt10</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans10</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt10-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt10-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt11</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans11</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt11-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt11-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt12</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans12</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt12-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt12-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt13</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans13</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt13-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt13-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt14</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans14</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt14-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt14-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt15</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans15</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt15-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt15-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt16</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans16</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt16-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt16-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt17</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans17</sans>
                <tans>b</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt17-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt17-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt18</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans18</sans>
                <tans>b+1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt18-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt18-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>a/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>-a^2/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>a/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>-1/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt6</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>a/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt6-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt6-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt7</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans7</sans>
                <tans>-a^2/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt7-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt7-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt8</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans8</sans>
                <tans>a/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt8-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt8-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt9</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans9</sans>
                <tans>(a^2-1)/(2*a^2-1)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt9-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt9-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Matrizen</text>
            </tag>
            <tag>
                <text>TE:3:Gau&#223;-Algorithmus</text>
            </tag>
            <tag>
                <text>TE:3:Inverse Matrix</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>