<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineares Gleichungssystem: Gau&#223;-Algorithmus</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bestimmen Sie f&#252;r
        <br/>
    </p>
    <p style="">
        \({@A@}={@m@}\) und &#160; \(b={@b@}\)
        <br/>
    </p>
    die L&#246;sung des linearen Gleichungssystems \( A\cdot x=b  \)
    <div class="editor-indent" style="margin-left: 30px;">
        <h5>
            <ul>
                <li>
                    <span style="font-weight: normal;">
                        \({@x1@}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x2@}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x3@}=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x4@}=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
                    </span>
                </li>
            </ul>
        </h5>
    </div>
    <p>
    </p>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Choose for
        <br/>
    </p>
    <p style="">
        \({@A@}={@m@}\) and \(b={@b@}\)
        <br/>
    </p>
    the solution of the linear equation system \( A\cdot x=b  \)
    <div class="editor-indent" style="margin-left: 30px;">
        <h5>
            <ul>
                <li>
                    <span style="font-weight: normal;">
                        \({@x1@}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x2@}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x3@}=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
                    </span>
                </li>
                <li>
                    <span style="font-weight: normal;">
                        \({@x4@}=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
                    </span>
                </li>
            </ul>
        </h5>
    </div>
    <p>
    </p>
    <p>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <h4>
        <strong>
            L&#246;sung
        </strong>
    </h4>
    <p>
        <br/>
    </p>
    <p>
        Mit Gauss:
    </p>
    <p>
        \(\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp;{@A1@}\\1&amp;1&amp;1&amp;1&amp;{@A1+1@}\\8&amp;4&amp;2&amp;1&amp;{@A1+2@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@}\end{array}\right)\xrightarrow{II+I, III+8I}\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp;{@A1@}\\0&amp;2&amp;0&amp;2&amp;{@2*A1+1@}\\0&amp;12&amp;-6&amp;9&amp;{@9*A1+2@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@} \end{array}\right)\xrightarrow{III-6II}\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp; {@A1@}\\0&amp;2&amp;0&amp;2&amp;{@2*A1+1@}\\0&amp;0&amp;-6&amp;-3&amp;{@-3*A1-4@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@}\end{array}\right)\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            IV
        </strong>
        <strong>
            :
        </strong>
        \(x_4={@A1+3@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            III
        </strong>
        <strong>
            :
        </strong>
        \(-6x_3-3x_4={@-3*A1-4@}\) mit IV:
    </p>
    <p>
        \(-6x_3-3\cdot {@A1+3@}={@-3*A1-4@}\)&#160; &#160; \(\Leftrightarrow\)&#160; &#160;\(x_3={@t3@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            II:
        </strong>
        \(2x_2+2x_4={@2*A1+1@}\)&#160; mit IV:
    </p>
    <p>
        \(2x_2+2\cdot {@t4@}={@2*A1+1@}\)&#160; &#160; \(\Leftrightarrow\)&#160; &#160;\(x_2={@t2@}\)
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            I:
        </strong>
        \(-x_1+x_2-x_3+x_4={@A1@}\)&#160; mit II, III, IV:
    </p>
    <p>
        \(-x_1+{@t2@}-{@t3@}+{@t4@}={@A1@}\)&#160; &#160; \(\Leftrightarrow\)&#160; &#160;\(x_1={@t1@}\)
    </p>
    <p>
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <h4>
        <strong>
            Solution
        </strong>
    </h4>
    <p>
        <br/>
    </p>
    <p>
        With Gauss:
    </p>
    <p>
        \(\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp;{@A1@}\\1&amp;1&amp;1&amp;1&amp;{@A1+1@}\\8&amp;4&amp;2&amp;1&amp;{@A1+2@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@}\end{array}\right)\xrightarrow{II+I, III+8I}\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp;{@A1@}\\0&amp;2&amp;0&amp;2&amp;{@2*A1+1@}\\0&amp;12&amp;-6&amp;9&amp;{@9*A1+2@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@} \end{array}\right)\xrightarrow{III-6II}\left(\begin{array}{cccc|c}-1&amp;1&amp;-1&amp;1&amp; {@A1@}\\0&amp;2&amp;0&amp;2&amp;{@2*A1+1@}\\0&amp;0&amp;-6&amp;-3&amp;{@-3*A1-4@}\\0&amp;0&amp;0&amp;1&amp;{@A1+3@}\end{array}\right)\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            IV
        </strong>
        <strong>
            :
        </strong>
        \(x_4={@A1+3@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            III
        </strong>
        <strong>
            :
        </strong>
        \(-6x_3-3x_4={@-3*A1-4@}\) with IV:
    </p>
    <p>
        \(-6x_3-3\cdot {@A1+3@}={@-3*A1-4@}\) \(\Leftrightarrow\)\(x_3={@t3@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            II:
        </strong>
        \(2x_2+2x_4={@2*A1+1@}\)with IV:
    </p>
    <p>
        \(2x_2+2\cdot {@t4@}={@2*A1+1@}\) \(\Leftrightarrow\)\(x_2={@t2@}\)
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
        <strong>
            I:
        </strong>
        \(-x_1+x_2-x_3+x_4={@A1@}\)with II, III, IV:
    </p>
    <p>
        \(-x_1+{@t2@}-{@t3@}+{@t4@}={@A1@}\) \(\Leftrightarrow\)\(x_1={@t1@}\)
    </p>
    <p>
        <br/>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>A1:rand(10);
A2:1+rand(30);
m : matrix( [-1,1,-1,1],[1,1,1,1],[8,4,2,1],[0,0,0,1] )

b : matrix([A1],[A1+1],[A1+2],[A1+3]);

ta:linsolve_by_lu(m,b)
tbb:ta[1];
t1:tbb[1,1];
t2:tbb[2,1];
t3:tbb[3,1];
t4:tbb[4,1];

/* Liste erzeugen */
liste:makelist(tbb);
f1:tbb[1,1];
f2:tbb[2,1];
f3:tbb[3,1];
f4:tbb[4,1];</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>t1</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>t2</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>t3</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>t4</tans>
            <boxsize>4</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans1</sans>
                <tans>t1</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans2</sans>
                <tans>t2</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans3</sans>
                <tans>t3</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumRelative</answertest>
                <sans>ans4</sans>
                <tans>t4</tans>
                <testoptions>0.05</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>TE:2:Lineare Gleichungssysteme</text>
            </tag>
            <tag>
                <text>te:3:L&#246;sungsmenge</text>
            </tag>
            <tag>
                <text>TE:3:Gau&#223;-Algorithmus</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:2</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>