<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Basiswechsel R^3 2</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <span style="font-size: 14px; font-weight: normal;">
        Betrachten Sie den Vektorraum \( \mathbb{R}^3\) mit den Basen
    </span>
</h4>
<h4>
    <span style="font-size: 14px; font-weight: normal;">
        \(\mathcal{B}_1=\{e_1,e_2,e_3\}\).
    </span>
    <br/>
</h4>
<p>
    \(\mathcal{B}_2=\big\{{@uv@},{@vv@},{@wv@}\big\}\)
</p>
<p>
    <br/>
</p>
<p>
    Berechnen Sie die Koordinaten des Vektors
</p>
<p>
</p>
<p>
    \({x} = {@xv@} \)
</p>
<p>
</p>
<p>
    in der Basis \(\mathcal{B}_2\), das hei&#223;t
</p>
<p>
    <br/>
</p>
<p>
    \({x}\,=\,a_1{@uv@}+a_2{@vv@}+a_3{@wv@}\)
</p>
<p>
</p>
<p>
    \(a_1=\) [[/lang]][[lang code="en"]]
    <h4>
        <span style="font-size: 14px; font-weight: normal;">
            View the vector space \( \mathbb{R}^3\) with the bases
        </span>
    </h4>
    <h4>
        <span style="font-size: 14px; font-weight: normal;">
            \(\mathcal{B}_1=\{e_1,e_2,e_3\}\).
        </span>
        <br/>
    </h4>
    <p>
        \(\mathcal{B}_2=\big\{{@uv@},{@vv@},{@wv@}\big\}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        Calculate vector coordinates
    </p>
    <p>
    </p>
    <p>
        \({x} = {@xv@} \)
    </p>
    <p>
    </p>
    <p>
        in the base \(\mathcal{B}_2\)That means
    </p>
    <p>
        <br/>
    </p>
    <p>
        \({x}\,=\,a_1{@uv@}+a_2{@vv@}+a_3{@wv@}\)
    </p>
    <p>
    </p>
    <p>
        \(a_1=\) [[/lang]][[input:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        \(a_2=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        \(a_2=\) [[/lang]][[input:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
        \(a_3=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        \(a_3=\) [[/lang]][[input:ans3]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[validation:ans3]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt3]][[lang code="de"]]
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    [[/lang]][[lang code="en"]]
</p>
<p>
    <br/>
</p>
<p>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <b>
        L&#246;sung
    </b>
</h4>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    Der Vektor&#160;\({x} = {@xv@}\) soll in Koordinaten \(a\) bez&#252;glich der gegebenen Basis&#160; \(\mathcal{B}_2\) umgerechnet werden \(\Big({a}=T^{-1}\cdot {x}\Big)\).
</p>
<p>
    \(T\) stellt die Transformationsmatrix dar, die in den Spalten die gegebenen Basisvektoren hat:
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \(T={@M@}\)
    </p>
    <p>
        <br/>
    </p>
</div>
<p>
    Daraus l&#228;sst sich die Inverse \(T^{-1}={@Minv@}\) berechnen.
</p>
<p>
    <br/>
</p>
<p>
    \({\lambda}=\begin{bmatrix} a_1\\ a_2\\ a_3 \end{bmatrix}=T^{-1}\cdot {x}\)
    <br/>
</p>
<p>
    <br/>
</p>
<p>
    \(={@Minv@} \cdot {@xv@}\)
</p>
<p>
    <br/>
</p>
<p>
    \(=\begin{bmatrix} {@Bx1@}\\{@Bx2@}\\{@Bx3@} \end{bmatrix}\)
</p>
<p>
    <br/>
</p>
<p>
    Koordinaten bez&#252;glich der Basis&#160;\(\mathcal{B}_2\): \(({@Bx1@},{@Bx2@},{@Bx3@})\)
    <br/>
</p>
[[/lang]][[lang code="en"]]
<h4>
    <b>
        Solution
    </b>
</h4>
<p>
    <b>
        <br/>
    </b>
</p>
<p>
    The vector\({x} = {@xv@}\) in coordinates \(a\) with respect to the given base \(\mathcal{B}_2\) conversion \(\Big({a}=T^{-1}\cdot {x}\Big)\).
</p>
<p>
    \(T\) represents the transformation matrix which has the given base vectors in the columns:
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \(T={@M@}\)
    </p>
    <p>
        <br/>
    </p>
</div>
<p>
    The Inverse \(T^{-1}={@Minv@}\) calculate.
</p>
<p>
    <br/>
</p>
<p>
    \({\lambda}=\begin{bmatrix} a_1\\ a_2\\ a_3 \end{bmatrix}=T^{-1}\cdot {x}\)
    <br/>
</p>
<p>
    <br/>
</p>
<p>
    \(={@Minv@} \cdot {@xv@}\)
</p>
<p>
    <br/>
</p>
<p>
    \(=\begin{bmatrix} {@Bx1@}\\{@Bx2@}\\{@Bx3@} \end{bmatrix}\)
</p>
<p>
    <br/>
</p>
<p>
    coordinates with respect to the base\(\mathcal{B}_2\): \(({@Bx1@},{@Bx2@},{@Bx3@})\)
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text>u1: 1+rand(3)
u2: -1-rand(4)
u3: 0
uv: matrix([u1],[u2],[u3]);
v1: 1
v2: 1+rand(3)
v3: 1
vv: matrix([v1],[v2],[v3]);
w1: -1-rand(3)
w2: 0
w3: 1+rand(3)
wv: matrix([w1],[w2],[w3]);

M: matrix([u1,v1,w1], [u2,v2,w2], [u3,v3,w3]);

r: rank(M);


x1: rand_with_step(-9,9,1);
x2: rand_with_step(-9,9,1);
x3: rand_with_step(-9,9,1);
xv: matrix([x1],[x2],[x3]);
Minv: invert(M);

Bx1: (Minv.[x1,x2,x3])[1][1];
Bx2: (Minv.[x1,x2,x3])[2][1];
Bx3: (Minv.[x1,x2,x3])[3][1];</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(Rang={@r@}, \lambda_1={@Bx1@}, \lambda_2={@Bx2@}. \lambda_3={@Bx3@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Richtige Antwort, gut gemacht!
</p>
[[/lang]][[lang code="en"]]
<p>
    Right answer, well done!
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Ihre Antwort ist teilweise korrekt.
</p>
[[/lang]][[lang code="en"]]
<p>
    Your answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Falsche Antwort.
</p>
[[/lang]][[lang code="en"]]
<p>
    Wrong answer.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>none</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>Bx1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>Bx2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>Bx3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>loes: [Bx1,Bx2,Bx3];
antwort: [ans1,ans2,ans3];</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>Bx1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>Bx2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>Bx3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Bloom:3</text>
            </tag>
            <tag>
                <text>Maier:1:2</text>
            </tag>
            <tag>
                <text>Maier:2:3</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>te:2:Vektorraum</text>
            </tag>
            <tag>
                <text>TE:3:Basiswechsel</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>