<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Dimension</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Welche Dimension hat der Untervektorraum des \(\mathbb{R}^4\), der von den folgenden Vektoren erzeugt wird?
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
        \(v = \begin{bmatrix} {@a@} \\ -1\\ 2 \\ 3\\ \end{bmatrix}, w = \begin{bmatrix} {0} \\ {0}\\ {2}\\-1 \end{bmatrix}, s = \begin{bmatrix} {0} \\ 2\\ {4} \\ {1} \end{bmatrix}, t = \begin{bmatrix} {@minus2@} \\ {8}\\ {6} \\ {-2} \end{bmatrix} \)
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
        Die Dimension betr&#228;gt [[input:ans1]]. [[validation:ans1]][[feedback:prt1]]
    </p>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        What dimension has the subvector space of the \(\mathbb{R}^4\)which is generated by the following vectors?
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
        \(v = \begin{bmatrix} {@a@} \\ -1\\ 2 \\ 3\\ \end{bmatrix}, w = \begin{bmatrix} {0} \\ {0}\\ {2}\\-1 \end{bmatrix}, s = \begin{bmatrix} {0} \\ 2\\ {4} \\ {1} \end{bmatrix}, t = \begin{bmatrix} {@minus2@} \\ {8}\\ {6} \\ {-2} \end{bmatrix} \)
        <br/>
    </p>
    <p>
        <br/>
    </p>
    <p>
    </p>
    <p>
        The dimension is [[input:ans1]]. [[validation:ans1]][[feedback:prt1]]
    </p>
    <p>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <h4>
        L&#246;sung
    </h4>
    <p>
        <br/>
    </p>
    <p>
        Die Dimension entspricht der Anzahl der linear unabh&#228;ngigen Vektoren.
    </p>
    <p>
        Das hei&#223;t: Die Dimension ist gleich dem Rang der Matrix, die \(v, \ w, \ s, \ t\) als Spalten besitzt.
    </p>
    <p>
        Umformung der Matrix in Zeilenstufenform:
    </p>
    <p>
        <br/>
    </p>
    <p>
        \({@A@}\xrightarrow{ \frac{I}{{@a@}},\frac{III}{2}}{@B@}\xrightarrow{II+I, III-I, IV-3I}{@C@}\)
    </p>
    <p>
        \(\xrightarrow{IV+III}{@D@}\xrightarrow{\frac{II}{2},\frac{IV}{3}, II \leftrightarrow III}{@E@}\xrightarrow{IV-III}{@F@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        Der Bildraum hat somit die Dimension 3.
    </p>
</span>
<span class="multilang" lang="en">
    <h4>
        Solution
    </h4>
    <p>
        <br/>
    </p>
    <p>
        The dimension corresponds to the number of linear independent vectors.
    </p>
    <p>
        That means: The dimension is equal to the rank of the matrix, \(v, \ w, \ s, \ t\) as a column.
    </p>
    <p>
        Forming the matrix in line-level form:
    </p>
    <p>
        <br/>
    </p>
    <p>
        \({@A@}\xrightarrow{ \frac{I}{{@a@}},\frac{III}{2}}{@B@}\xrightarrow{II+I, III-I, IV-3I}{@C@}\)
    </p>
    <p>
        \(\xrightarrow{IV+III}{@D@}\xrightarrow{\frac{II}{2},\frac{IV}{3}, II \leftrightarrow III}{@E@}\xrightarrow{IV-III}{@F@}\)
    </p>
    <p>
        <br/>
    </p>
    <p>
        The image space thus has the dimension 3.
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>a:rand_with_step(1,5,1);
minus2:-2*a;
A: matrix([a,0,0,minus2],[-1,0,2,8],[2,2,4,6],[3,-1,1,-2])
B: matrix([1,0,0,-2],[-1,0,2,8],[1,1,2,3],[3,-1,1,-2])
C: matrix([1,0,0,-2],[0,0,2,6],[0,1,2,5],[0,-1,1,4])
D: matrix([1,0,0,-2],[0,0,2,6],[0,1,2,5],[0,0,3,9])
E: matrix([1,0,0,-2],[0,1,2,5], [0,0,1,3], [0,0,1,3])
F: matrix([1,0,0,-2],[0,1,2,5], [0,0,1,3], [0,0,0,0])</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>te:2:Vektorraum</text>
            </tag>
            <tag>
                <text>TE:3:Dimension</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:2</text>
            </tag>
            <tag>
                <text>Maier:1:3</text>
            </tag>
            <tag>
                <text>Maier:2:1</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>