<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lineare Algebra: Lineare Abh&#228;ngigkeit (MC)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Sind die angegebenen Vektoren linear abh&#228;ngig voneinander?
    <span style="font-size: 20px;">
        <b>
            <br/>
        </b>
    </span>
</p>
<p>
</p>
<p>
</p>
<p>
    <b>
        (a)
    </b>
    \(\begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@aplus1@} \\ {@bminus1@}\\ {@c@}\\ \end{bmatrix} , \begin{bmatrix} {@aplus2@} \\ {@b@}\\ {@cplus2@}\\ \end{bmatrix} \)
    <br/>
</p>
<p>
    [[/lang]][[lang code="en"]]
    <p>
        Are the specified vectors linearly dependent on each other?
        <span style="font-size: 20px;">
            <b>
                <br/>
            </b>
        </span>
    </p>
    <p>
    </p>
    <p>
    </p>
    <p>
        <b>
            (a)
        </b>
        \(\begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@aplus1@} \\ {@bminus1@}\\ {@c@}\\ \end{bmatrix} , \begin{bmatrix} {@aplus2@} \\ {@b@}\\ {@cplus2@}\\ \end{bmatrix} \)
        <br/>
    </p>
    <p>
        [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]]
        <br/>
    </p>
    <p>
        <b>
            (b)
        </b>
        \( \begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@d@} \\ {@e@}\\ {@f@}\\ \end{bmatrix} , \begin{bmatrix} {@aminus2@} \\ {@bminus2@}\\ {@cminus2@}\\ \end{bmatrix} \)
        <br/>
    </p>
    <p>
        [[/lang]][[lang code="en"]]
        <br/>
    </p>
    <p>
        <b>
            (b)
        </b>
        \( \begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@d@} \\ {@e@}\\ {@f@}\\ \end{bmatrix} , \begin{bmatrix} {@aminus2@} \\ {@bminus2@}\\ {@cminus2@}\\ \end{bmatrix} \)
        <br/>
    </p>
    <p>
        [[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
    </p>
    <p>
    </p>
    [[/lang]][[lang code="en"]]
</p>
<p>
</p>
<p>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<h4>
    <b>
        L&#246;sung
    </b>
</h4>
<p>
    <br/>
</p>
<p>
</p>
<p>
    Es wird eine Umformung der Matrix mit den Vektoren als Spalten in Zeilenstufenform durchgef&#252;hrt.
</p>
<p>
    Ist der Rang voll (also keine Nullzeile am Ende), sind die Vektoren linear unabh&#228;ngig, anderenfalls sind sie linear abh&#228;ngig.
</p>
<br/>
<p>
</p>
<p>
    <b>
        <b>
            (a)
        </b>
    </b>
    \(\begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@aplus1@} \\ {@bminus1@}\\ {@c@}\\ \end{bmatrix} , \begin{bmatrix} {@aplus2@} \\ {@b@}\\ {@cplus2@}\\ \end{bmatrix} \)
    <br/>
</p>
<p>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \({@A@} \xrightarrow{{@a@}II-{@b@}I, {@a@}III- {@c@}I} {@B@} \xrightarrow{{@a1@}III+({@a2@})II} {@C@}\)
    </p>
</div>
<p>
    <br/>
</p>
<p>
    <b>
        (b)
    </b>
    \( \begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@d@} \\ {@e@}\\ {@f@}\\ \end{bmatrix} , \begin{bmatrix} {@aminus2@} \\ {@bminus2@}\\ {@cminus2@}\\ \end{bmatrix} \)
    <br/>
</p>
<p>
    <br/>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \({@D@} \xrightarrow{{@a@}II-{@b@}I, {@a@}III- {@c@}I} {@E@} \xrightarrow{{@b1@}III+({@b2@})II} {@F@}\)
    </p>
</div>
[[/lang]][[lang code="en"]]
<h4>
    <b>
        Solution
    </b>
</h4>
<p>
    <br/>
</p>
<p>
</p>
<p>
    A transformation of the matrix with the vectors is carried out as columns in row step form.
</p>
<p>
    If the rank is full (i.e. no zero line at the end), the vectors are linearly independent, otherwise they are linearly dependent.
</p>
<br/>
<p>
</p>
<p>
    <b>
        <b>
            (a)
        </b>
    </b>
    \(\begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@aplus1@} \\ {@bminus1@}\\ {@c@}\\ \end{bmatrix} , \begin{bmatrix} {@aplus2@} \\ {@b@}\\ {@cplus2@}\\ \end{bmatrix} \)
    <br/>
</p>
<p>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \({@A@} \xrightarrow{{@a@}II-{@b@}I, {@a@}III- {@c@}I} {@B@} \xrightarrow{{@a1@}III+({@a2@})II} {@C@}\)
    </p>
</div>
<p>
    <br/>
</p>
<p>
    <b>
        (b)
    </b>
    \( \begin{bmatrix} {@a@} \\ {@b@}\\ {@c@}\\ \end{bmatrix} ,\begin{bmatrix} {@d@} \\ {@e@}\\ {@f@}\\ \end{bmatrix} , \begin{bmatrix} {@aminus2@} \\ {@bminus2@}\\ {@cminus2@}\\ \end{bmatrix} \)
    <br/>
</p>
<p>
    <br/>
</p>
<div class="editor-indent" style="margin-left: 30px;">
    <p>
        \({@D@} \xrightarrow{{@a@}II-{@b@}I, {@a@}III- {@c@}I} {@E@} \xrightarrow{{@b1@}III+({@b2@})II} {@F@}\)
    </p>
</div>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a:rand_with_step(1,5,1);
aplus1:a+1;
aplus2:a+2;
b:rand_with_step(1,5,1);
bminus1:b-1;
c: rand_with_prohib(1,5,[a+b]);
cplus2:c+2;
d:rand_with_step(1,5,1);
e:rand_with_step(1,5,1);
f:rand_with_step(1,5,1);
aminus2:a-2*d;
bminus2:b-2*e;
cminus2: rand_with_prohib(c-2*f,c-2*f,[a,b]);
mm:matrix([1,3,-1],[0,2,1],[0,-1,2])


tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
Tans2 : 2;
AnsList : [[1, false, tr("ja","yes")], [2, true, tr("nein","no")]];


Tans3 : 1;
AnsList2 : [[1, true, tr("ja","yes")], [2, false, tr("nein","no")]];

A: matrix([a,a+1,a+2],[b,b-1,b],[c,c,c+2])
B: matrix([a,a+1,a+2],[0,a*(b-1)-b*(a+1),a*b-b*(a+2)], [0, a*c-c*(a+1),a*(c+2)-c*(a+2)])
C: matrix([a,a+1,a+2],[0,a*(b-1)-b*(a+1),a*b-b*(a+2)], [0,0 ,(a*(b-1)-b*(a+1))*(a*(c+2)-c*(a+2))-(a*c-c*(a+1))*(a*b-b*(a+2))])
a1: (a*(b-1)-b*(a+1))
a2: (a*c-c*(a+1))

D: matrix([a,d,aminus2], [b,e,bminus2], [c,f,cminus2])
E: matrix([a,d,aminus2], [0, a*e-b*d, a*bminus2-b*aminus2], [0, a*f-c*d, a*cminus2-c*aminus2])
F: matrix([a,d,aminus2], [0, a*e-b*d, a*bminus2-b*aminus2], [0, 0, (a*e-b*d)*(a*cminus2-c*aminus2)-( a*f-c*d)*( a*bminus2-b*aminus2)])
b1:a*e-b*d
b2: a*f-c*d]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>n</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>radio</type>
            <tans>AnsList </tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>radio</type>
            <tans>AnsList2 </tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>Tans2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>Tans3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>TE:1:Lineare Algebra</text>
            </tag>
            <tag>
                <text>te:2:Vektorraum</text>
            </tag>
            <tag>
                <text>TE:3:Lineare Unabh&#228;ngigkeit</text>
            </tag>
            <tag>
                <text>src:F07-INT</text>
            </tag>
            <tag>
                <text>Bloom:2</text>
            </tag>
            <tag>
                <text>Maier:1:3</text>
            </tag>
            <tag>
                <text>Maier:2:1</text>
            </tag>
            <tag>
                <text>Maier:3:1</text>
            </tag>
            <tag>
                <text>Maier:4:1</text>
            </tag>
            <tag>
                <text>Maier:5:1</text>
            </tag>
            <tag>
                <text>Maier:6:1</text>
            </tag>
            <tag>
                <text>Maier:7:1</text>
            </tag>
            <tag>
                <text>praxiserprobt</text>
            </tag>
            <tag>
                <text>qs:visuellGeprueft</text>
            </tag>
            <tag>
                <text>qs:sprachlichGeprueft</text>
            </tag>
            <tag>
                <text>qs:TaggingGeprueft</text>
            </tag>
            <tag>
                <text>qs:importIliasGeprueft</text>
            </tag>
            <tag>
                <text>qs:importMoodleGeprueft</text>
            </tag>
        </tags>
    </question>
    

</quiz>