<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Linearisierten Fehler [{"DOMAINUID":"265F185"}] [{"DOMAINUID":"265F185"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Gegeben sind die Volumen \(V\) und die Oberfl&#228;che \(O\) eines K&#246;rpers:
        <br/>
        <span style="font-size: 0.9375rem;">
            \( V = \displaystyle\frac{1}{3} \cdot \pi \cdot r^2 \cdot&#160;h\)
        </span>
        <br/>
    </p>
    <p>
        \(O = \pi \cdot&#160;r \cdot&#160;\sqrt{r^2+h^2} + \pi \cdot&#160;r^2 \)
    </p>
    <p>
        Mit welchen absoluten linearisierten Fehlern muss man bei der Ermittlung des Volumens und der Oberfl&#228;che rechnen, wenn der Radius \( r = {@radius@}\,cm\) und die H&#246;he \(h = {@hoehe@}\,cm\) mit einer Genauigkeit von \( \pm{@fehler@}\,cm \) gemessen wurden?
    </p>
    <p>
        <strong>
            Hinweise:
        </strong>
        Benutzen Sie keine Dezimahlzahlen. Schreiben Sie pi f&#252;r \(\pi\).
    </p>
    <p>
        \( \Delta V = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        \( \Delta O = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
    <p>
        Relative Fehler von \(V =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </p>
    <p>
        Relative Fehler von \(O =\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The volume \(V\) and the surface \(O\) a body:
        <br/>
        <span style="font-size: 0.9375rem;">
            \( V = \displaystyle\frac{1}{3} \cdot \pi \cdot r^2 \cdot&#160;h\)
        </span>
        <br/>
    </p>
    <p>
        \(O = \pi \cdot&#160;r \cdot&#160;\sqrt{r^2+h^2} + \pi \cdot&#160;r^2 \)
    </p>
    <p>
        What absolute linearized errors must be expected when determining the volume and surface if the radius \( r = {@radius@}\,cm\) and the height \(h = {@hoehe@}\,cm\) with an accuracy of \( \pm{@fehler@}\,cm \) were measured?
    </p>
    <p>
        <strong>
            Notes:
        </strong>
        Don't use decirculum numbers. Write pi for \(\pi\).
    </p>
    <p>
        \( \Delta V = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </p>
    <p>
        \( \Delta O = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
    </p>
    <p>
        Relative errors of \(V =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
    </p>
    <p>
        Relative errors of \(O =\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>radius : rand_with_step(90, 270, 30);
fehler : rand_with_step(1, 3, 1);
hoehe : radius/3*4;


V : 1/3 * %pi * radius^2 * hoehe;

O: %pi*radius*sqrt(radius^2+hoehe^2)+%pi*radius^2;

deltav : (2/3*radius*hoehe*fehler + 1/3*radius^2*fehler)*%pi;

deltao : ((sqrt(radius^2+hoehe^2) + radius^2/sqrt(radius^2+hoehe^2) + 2*radius)*fehler + radius*hoehe/sqrt(radius^2+hoehe^2)*fehler)*%pi;

relv : deltav / V;

relo : deltao / O;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\[ V = \displaystyle\frac{1}{3} \pi r^2 h \quad,\quad  O = \pi r \sqrt{r^2+h^2} + \pi r^2 \]
\[{@deltav@}, {@deltao@}, {@relv@}, {@relo@} \]</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Richtig
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Right
    </p>
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Teilweise richtig
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Partly correct
    </p>
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Antwort ist nicht richtig
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The answer is not correct
    </p>
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>deltav</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>deltao</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>relv</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>relo</tans>
            <boxsize>40</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>deltav</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>deltao</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>relv</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>relo</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-0-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-0-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>