<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Mengen: Symmetrische Differenz  [{"DOMAINUID":"266DCB3"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die symmetrische Differenz zweier Mengen \(A\) und \(B\) ist definiert als \(A\,\triangle\, B=(A\setminus B)\cup(B\setminus A)\)
    <br/>
</p>
<p>
    Seien \( A=\{{@a@},{@a*2@},{@a*3@},\ldots\} \) und \( B=\{{@b@},{@b*2@},{@b*3@},\ldots\} \) gegeben.
</p>
<p>
    Bestimmen Sie eine mindestens \(3\)-elementige Menge \( C\subseteq A\cup B \), so dass
</p>
<p>
    \( C = \) [[/lang]][[lang code="en"]]
    <p>
        The symmetrical difference of two sets \(A\) and \(B\) is defined as \(A\,\triangle\, B=(A\setminus B)\cup(B\setminus A)\)
        <br/>
    </p>
    <p>
        Be \( A=\{{@a@},{@a*2@},{@a*3@},\ldots\} \) and \( B=\{{@b@},{@b*2@},{@b*3@},\ldots\} \) given.
    </p>
    <p>
        Determine at least one \(3\)-elementary set \( C\subseteq A\cup B \)so that
    </p>
    <p>
        \( C = \) [[/lang]][[input:ans1]][[lang code="de"]] \( \subseteq A\cap B \) [[/lang]][[lang code="en"]] \( \subseteq A\cap B \) [[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt1]][[lang code="de"]]
    </p>
    <p>
        Bestimmen Sie eine mindestens \(3\)-elementige Menge \( D\subseteq A\cup B \), so dass
    </p>
    <p>
        \( D = \) [[/lang]][[lang code="en"]]
    </p>
    <p>
        Determine at least one \(3\)-elementary set \( D\subseteq A\cup B \)so that
    </p>
    <p>
        \( D = \) [[/lang]][[input:ans2]][[lang code="de"]] \( \subseteq A\,\triangle\, B \) [[/lang]][[lang code="en"]] \( \subseteq A\,\triangle\, B \) [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[feedback:prt2]][[lang code="de"]]
    </p>
    <p>
        <strong>
            Hinweis:
            <br/>
        </strong>
    </p>
    <ul>
        <li>
            Schreiben Sie die Menge mit geschweiften Klammern. Trennen Sie die L&#246;sungen mit Komma, z.B. {1,3}.
        </li>
        <li>
            Schreiben Sie die leere Menge als { }.
        </li>
    </ul>
    [[/lang]][[lang code="en"]]
</p>
<p>
    <strong>
        Note:
        <br/>
    </strong>
</p>
<ul>
    <li>
        Write the set with curly brackets. Separate the solutions with comma, e.g. {1,3}.
    </li>
    <li>
        Write the empty set as { }.
    </li>
</ul>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        L&#246;sungsweg:
    </b>
</p>
<p>
</p>
<ul>
    <li>
        \( C \subseteq A\cap B \) bedeutet, dass die Elemente von \(C\) in den Mengen \(A\)
        <u>
            und
        </u>
        \(B\) enthalten sein m&#252;ssen. Da \(A\) offenbar alle Vielfachen von {@a@} und \(B\) alle Vielfachen von {@b@} enth&#228;lt, gilt: F&#252;r alle \(k\in \mathbb{N}\) liegt \(k\cdot {@a@} \cdot {@b@}\) im Schnitt von \(A\) und \(B\). Ein Beispiel daf&#252;r ist die Menge \(&#160; \{ 1\cdot {@a@} \cdot {@b@},2\cdot {@a@} \cdot {@b@}, 3\cdot {@a@} \cdot {@b@} \} ={@set(a*b,a*b*2,a*b*3)@}\).
    </li>
</ul>
<p>
</p>
<p>
</p>
<ul>
    <li>
        Zeichnet man f&#252;r die Menge \(A\,\triangle\, B =(A\setminus B)\cup(B\setminus A)\) ein Venn-Diagramm, sieht man schnell, dass \(A\,\triangle\, B = (A\cup B)\setminus (A\cap B)\) gilt. Dies bedeutet, dass man hier einfach nur Elemente der Vereinigung \(A\cup B\) angeben muss, die nicht im Durchschnitt \(A\cap B\) liegen. Ein Beispiel daf&#252;r ist die Menge \({@set(a,b,13*b)@}\).
    </li>
</ul>
<p>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        Solution:
    </b>
</p>
<p>
</p>
<ul>
    <li>
        \( C \subseteq A\cap B \) means that the elements of \(C\) in sets \(A\)
        <u>
            and
        </u>
        \(B\) must be included. There. \(A\) apparently all multiples of {@a@} and \(B\) all multiples of {@b@} for all: \(k\in \mathbb{N}\) lies \(k\cdot {@a@} \cdot {@b@}\) in section \(A\) and \(B\). An example is the set \(&#160; \{ 1\cdot {@a@} \cdot {@b@},2\cdot {@a@} \cdot {@b@}, 3\cdot {@a@} \cdot {@b@} \} ={@set(a*b,a*b*2,a*b*3)@}\).
    </li>
</ul>
<p>
</p>
<p>
</p>
<ul>
    <li>
        Signs for set \(A\,\triangle\, B =(A\setminus B)\cup(B\setminus A)\) a Venn chart, you can quickly see that \(A\,\triangle\, B = (A\cup B)\setminus (A\cap B)\) Article 2 This means that here you simply have elements of union \(A\cup B\) not on average \(A\cap B\) lie. An example is the set \({@set(a,b,13*b)@}\).
    </li>
</ul>
<p>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>liste:[2,2,3,5,7];
a1:rand(liste);liste:delete(a1,liste);
a2:rand(liste);liste:delete(a2,liste);
a3:rand(liste);liste:delete(a3,liste);
a:a1*a2*a3;
b1:a1;
liste:append(liste,[11]);
b2:rand(liste);
b:b1*b2;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<br/>
<br/>
[[/lang]][[lang code="en"]]
<br/>
<br/>
[[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@a1@}</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>{a*b,a*b*2,a*b*3}</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>{a,b,13*b}</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[liste1:listify(ans1);
rest_a:sum(map(lambda([x],mod(x,a)),liste1)[i],i,1,length(liste1));
falsche_a:setify(sublist(liste1,lambda([x],mod(x,a)>0)));
rest_b:sum(map(lambda([x],mod(x,b)),liste1)[i],i,1,length(liste1));
falsche_b:setify(sublist(liste1,lambda([x],mod(x,b)>0)));]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>rest_a</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Zahlen aus der Menge \({@falsche_a@}\) sind keine Elemente von \(A\) und daher auch keine Elemente von \(A\cap B\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The numbers from the set \({@falsche_a@}\) are not elements of \(A\) and therefore no elements of \(A\cap B\).
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>rest_b</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Zahlen aus der Menge \({@falsche_b@}\) sind keine Elemente von \(B\) und daher auch keine Elemente von \(A\cap B\).
</p>
[[/lang]][[lang code="en"]]
<p>
    The numbers from the set \({@falsche_b@}\) are not elements of \(B\) and therefore no elements of \(A\cap B\).
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[liste2:listify(ans2);
nicht_in_A_oder_B:setify(sublist(liste2,lambda([x],mod(x,a)>0 and mod(x,b)>0)));
reste_A_und_B:sum(map(lambda([x],mod(x,a)*mod(x,b)),liste2)[i],i,1,length(liste2));

in_A_und_B:setify(sublist(liste2,lambda([x],mod(x,a)=0 and mod(x,b)=0)));
p_in_A_und_B:emptyp(in_A_und_B);
reste_A_oder_B:sum(map(lambda([x],mod(x,a)+mod(x,b)),liste2)[i],i,1,length(liste2));]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>reste_A_und_B</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Zahlen aus der Menge \({@nicht_in_A_oder_B@}\) sind weder Elemente von A noch von B und daher auch keine Elemente von \(A\, \triangle\, B\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The numbers from the set \({@nicht_in_A_oder_B@}\) are neither elements A nor B and therefore no elements of \(A\, \triangle\, B\).
    <br/>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>p_in_A_und_B</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Die Zahlen aus der Menge \({@in_A_und_B@}\) sind Elemente von \( A \cap B \) und daher keine Elemente von \(A\,\triangle\, B\).
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    The numbers from the set \({@in_A_und_B@}\) are elements of \( A \cap B \) and therefore no elements of \(A\,\triangle\, B\).
    <br/>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>+</falsescoremode>
                <falsescore>0.5000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
            <tag>
                <text>R&#252;ckmeldebaum</text>
            </tag>
        </tags>
    </question>
    

</quiz>