<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Mengen angeben um Ungleichung zu beweisen [{"DOMAINUID":"3D729C2"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Zeigen Sie, dass f&#252;r drei Mengen \(A, B, C\) im Allgemeinen \[(A\cap B)\cup C \ne (A\cap C)\cup B\] gilt, indem Sie drei Mengen \(A, B, C\) angeben, bei denen diese Ungleichung erf&#252;llt ist.
    <br/>
</p>
<p>
    \(A=\) [[/lang]][[lang code="en"]]
    <p>
        Show that for three sets \(A, B, C\) in general \[(A\cap B)\cup C \ne (A\cap C)\cup B\] applies by providing three sets \(A, B, C\) for which this imbalance is fulfilled.
        <br/>
    </p>
    <p>
        \(A=\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]]
    </p>
    <p>
        \(B=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        \(B=\) [[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]]
    </p>
    <p>
        \(C=\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        \(C=\) [[/lang]][[input:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]]
    </p>
    <p>
        <span>
            <span style="">
                <strong style="">
                    Hinweis:
                </strong>
                <span style="">
                    Eine Menge \(\{a,b,c\}\) k&#246;nnen Sie mit
                </span>
            </span>
            <code style="">
                {a,b,c}
            </code>
            eingeben.
        </span>
    </p>
    [[/lang]][[lang code="en"]]
</p>
<p>
    <span>
        <span style="">
            <strong style="">
                Note:
            </strong>
            <span style="">
                A set \(\{a,b,c\}\) can be entered as
            </span>
        </span>
        <code style="">
            {a,b,c}
        </code>
        .
    </span>
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        L&#246;sungsidee:
    </b>
</p>
<p>
    Ein Beispiel f&#252;r drei Mengen, an denen deutlich wird, dass im Allgemeinen \[(A\cap B)\cup C \ne (A\cap C)\cup B\] gilt, sind die Mengen \(A={@MusterA@}\), \( B={@MusterB@}\) und \(C={@MusterC@}\), denn:
    <br/>
</p>
<ul>
    <li>
        \((A\cap B)\cup C = {@intersection(MusterA,MusterB)@}\cup {@MusterC@} = {@union( intersection(MusterA, MusterB), MusterC )@}\)
    </li>
    <li>
        \((A\cap C)\cup B = \emptyset\cup {@MusterB@} = {@union( intersection(MusterA, MusterC), MusterB )@}\)
    </li>
</ul>
<p>
    Und diese beiden Mengen sind unterschiedlich.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        Solution idea:
    </b>
</p>
<p>
    An example of three sets where it becomes clear that in general \[(A\cap B)\cup C \ne (A\cap C)\cup B\] applies: the sets \(A={@MusterA@}\), \( B={@MusterB@}\) and \(C={@MusterC@}\)for:
    <br/>
</p>
<ul>
    <li>
        \((A\cap B)\cup C = {@intersection(MusterA,MusterB)@}\cup {@MusterC@} = {@union( intersection(MusterA, MusterB), MusterC )@}\)
    </li>
    <li>
        \((A\cap C)\cup B = \emptyset\cup {@MusterB@} = {@union( intersection(MusterA, MusterC), MusterB )@}\)
    </li>
</ul>
<p>
    And these two sets are different.
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>/*Erstellt von Jonas Lache, Ruhr-Universit&#228;t Bochum*/
MusterA: set(1,2);
MusterB: set(1,3);
MusterC: set(4);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text/>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>MusterA</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>MusterB</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>MusterC</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>LHS: union( intersection(ans1, ans2), ans3 );
RHS: union( intersection(ans1, ans3), ans2 );</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>LHS</sans>
                <tans>RHS</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    F&#252;r Ihre Eingabe gilt
    <span style="font-size: 0.9375rem;">
        \( (A\cap B)\cup C={@LHS@} \) und
    </span>
    <span style="font-size: 0.9375rem;">
        \( (A\cap C)\cup B={@RHS@}\)
    </span>
    <span style="font-size: 0.9375rem;">
        und somit \((A\cap B)\cup C = (A\cap C)\cup B\), wof&#252;r Sie eigentlich ein Gegenbeispiel finden sollten.
    </span>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your entry is valid
    <span style="font-size: 0.9375rem;">
        \( (A\cap B)\cup C={@LHS@} \) and
    </span>
    <span style="font-size: 0.9375rem;">
        \( (A\cap C)\cup B={@RHS@}\)
    </span>
    <span style="font-size: 0.9375rem;">
        and thus \((A\cap B)\cup C = (A\cap C)\cup B\), for which you should actually find an example.
    </span>
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>1.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    F&#252;r Ihre Eingabe gilt
    <span style="font-size: 0.9375rem;">
        \((A\cap B)\cup C={@LHS@}\) und
    </span>
    <span style="font-size: 0.9375rem;">
        \((A\cap C)\cup B={@RHS@}\)
    </span>
    <span style="font-size: 0.9375rem;">
        und somit \((A\cap B)\cup C \ne (A\cap C)\cup B\), sehr gut!
    </span>
</p>
[[/lang]][[lang code="en"]]
<p>
    Your entry is valid
    <span style="font-size: 0.9375rem;">
        \((A\cap B)\cup C={@LHS@}\) and
    </span>
    <span style="font-size: 0.9375rem;">
        \((A\cap C)\cup B={@RHS@}\)
    </span>
    <span style="font-size: 0.9375rem;">
        and thus \((A\cap B)\cup C \ne (A\cap C)\cup B\)Very good!
    </span>
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>R&#252;ckmeldebaum</text>
            </tag>
            <tag>
                <text>L&#246;sungsidee</text>
            </tag>
        </tags>
    </question>
    

</quiz>