<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Mengenoperationen durchf&#252;hren [{"DOMAINUID":"1E7DF00"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Gegeben sind die Mengen \(A={@MA@}\), \(B={@MB@}\) und \(C={@MC@}\).
    <br/>
    Geben Sie die folgenden Mengen an.
    <br/>
</p>
<p>
</p>
<p>
    <strong>
        Hinweis:
    </strong>
</p>
<ul>
    <li>
        Schreiben Sie die Menge mit geschweiften Klammern. Trennen Sie die L&#246;sungen mit Komma, z.B. {1,3}.
    </li>
    <li>
        Schreiben Sie die leere Menge als { }.
    </li>
</ul>
<p>
    a) \( A \setminus (B\cup C) =\) [[/lang]][[lang code="en"]]
    <p>
        The sets \(A={@MA@}\), \(B={@MB@}\) and \(C={@MC@}\).
        <br/>
        Specify the following sets.
        <br/>
    </p>
    <p>
    </p>
    <p>
        <strong>
            Note:
        </strong>
    </p>
    <ul>
        <li>
            Write the set with curly brackets. Separate the solutions with comma, e.g. {1,3}.
        </li>
        <li>
            Write the empty set as { }.
        </li>
    </ul>
    <p>
        (a) \( A \setminus (B\cup C) =\) [[/lang]][[input:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans1]][[lang code="de"]]
    </p>
    <p>
        b) \( (B\cup C) \setminus A =\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        (b) \( (B\cup C) \setminus A =\) [[/lang]][[input:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]]
    </p>
    <p>
        c) \( (B\cup C) \cap A =\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        (c) \( (B\cup C) \cap A =\) [[/lang]][[input:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]]
    </p>
    <p>
        d) \( (A\cap B) \setminus C =\) [[/lang]][[lang code="en"]]
    </p>
    <p>
        (d) \( (A\cap B) \setminus C =\) [[/lang]][[input:ans4]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans4]][[lang code="de"]]
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <b>
        Die Aufgabe kann wie folgt gel&#246;st werden:
    </b>
</p>
<p>
    Bei allen Aufgaben ist es sinnvoll, Schritt f&#252;r Schritt vorzugehen.
</p>
<p>
    a) Zun&#228;chst bestimmen wir \(B\cup C\). Es gilt \(B\cup C={@union(MB, MC)@}\).
    <br/>
    Jetzt bestimmen wir noch \( A \setminus (B\cup C) = {@MA@} \setminus {@union(MB, MC)@} = {@TA1@}\)
</p>
<p>
    b) \(B\cup C\) haben wir schon in Aufgabe a) bestimmt.
    <br/>
    Jetzt bestimmen wir noch \( (B\cup C) \setminus A ={@union(MB, MC)@} \setminus&#160;{@MA@}\).
    <br/>
    Hier kommt die Menge \({@TA2@}\) heraus.
</p>
<p>
    c) Auch hier k&#246;nnen wir wieder \(B\cup C\) aus a) verwenden.
    <br/>
    Jetzt bestimmen wir noch \( (B\cup C) \cap A = {@union(MB, MC)@} \cap {@MA@}\).
    <br/>
    Heraus kommt die Menge \({@TA3@}\).
</p>
<p>
    d) Hier bestimmen wir zun&#228;chst \( A\cap B \), das Ergebnis ist {@es(intersection(MA, MB))@}.
    <br/>
    Vereinigt man diese Menge nun noch mit&#160; \(C={@MC@}\), kommt als Ergebnis \( (A\cap B) \setminus C = {@TA4@}\) heraus.
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
    <b>
        The object can be achieved as follows:
    </b>
</p>
<p>
    In all tasks it is appropriate to proceed step by step.
</p>
<p>
    a) First of all we determine \(B\cup C\). The following shall apply: \(B\cup C={@union(MB, MC)@}\).
    <br/>
    Let's go \( A \setminus (B\cup C) = {@MA@} \setminus {@union(MB, MC)@} = {@TA1@}\)
</p>
<p>
    (b) \(B\cup C\) we have already determined in task a).
    <br/>
    Let's go \( (B\cup C) \setminus A ={@union(MB, MC)@} \setminus&#160;{@MA@}\).
    <br/>
    Here comes the set \({@TA2@}\) out.
</p>
<p>
    (c) Here we can again \(B\cup C\) (a) use.
    <br/>
    Let's go \( (B\cup C) \cap A = {@union(MB, MC)@} \cap {@MA@}\).
    <br/>
    Out comes the set \({@TA3@}\).
</p>
<p>
    (d) Here we determine \( A\cap B \), the result is {@es(intersection(MA, MB))@}.
    <br/>
    This set is now combined with \(C={@MC@}\), comes as a result \( (A\cap B) \setminus C = {@TA4@}\) out.
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*Erstellt von Jonas Lache, Ruhr-Universit&#228;t Bochum*/

a: rand([3,5,7]);

MA: setify(makelist(a*[1,2,3,4,5][k],k,1,5));
MB: setify(makelist(2*a*[1,2,3,4][k],k,1,4));
MC: setify(makelist(6*a*[1,2,3,4][k],k,1,4));

TA1: setdifference(MA, union(MB, MC) );
TA2: setdifference( union(MB, MC), MA);
TA3: intersection( union(MB, MC), MA);
TA4: setdifference( intersection(MA, MB), MC);

/*Funktion, die wenn eine Menge leer ist, nicht {}, sondern \emptyset zur&#252;ck gibt*/
es(M):= if M=setify([]) then "\(\emptyset\)" else M;]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(A={@MA@}\), \(B={@MB@}\) und \(C={@MC@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>TA1</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>TA2</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>TA3</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>TA4</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>TA1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>TA2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>TA3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>TA4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>Musterl&#246;sung</text>
            </tag>
        </tags>
    </question>
    

</quiz>