<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Mengen angeben um Ungleichung zu beweisen [{"DOMAINUID":"3D729C2"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Zeigen Sie, dass f&#252;r drei Mengen \(A, B, C\) im Allgemeinen \[(A\cap B)\cup C \ne (A\cap C)\cup B\] gilt, indem Sie drei Mengen \(A, B, C\) angeben, bei denen diese Ungleichung erf&#252;llt ist.
        <br/>
    </p>
    <p>
        \(A=\) [[input:ans1]] [[validation:ans1]]
    </p>
    <p>
        \(B=\) [[input:ans2]] [[validation:ans2]]
    </p>
    <p>
        \(C=\) [[input:ans3]] [[validation:ans3]]
    </p>
    <p>
        <span>
            <span style="">
                <strong style="">
                    Hinweis:
                </strong>
                <span style="">
                    Eine Menge \(\{a,b,c\}\) k&#246;nnen Sie mit
                </span>
            </span>
            <code style="">
                {a,b,c}
            </code>
            eingeben.
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Show that for three quantities \(A, B, C\) in general \[(A\cap B)\cup C \ne (A\cap C)\cup B\] by applying three quantities \(A, B, C\) indicate that this imbalance is fulfilled.
        <br/>
    </p>
    <p>
        \(A=\) [[input:ans1]] [[validation:ans1]]
    </p>
    <p>
        \(B=\) [[input:ans2]] [[validation:ans2]]
    </p>
    <p>
        \(C=\) [[input:ans3]] [[validation:ans3]]
    </p>
    <p>
        <span>
            <span style="">
                <strong style="">
                    Note:
                </strong>
                <span style="">
                    A lot \(\{a,b,c\}\) you can
                </span>
            </span>
            <code style="">
                {a,b,c}
            </code>
            enter.
        </span>
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <b>
            L&#246;sungsidee:
        </b>
    </p>
    <p>
        Ein Beispiel f&#252;r drei Mengen, an denen deutlich wird, dass im Allgemeinen \[(A\cap B)\cup C \ne (A\cap C)\cup B\] gilt, sind die Mengen \(A={@MusterA@}\), \( B={@MusterB@}\) und \(C={@MusterC@}\), denn:
        <br/>
    </p>
    <ul>
        <li>
            \((A\cap B)\cup C = {@intersection(MusterA,MusterB)@}\cup {@MusterC@} = {@union( intersection(MusterA, MusterB), MusterC )@}\)
        </li>
        <li>
            \((A\cap C)\cup B = \emptyset\cup {@MusterB@} = {@union( intersection(MusterA, MusterC), MusterB )@}\)
        </li>
    </ul>
    <p>
        Und diese beiden Mengen sind unterschiedlich.
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <b>
            Solution idea:
        </b>
    </p>
    <p>
        An example of three quantities where it becomes clear that in general \[(A\cap B)\cup C \ne (A\cap C)\cup B\] the quantities \(A={@MusterA@}\), \( B={@MusterB@}\) and \(C={@MusterC@}\)for:
        <br/>
    </p>
    <ul>
        <li>
            \((A\cap B)\cup C = {@intersection(MusterA,MusterB)@}\cup {@MusterC@} = {@union( intersection(MusterA, MusterB), MusterC )@}\)
        </li>
        <li>
            \((A\cap C)\cup B = \emptyset\cup {@MusterB@} = {@union( intersection(MusterA, MusterC), MusterB )@}\)
        </li>
    </ul>
    <p>
        And these two quantities are different.
        <br/>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2021120900</text>
        </stackversion>
        <questionvariables>
            <text>/*Erstellt von Jonas Lache, Ruhr-Universit&#228;t Bochum*/
MusterA: set(1,2);
MusterB: set(1,3);
MusterC: set(4);</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    [[feedback:prt1]]
</span>
<span class="multilang" lang="en">
    [[feedback:prt1]]
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text/>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>MusterA</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>MusterB</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>MusterC</tans>
            <boxsize>25</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>LHS: union( intersection(ans1, ans2), ans3 );
RHS: union( intersection(ans1, ans3), ans2 );</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>LHS</sans>
                <tans>RHS</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        F&#252;r Ihre Eingabe gilt
        <span style="font-size: 0.9375rem;">
            \( (A\cap B)\cup C={@LHS@} \) und
        </span>
        <span style="font-size: 0.9375rem;">
            \( (A\cap C)\cup B={@RHS@}\)
        </span>
        <span style="font-size: 0.9375rem;">
            und somit \((A\cap B)\cup C = (A\cap C)\cup B\), wof&#252;r Sie eigentlich ein Gegenbeispiel finden sollten.
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your entry is valid
        <span style="font-size: 0.9375rem;">
            \( (A\cap B)\cup C={@LHS@} \) and
        </span>
        <span style="font-size: 0.9375rem;">
            \( (A\cap C)\cup B={@RHS@}\)
        </span>
        <span style="font-size: 0.9375rem;">
            and thus \((A\cap B)\cup C = (A\cap C)\cup B\), for which you should actually find an example.
        </span>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>1.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        F&#252;r Ihre Eingabe gilt
        <span style="font-size: 0.9375rem;">
            \((A\cap B)\cup C={@LHS@}\) und
        </span>
        <span style="font-size: 0.9375rem;">
            \((A\cap C)\cup B={@RHS@}\)
        </span>
        <span style="font-size: 0.9375rem;">
            und somit \((A\cap B)\cup C \ne (A\cap C)\cup B\), sehr gut!
        </span>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your entry is valid
        <span style="font-size: 0.9375rem;">
            \((A\cap B)\cup C={@LHS@}\) and
        </span>
        <span style="font-size: 0.9375rem;">
            \((A\cap C)\cup B={@RHS@}\)
        </span>
        <span style="font-size: 0.9375rem;">
            and thus \((A\cap B)\cup C \ne (A\cap C)\cup B\)Very good!
        </span>
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <tags>
            <tag>
                <text>R&#252;ckmeldebaum</text>
            </tag>
            <tag>
                <text>L&#246;sungsidee</text>
            </tag>
        </tags>
    </question>
    

</quiz>