<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Polynomdivision (Idee AP)-2 (Kopie) [{"DOMAINUID":"88CAE9F"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    F&#252;hren Sie die nachstehende Polynomdivision durch:
</p>
<p>
    \( p(x) = ({@p@}\)) : (\( x-{@x0@}\) )= [[/lang]][[lang code="en"]]
    <p>
    </p>
    <p>
        Perform the following polynomial division:
    </p>
    <p>
        \( p(x) = ({@p@}\): (\( x-{@x0@}\) )= [[/lang]][[input:ans1]][[lang code="de"]] \(\cdot x^2 +\) [[/lang]][[lang code="en"]] \(\cdot x^2 +\) [[/lang]][[input:ans2]][[lang code="de"]]\( \cdot &#160;x +\) [[/lang]][[lang code="en"]]\( \cdot &#160;x +\) [[/lang]][[input:ans3]][[lang code="de"]]
    </p>
    <div>
        <br/>
    </div>
    <div>
        Entscheiden Sie, ob die folgende Aussage wahr oder falsch ist!
    </div>
    <div>
        "Das obige Polynom \(p(x)\) l&#228;sst sich &#252;ber den reellen Zahlen vollst&#228;ndig in Linearfaktoren zerlegen."
    </div>
    <p>
        [[/lang]][[lang code="en"]]
    </p>
    <div>
        <br/>
    </div>
    <div>
        Decide whether the following statement is true or wrong!
    </div>
    <div>
        "The above polynomial \(p(x)\) can be completely divided into linear factors via real numbers."
    </div>
    <p>
        [[/lang]][[input:ans4]][[lang code="de"]]
    </p>
    <p>
        <br/>
    </p>
    <div>
        <span style="line-height: 21px;">
            <span style="line-height: 20px;">
                [[/lang]][[lang code="en"]]
            </span>
        </span>
    </div>
</p>
<p>
    <br/>
</p>
<div>
    <span style="line-height: 21px;">
        <span style="line-height: 20px;">
            [[/lang]][[validation:ans1]][[lang code="de"]]
        </span>
        <br/>
    </span>
</div>
<div>
    <span style="line-height: 21px;">
        <span style="line-height: 20px;">
            [[/lang]][[lang code="en"]]
        </span>
        <br/>
    </span>
</div>
<div>
    <span style="line-height: 21px;">
        <span style="line-height: 20px;">
            [[/lang]][[validation:ans2]][[lang code="de"]]
        </span>
        <br/>
    </span>
</div>
<div>
    <span style="line-height: 21px; display: inline; float: none;">
        <span style="line-height: 20px; display: inline; float: none;">
            <span style="display: inline; float: none;">
                [[/lang]][[lang code="en"]]
            </span>
            <br/>
        </span>
    </span>
</div>
<div>
    <span style="line-height: 21px; display: inline; float: none;">
        <span style="line-height: 20px; display: inline; float: none;">
            <span style="display: inline; float: none;">
                [[/lang]][[validation:ans3]][[lang code="de"]]
            </span>
            <br/>
        </span>
    </span>
</div>
<div>
    <span style="line-height: 21px;">
        <span style="line-height: 20px;">
            [[/lang]][[lang code="en"]]
        </span>
        <br/>
    </span>
</div>
<div>
    <span style="line-height: 21px;">
        <span style="line-height: 20px;">
            [[/lang]][[validation:ans4]][[lang code="de"]]
        </span>
    </span>
</div>
[[/lang]][[lang code="en"]]
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text/>
        </stackversion>
        <questionvariables>
            <text><![CDATA[a : rand(4)+1;
b : rand(2)*8-5;
c : rand(4)+1;
v1 : rand(2)*2-1;
v2 : rand(2)*2-1;
x0 : c;
q : expand((x+v1*a)^2+b);
p : expand(q*(x-x0));
a3 : coeff(p,x,3);
a2 : coeff(p,x,2);
a1 : coeff(p,x,1);
a0 : coeff(p,x,0);
b3 : coeff(q,x,2);
b2 : coeff(q,x,1);
b1 : coeff(q,x,0);
ta4 : is (b<0)]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>a2</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Richtige Antwort, gut gemacht![[/lang]][[lang code="en"]]Right answer, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Ihre Antwort ist teilweise korrekt.[[/lang]][[lang code="en"]]Your answer is partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]Falsche Antwort.[[/lang]][[lang code="en"]]Wrong answer.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>b3</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>b2</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>b1</tans>
            <boxsize>3</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>boolean</type>
            <tans>ta9</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>b3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Der f&#252;hrende Koeffizient des Polynoms ist falsch.
</p>
[[/lang]][[lang code="en"]]
<p>
    The leading coefficient of the polynomial is incorrect.
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>b2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.3000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>b1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.3000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>ta4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>