<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Potenz und Br&#252;che</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Vereinfachen Sie soweit wie m&#246;glich.
</p>
<p>
    \(\left( \left({@v@}-\dfrac{1}{{@v@}} \right)^{2}+2\right) : {\dfrac{1}{{@v^2@}}}\,=\,\)&#160; [[/lang]][[lang code="en"]]
    <p>
        Simplify as far as possible.
    </p>
    <p>
        \(\left( \left({@v@}-\dfrac{1}{{@v@}} \right)^{2}+2\right) : {\dfrac{1}{{@v^2@}}}\,=\,\) [[/lang]][[input:ans]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans]][[lang code="de"]]
        <span style="font-size: 0.9375rem;">
            [[/lang]][[lang code="en"]]
            <span style="font-size: 0.9375rem;">
                [[/lang]][[feedback:prt1]][[lang code="de"]]
            </span>
        </span>
    </p>
    [[/lang]][[lang code="en"]]
</p>
[[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    <strong>
        L&#246;sungsweg:
    </strong>
</p>
<p>
    <strong>
        <br/>
    </strong>
</p>
<p>
    \[ \left(\left({@v@}-\dfrac{1}{{@v@}}\right)^{2}+2\right) : {\dfrac{1}{{@v^2@}}} = {\left({{@v@}}^{2}-2+\dfrac{1}{{{@v@}}^{2}}+2\right)} : {\dfrac{1}{{@v^2@}}} = {\left({{@v@}}^{2}+\dfrac{1}{{{@v@}}^{2}}\right)}\cdot{{@v^2@}} = \displaystyle {@ta@}. \]
</p>
[[/lang]][[lang code="en"]]
<p>
    <strong>
        Solution:
    </strong>
</p>
<p>
    <strong>
        <br/>
    </strong>
</p>
<p>
    \[ \left(\left({@v@}-\dfrac{1}{{@v@}}\right)^{2}+2\right) : {\dfrac{1}{{@v^2@}}} = {\left({{@v@}}^{2}-2+\dfrac{1}{{{@v@}}^{2}}+2\right)} : {\dfrac{1}{{@v^2@}}} = {\left({{@v@}}^{2}+\dfrac{1}{{{@v@}}^{2}}\right)}\cdot{{@v^2@}} = \displaystyle {@ta@}. \]
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>v:rand([x,y,z,t]);
ta:ratsimp(((v-1/v)^2+2)/(1/v^2));</text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\[\left( \left({@v@}-\frac{1}{{@v@}} \right)^{2}+2\right)\div{\frac{1}{{@v^2@}}} = {@ta@} .\]</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is correct.
</p>
[[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist teilweise richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is partly correct.
</p>
[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Die Antwort ist nicht richtig.
</p>
[[/lang]][[lang code="en"]]
<p>
    The answer is not correct.
</p>
[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans</name>
            <type>algebraic</type>
            <tans>ta</tans>
            <boxsize>30</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>1</requirelowestterms>
            <checkanswertype>1</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>EqualComAss</answertest>
                <sans>ans</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>1-0-T </trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>1-0-F </falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans</sans>
                <tans>ta</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>1-1-T </trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Sie k&#246;nnen die Antwort noch weiter k&#252;rzen oder zusammenfassen.
</p>
[[/lang]][[lang code="en"]]
<p>
    You can further shorten or summarize the answer.
</p>
[[/lang]]
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>1-1-F </falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
[[lang code="de"]]
<p>
    Wenn Sie eine Kombination von mehreren Rechenoperationen haben, sollten Sie die Operationen nacheinander durchf&#252;hren.
</p>
[[/lang]][[lang code="en"]]
<p>
    If you have a combination of multiple computing operations, you should perform the operations one after the other.
</p>
[[/lang]]
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>144444624</deployedseed>
        <deployedseed>521266647</deployedseed>
        <deployedseed>687378592</deployedseed>
        <deployedseed>1827789409</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans</name>
                <value>ta</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>1-0-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans</name>
                <value>((v-1/v)^2+2)/(1/v^2)</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.1000000</expectedpenalty>
                <expectedanswernote>1-1-T </expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>