<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Verletzung der Injektivit&#228;tsbedingung   [{"DOMAINUID":"B613986"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Betrachten Sie die Funktion \(f:\mathbb{R}\to\mathbb{R}\) mit \(f(x)={@f(x)@}\).
        <br/>
    </p>
    <p>
        Geben Sie zwei Stellen an, an denen die Injektivit&#228;tsbedingung f&#252;r \(f\) verletzt ist.
        <br/>
    </p>
    <p>
        \(x_1 = \) [[input:ans1]] [[validation:ans1]]
    </p>
    <p>
        \(x_2 = \) [[input:ans2]] [[validation:ans2]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        View the function \(f:\mathbb{R}\to\mathbb{R}\) with \(f(x)={@f(x)@}\).
        <br/>
    </p>
    <p>
        Specify two locations where the injector condition for \(f\) is hurt.
        <br/>
    </p>
    <p>
        \(x_1 = \) [[input:ans1]] [[validation:ans1]]
    </p>
    <p>
        \(x_2 = \) [[input:ans2]] [[validation:ans2]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
    </p>
    <p>
        <b>
            Hinweis:
        </b>
    </p>
    <p>
        Eine Funktion \(f\colon A\to B, \; \; x\mapsto f(x)\) hei&#223;t injektiv, falls f&#252;r alle \(a_1,a_2\in A\) mit \(f(a_1)=f(a_2)\) gilt, dass \(a_1=a_2\) ist.
    </p>
    <p>
        In dieser Aufgabe sollen zwei Stellen angegeben werden, die diese Bedingung verletzen, also zwei Stellen \(a_1,a_2\in A\) mit \(f(a_1)=f(a_2)\), aber \(a_1\ne a_2\).
    </p>
    <p>
        Ein Beispiel daf&#252;r ist \(a_1={@sol1@}\) und \(a_2={@sol2@}\), denn es gilt \(f(a_1)={@f(sol1)@}=f(a_2)\), aber \({@sol1@}\ne {@sol2@}\).
        <br/>
    </p>
    <br/>
    <p>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
    </p>
    <p>
        <b>
            Note:
        </b>
    </p>
    <p>
        A function \(f\colon A\to B, \; \; x\mapsto f(x)\) is injective if for all \(a_1,a_2\in A\) with \(f(a_1)=f(a_2)\) that: \(a_1=a_2\) is.
    </p>
    <p>
        The purpose of this task is to specify two points which violate this condition, i.e. two points \(a_1,a_2\in A\) with \(f(a_1)=f(a_2)\)but \(a_1\ne a_2\).
    </p>
    <p>
        An example is \(a_1={@sol1@}\) and \(a_2={@sol2@}\), for it applies \(f(a_1)={@f(sol1)@}=f(a_2)\)but \({@sol1@}\ne {@sol2@}\).
        <br/>
    </p>
    <br/>
    <p>
    </p>
</span>
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text>n1:2+rand(3);
n2:rand([-5,-4,-3,-2,-1,1,2,3,4,5]);
f(x):=x^2+n1*x+n2;

sol1:1;
sol2:-sol1-n1;</text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    [[feedback:prt1]]
</span>
<span class="multilang" lang="en">
    [[feedback:prt1]]
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(f(x)={@f(x)@}\)</text>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    <strong>
        Richtige Antwort, gut gemacht!
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    <strong>
        Right answer, well done!
    </strong>
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    <strong>
        Ihre Antwort ist teilweise korrekt.
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    <strong>
        Your answer is partly correct.
    </strong>
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    <strong>
        Falsche Antwort.
    </strong>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    <strong>
        Wrong answer.
    </strong>
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>sol1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>sol2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>3</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>f1:f(ans1);
f2:f(ans2);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>ans2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    Sie haben zweimal die gleiche \(x\)-Stelle angegeben.
    <br/>
</span>
<span class="multilang" lang="en">
    You have twice the same \(x\)- Point.
    <br/>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>f1</sans>
                <tans>f2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Es gilt n&#228;mlich: \(f({@ans1@})={@f1@}=f({@ans2@})\).
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        It is true that: \(f({@ans1@})={@f1@}=f({@ans2@})\).
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Es gilt n&#228;mlich: \(f({@ans1@})={@f1@}\neq{@f2@}=f({@ans2@})\).
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        It is true that: \(f({@ans1@})={@f1@}\neq{@f2@}=f({@ans2@})\).
        <br/>
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>