<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text><![CDATA[Welche Aussagen sind &#228;quivalent? [{"DOMAINUID":"130EE58"}]]]></text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
    Seien \(A\), \(B\) und \(C\) logische Aussagen. Geben Sie jeweils an, ob die beiden Aussagen &#228;quivalent zu einander sind.
</p>
<table>
    <thead>
        <tr>
            <th scope="col">
                <span style="font-weight: normal;">
                    a)
                </span>
                <br/>
            </th>
            <th scope="col">
                <span style="font-weight: normal;">
                    \( A \Rightarrow B \)
                    <span>
                        und
                        <span>
                            \( B \Rightarrow A \)
                        </span>
                    </span>
                </span>
                <br/>
            </th>
            <th scope="col">
                <br/>
            </th>
            <th scope="col">
                <br/>
            </th>
            <th scope="col">
                <span style="font-weight: normal;">
                    [[/lang]][[lang code="en"]]
                    <p>
                        Let \(A\), \(B\) and \(C\) be logical statements. Specify whether the two statements are equivalent to each other.
                    </p>
                    <table>
                        <thead>
                            <tr>
                                <th scope="col">
                                    <span style="font-weight: normal;">
                                        (a)
                                    </span>
                                    <br/>
                                </th>
                                <th scope="col">
                                    <span style="font-weight: normal;">
                                        \( A \Rightarrow B \)
                                        <span>
                                            and
                                            <span>
                                                \( B \Rightarrow A \)
                                            </span>
                                        </span>
                                    </span>
                                    <br/>
                                </th>
                                <th scope="col">
                                    <br/>
                                </th>
                                <th scope="col">
                                    <br/>
                                </th>
                                <th scope="col">
                                    <span style="font-weight: normal;">
                                        [[/lang]][[input:ans1]][[lang code="de"]]
                                    </span>
                                </th>
                            </tr>
                        </thead>
                        <tbody>
                            <tr>
                                <td>
                                    b)
                                    <br/>
                                </td>
                                <td>
                                    <span>
                                        \( A \Rightarrow B \)
                                        <span>
                                            und&#160;\(&#160;\lnot B \Rightarrow \lnot A\)
                                        </span>
                                    </span>
                                </td>
                                <td>
                                    <span>
                                    </span>
                                    <br/>
                                </td>
                                <td>
                                    <br/>
                                </td>
                                <td>
                                    [[/lang]][[lang code="en"]]
                                </td>
                            </tr>
                        </tbody>
                    </table>
                </span>
            </th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>
                (b)
                <br/>
            </td>
            <td>
                <span>
                    \( A \Rightarrow B \)
                    <span>
                        and\(&#160;\lnot B \Rightarrow \lnot A\)
                    </span>
                </span>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[input:ans2]][[lang code="de"]]
            </td>
        </tr>
        <tr>
            <td>
                c)
                <br/>
            </td>
            <td>
                \( A \lor B \)
                <span>
                    und \( B \lor A \)
                </span>
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
        <tr>
            <td>
                (c)
                <br/>
            </td>
            <td>
                \( A \lor B \)
                <span>
                    and \( B \lor A \)
                </span>
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[input:ans3]][[lang code="de"]]
            </td>
        </tr>
        <tr>
            <td>
                d)
                <br/>
            </td>
            <td>
                \( A \lor B \lor C \)
                <span>
                    und
                </span>
                \( (A \lor B) \lor C \)
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
        <tr>
            <td>
                (d)
                <br/>
            </td>
            <td>
                \( A \lor B \lor C \)
                <span>
                    and
                </span>
                \( (A \lor B) \lor C \)
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[input:ans4]][[lang code="de"]]
            </td>
        </tr>
        <tr>
            <td>
                e)
                <br/>
            </td>
            <td>
                \( A \land (B \lor C) \)
                <span>
                    und \( (A \land B) \lor C \)
                    <br/>
                </span>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
        <tr>
            <td>
                (e)
                <br/>
            </td>
            <td>
                \( A \land (B \lor C) \)
                <span>
                    and \( (A \land B) \lor C \)
                    <br/>
                </span>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[input:ans5]][[lang code="de"]]
            </td>
        </tr>
        <tr>
            <td>
                f)
                <br/>
            </td>
            <td>
                \( \lnot (A \land B) \)
                <span>
                    und&#160; \( \lnot A \lor \lnot B \)
                </span>
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[lang code="en"]]
            </td>
        </tr>
        <tr>
            <td>
                (f)
                <br/>
            </td>
            <td>
                \( \lnot (A \land B) \)
                <span>
                    and \( \lnot A \lor \lnot B \)
                </span>
                <br/>
            </td>
            <td>
                <span>
                </span>
                <br/>
            </td>
            <td>
                <br/>
            </td>
            <td>
                [[/lang]][[input:ans6]][[lang code="de"]]
            </td>
        </tr>
    </tbody>
</table>
[[/lang]][[lang code="en"]]
[[/lang]][[validation:ans1]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans2]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans3]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans4]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans5]][[lang code="de"]] [[/lang]][[lang code="en"]] [[/lang]][[validation:ans6]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text><![CDATA[
[[lang code="de"]]
<p>
</p>
<p>
    <strong>
        Hinweise:
    </strong>
</p>
<p>
    a) Die Aussagen
    <span>
        \( A \Rightarrow B \)
        <span>
            und
            <span>
                \( B \Rightarrow A \) sind
                <strong>
                    <span>
                        nicht
                    </span>
                    &#228;quivalent
                </strong>
                .
                <br/>
            </span>
        </span>
    </span>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
    <u>
        Beispiel:
    </u>
    \(A\) =
    <i>
        "Drau&#223;en regnet es"
    </i>
    und \(B\) =
    <i>
        "Die Stra&#223;e ist nass"
    </i>
    <br/>
    Dann gilt zwar \(A\Rightarrow B\), aber
    <u>
        nicht
    </u>
    \(B\Rightarrow A\), denn die Stra&#223;e k&#246;nnte ja auch aus anderen Gr&#252;nden nass sein.
    <br/>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
</p>
<p>
    <span>
        <span>
            <span>
                b) Die Aussagen
                <span>
                    \( A \Rightarrow B \)
                    <span>
                        und&#160;\(&#160;\lnot B \Rightarrow \lnot A\) sind
                        <strong>
                            &#228;quivalent
                        </strong>
                    </span>
                </span>
            </span>
        </span>
    </span>
    . Diese Regel nennt man Kontrapositionsgesetz.
    <span>
        <span>
            <span>
                <span>
                </span>
            </span>
        </span>
    </span>
    <br/>
</p>
<ol>
</ol>
<ol>
</ol>
c) In der Aussagenlogik ist die
<span>
    Disjunktion ("\(\lor\)") kommutativ, daher sind&#160;\( A \lor B \)
    <span>
        und \( B \lor A \)
    </span>
    <span>
        <span>
            <strong>
                &#228;quivalent
            </strong>
            .
        </span>
    </span>
</span>
<p>
</p>
<p>
    d) Die Aussagen \( A \lor B \lor C \)
    <span>
        und
    </span>
    \( (A \lor B) \lor C \) sind
    <strong>
        &#228;quivalent
    </strong>
    . Dies sieht man, wenn man eine Wahrheitstafel anfertigt.
</p>
<p>
    e) Die Aussagen \( A \land (B \lor C) \)
    <span>
        und \( (A \land B) \lor C \) sind
        <strong>
            <span>
                nicht
            </span>
            &#228;quivalent
        </strong>
        . Dies sieht man, wenn man eine Wahrheitstafel anfertigt.
    </span>
    <br/>
</p>
<p>
    f) Die Aussagen \( \lnot (A \land B) \)
    <span>
        und&#160; \( \lnot A \lor \lnot B \) sind
        <strong>
            &#228;quivalent
        </strong>
        . Diese Regel ist eines der De Morganschen Gesetze der Aussagenlogik.
    </span>
    <br/>
</p>
[[/lang]][[lang code="en"]]
<p>
</p>
<p>
    <strong>
        Notes:
    </strong>
</p>
<p>
    (a) The statements
    <span>
        \( A \Rightarrow B \)
        <span>
            and
            <span>
                \( B \Rightarrow A \) are
                <strong>
                    <span>
                        not
                    </span>
                    equivalent
                </strong>
                .
                <br/>
            </span>
        </span>
    </span>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
    <u>
        Example:
    </u>
    \(A\) =
    <i>
        "It's raining outside"
    </i>
    and \(B\) =
    <i>
        "The road is wet"
    </i>
    <br/>
    This is true \(A\Rightarrow B\)but
    <u>
        not
    </u>
    \(B\Rightarrow A\), because the road could also be wet for other reasons.
    <br/>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
    <span>
        <span>
            <span>
            </span>
        </span>
    </span>
</p>
<p>
    <span>
        <span>
            <span>
                (b) The statements
                <span>
                    \( A \Rightarrow B \)
                    <span>
                        and\(&#160;\lnot B \Rightarrow \lnot A\) are
                        <strong>
                            equivalent
                        </strong>
                    </span>
                </span>
            </span>
        </span>
    </span>
    . This rule is called contraposition law.
    <span>
        <span>
            <span>
                <span>
                </span>
            </span>
        </span>
    </span>
    <br/>
</p>
<ol>
</ol>
<ol>
</ol>
(c) In the statement logic,
<span>
    Disjunction ("\(\lor\)") commutative, so are\( A \lor B \)
    <span>
        and \( B \lor A \)
    </span>
    <span>
        <span>
            <strong>
                equivalent
            </strong>
            .
        </span>
    </span>
</span>
<p>
</p>
<p>
    (d) The statements \( A \lor B \lor C \)
    <span>
        and
    </span>
    \( (A \lor B) \lor C \) are
    <strong>
        equivalent
    </strong>
    . You can see this when you make a truth table.
</p>
<p>
    (e) The statements \( A \land (B \lor C) \)
    <span>
        and \( (A \land B) \lor C \) are
        <strong>
            <span>
                not
            </span>
            equivalent
        </strong>
        . You can see this when you make a truth table.
    </span>
    <br/>
</p>
<p>
    (f) The statements \( \lnot (A \land B) \)
    <span>
        and \( \lnot A \lor \lnot B \) are
        <strong>
            equivalent
        </strong>
        . This rule is one of the De Morgan laws of the statement logic.
    </span>
    <br/>
</p>
[[/lang]]
]]></text>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023072101</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*Erstellt von Jonas Lache, Ruhr-Universit&#228;t Bochum*/
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
L1: [[false,true,tr("nicht &#228;quivalent","not equivalent")],[true,false,tr("&#228;quivalent","equivalent")]];
L2: [[false,false,tr("nicht &#228;quivalent","not equivalent")],[true,true,tr("&#228;quivalent","equivalent")]];
L3: [[false,false,tr("nicht &#228;quivalent","not equivalent")],[true,true,tr("&#228;quivalent","equivalent")]];
L4: [[false,false,tr("nicht &#228;quivalent","not equivalent")],[true,true,tr("&#228;quivalent","equivalent")]];
L5: [[false,true,tr("nicht &#228;quivalent","not equivalent")],[true,false,tr("&#228;quivalent","equivalent")]];
L6: [[false,false,tr("nicht &#228;quivalent","not equivalent")],[true,true,tr("&#228;quivalent","equivalent")]];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
[[lang code="de"]][[/lang]][[lang code="en"]][[/lang]][[feedback:prt1]][[lang code="de"]][[/lang]][[lang code="en"]][[/lang]]
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text/>
        </questionnote>
        <questiondescription format="html">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
Alle Antworten sind richtig, gut gemacht![[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:green;">
    <i class="fa fa-check">
    </i>
</span>
All answers are right, well done![[/lang]]
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Ihre Antworten sind teilweise richtig.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:orange;">
    <i class="fa fa-adjust">
    </i>
</span>
Your answers are partly correct.[[/lang]]
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
[[lang code="de"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
Alle Antworten sind falsch.[[/lang]][[lang code="en"]]
<span style="font-size: 1.5em; color:red;">
    <i class="fa fa-times">
    </i>
</span>
All answers are wrong.[[/lang]]
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>radio</type>
            <tans>L1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>radio</type>
            <tans>L2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>radio</type>
            <tans>L3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>radio</type>
            <tans>L4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>radio</type>
            <tans>L5</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans6</name>
            <type>radio</type>
            <tans>L6</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans5</sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>5</truenextnode>
                <trueanswernote>prt1-5-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prt1-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans6</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.167</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-6-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>